```python
# Initialize Otter
import otter
grader = otter.Notebook("lab01_.ipynb")
```

<a href="https://colab.research.google.com/github/develop_otter/material_20242/blob/main/lab/lab01/lab01_.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

<p><img alt="Colaborat" style="height: 220px;" src="Escudo-UdeA.svg.png" align="left" hspace="10px" vspace="0px"></p>

# Fundamentos en ciencias de datos

¡Bienvenido a al curso *Fundamentos de la ciencia de datos*! Cada semana completarás dos tareas de laboratorio como esta. No se pueden aprender temas técnicos sin  práctica, por lo que los laboratorios son una parte importante del curso.

Antes de comenzar, hay algunos detalles administrativos.

Las sesiones de laboratorio semanales están diseñadas para desarrollar tus habilidades con conceptos computacionales e inferenciales. Estas tareas de laboratorio son una parte obligatoria del curso y se entregarán el miércoles y el viernes a las 11 p. m.

Las sesiones de laboratorio no se transmiten por web. Estas son las políticas para obtener el crédito completo (la práctica de laboratorio vale el 10% de tu calificación final):

**Asistencia regular al laboratorio**: Para recibir notas por el laboratorio, debes asistir y completar la tarea de laboratorio durante el período de laboratorio y ser revisado por el personal del curso. Puedes quedarte todo el período y hacer un progreso significativo. Los profesores  revisarán tu trabajo para otorgarte el crédito completo, incluso si no has completado toda la tarea.


**Laboratorios a tu propio ritmo**: Si optas por trabajar de forma independiente (consulte las [políticas del curso] (http://data8.org/fa23/policies), puedes enviar la tarea completa antes de las 11 p.m. de la fecha limite. Debes terminar todo el laboratorio y aprobar todas las pruebas del autocalificador (100% de las pruebas aprobadas) para recibir tu nota. No se otorgará nota parcial.


Para el Laboratorio 1, **es necesario** todo completar el laboratorio.

Colaborar en los laboratorios está más que bien: ¡se recomienda! No deberías quedarte atascado más de unos minutos, así que pide ayuda a un instructor o compañero de clase. Explicar las cosas también es beneficioso; la mejor manera de solidificar tu conocimiento sobre un tema es explicarlo. Sin embargo, no te limites a compartir respuestas.


Puede leer más sobre las [políticas del curso](http://data8.org/fa23/policies) en el [sitio web del curso](http://data8.org/fa23/).


### En el laboratorio de hoy, aprenderá a:

1. Navegar por los cuadernos de Jupyter (como este);
2. Escribir y evaluar algunas *expresiones* básicas en Python, el lenguaje  del curso;
3. Llamar a *funciones* para usar el código escrito por otras personas
4. Dividir el código Python en partes más pequeñas para comprenderlo mejor.

Esta práctica de laboratorio cubre partes del [Capítulo 3](http://www.inferentialthinking.com/chapters/03/programming-in-python.html) del libro de texto en línea. Deberías leer los ejemplos del libro, pero no ahora. En lugar de eso, ¡comencemos!

```python
import otter
import pandas as pd
import numpy as np
%matplotlib inline
import otter
import matplotlib.pylab as plt
grader = otter.Notebook()
```

# 0. Formulario de pólizas obligatorias
¡Leer el texto anterior le ahorrará tiempo!

Complete [este](https://forms.gle/knEhFm28SprQi98E8) formulario antes de continuar. Utilice su correo electrónico de la UdeA para acceder al formulario. En la parte final del formulario, habrá una palabra secreta que deberá ingresar en el cuadro a continuación. Recuerde poner la palabra secreta entre comillas al ingresarla (es decir, "hola"). ¡Las comillas indican que es del tipo String!

**Pregunta  1:** Ingrese la palabra secreta

```python
secret_word = "..."
```

```python
grader.check("q0")
```

# 1. Jupyter Notebooks
Esta página web se llama Jupyter Notebook. Un cuaderno es un lugar para escribir programas, ver sus resultados y escribir texto.

## 1.1. Celdas de texto
En un cuaderno, cada rectángulo que contiene texto o código se llama **celda**. Las **celdas** de texto (como ésta) se pueden editar haciendo doble clic en ellas. Están escritos en un formato simple llamado [Markdown](https://www.markdownguide.org/basic-syntax/) para agregar formato y encabezados de sección. No es necesario que aprendas Markdown, pero quizás quieras hacerlo.

Después de editar una celda de texto, haga clic en el botón "ejecutar celda" en la parte superior que se parece a " ▶ " o mantenga presionado `shift` + `enter` para confirmar cualquier cambio. (Intente no borrar las instrucciones de la práctica).

**Pregunta 1.1.1.** Este párrafo está en su propia celda de texto. Intente editarlo para que esta oración sea la última oración del párrafo y luego haga clic en el botón "ejecutar celda" ▶ o mantenga presionado `shift` + `enter`. Esta frase, por ejemplo, debería suprimirse. Éste también debería hacerlo.

## 1.2. Celdas de código
Otras celdas contienen código en el lenguaje Python 3. Al ejecutar una celda de código se ejecutará todo el código que contiene. Notese que los comandos propios de python estan en inglés y aparecen en color verde: 

Para ejecutar el código en una celda de código, primero haga clic en esa celda para activarla. Se resaltará con un pequeño rectángulo verde o azul. A continuación, presione ▶ o mantenga presionado `shift` + `enter`.

Intente ejecutar esta celda:

```python
print("hola, mundo")
```

y este otra, 

```python
print("Hello, World!")
```

y este 

```python
print("\N{WAVING HAND SIGN}, \N{EARTH GLOBE ASIA-AUSTRALIA}!")
```

El componente fundamental del código Python es una expresión. Las celdas pueden contener varias líneas con múltiples expresiones. Cuando ejecuta una celda, las líneas de código se ejecutan en el orden en que aparecen. Cada expresión "imprimir" imprime una línea. Ejecute la siguiente celda y observe el orden de la salida.

```python
print("Primero se imprime esta línea,")
print("y luego esta.")
```

**Pregunta 1.2.1.** Cambie la celda de arriba para que imprima:

     Primero esta línea,
     luego todo 🌏,
     y luego este.

*Pista:* Si estás atrapado en el símbolo de la Tierra durante más de unos minutos, intenta hablar con un vecino o un instructor. Esa es una buena idea para cualquier problema de laboratorio.

## 1.3. Escribir cuadernos Jupyter

Puede utilizar los cuadernos de Jupyter para sus propios proyectos o documentos. Cuando crees tu propio cuaderno, necesitarás crear tus propias celdas para texto y código.

Para agregar una celda, haga clic en el botón + en la barra de menú. Comenzará como una celda de texto. Puedes cambiarlo a una celda de código haciendo clic dentro de él para que quede resaltado, haciendo clic en el cuadro desplegable al lado del botón con triángulos dobles en la barra de menú y eligiendo "Código".

**Pregunta 1.3.1.** Agregue una celda de código debajo de esta. Escriba el código que imprima:
   
     ¡Una celda completamente nueva! ♪🌏♪

(Ese símbolo de nota musical es como el símbolo de la Tierra. Su nombre largo es `\N{OCTAVA NOTA}`.)

Ejecuta en una celda para verificar que funciona.

```python
"Escribela aqui..."
```

Este tambien es un print:

```python
print("  /\\                 /\\")
print(" / \\'._   (\\_/)   _.'/ \\")
print(" |.''._'--(o.o)--'_.''.|")
print("  \\_ / `;=/ \" \\=;` \\ _/")
print("    `\\__| \\___/ |__/`")
print("         \\(_|_)/")

```

**Pregunta 1.3.2.**: 

Intenta pintar un gato con simbolos empleando un print, para ello copia los siguientes simbolos en una celda, análogo al muercielago mostrado en la anterior:
```bash
 /\_/\  
( o.o ) 
 > ^ <
```

```python
"Escribe la solución aqui..."
```

## 1.4. Errores
Python es un lenguaje, y al igual que los lenguajes humanos naturales tiene reglas. Se diferencia del lenguaje natural en dos aspectos importantes:

1. Las reglas son *simples*. Puedes aprender la mayoría de ellos en unas pocas semanas y adquirir un dominio razonable del idioma en un semestre.
   
2. Las reglas son *rígidas*. Si domina un lenguaje natural, puede comprender a un hablante que no lo domina, pasando por alto los pequeños errores. Una computadora que ejecuta código Python no es lo suficientemente inteligente para hacer eso.

Siempre que escribas código, cometerás errores. Cuando ejecuta una celda de código que tiene errores, Python a veces generará mensajes de error para indicarle qué hizo mal.

Los errores están bien; Incluso los programadores experimentados cometen muchos errores. Cuando cometes un error, sólo tienes que encontrar el origen del problema, solucionarlo y seguir adelante.

```python
print("This line is missing something."
```

Deberías ver algo como esto (menos las anotaciones y el resultado de error):

<img src="error.jpg"/>



La última línea del resultado de error intenta indicarle qué salió mal. La *sintaxis* de un lenguaje es su estructura, y este `SyntaxError` te indica que has creado una estructura ilegal. "`EOF`" significa "fin de archivo", por lo que el mensaje dice que Python esperaba que escribieras algo más (en este caso, un paréntesis derecho) antes de terminar la celda.

Hay mucha terminología en los lenguajes de programación, pero no es necesario saberla toda para programar de forma eficaz. Si ve un mensaje crítico como este, a menudo podrá arreglárselas sin descifrarlo. (Por supuesto, si está frustrado, pida ayuda a un compañero o a un instructor). Sin embargo, el paquete de errores del curso puede ayudar a explicar los errores y excepciones de Python y sus causas comunes, así como recursos a los que puede consultar para obtener ayuda adicional. En el ejemplo anterior, vemos que una fuente de un "SyntaxError" tiene que ver con los paréntesis.


Intente corregir el código anterior para que pueda ejecutar la celda y ver el mensaje deseado en lugar de un error.


**Nota:** En programación los mensajes aparecen en inglés, de ser necesario busca las palabras que no entiendas en un traductor, mientras se acostumbra a la terminología.

## 1.5. El núcleo
El kernel es un programa que ejecuta el código dentro de los cuadernos de trabajo y genera los resultados. En la parte superior derecha de su ventana, puede ver un círculo que indica el estado de su kernel. Si el círculo está vacío (⚪), el kernel está inactivo y listo para ejecutar código. Si el círculo está lleno (⚫), el núcleo está ocupado ejecutando algún código.

Al lado de cada celda de código, verá un texto que dice "[ ]". Antes de ejecutar la celda, verá "[ ]". Cuando la celda se esté ejecutando, verá `[*]`. Si ve un asterisco (\*) junto a una celda que no desaparece, es probable que el código dentro de la celda esté tardando demasiado en ejecutarse y podría ser un buen momento para interrumpir el kernel (que se analiza a continuación). . Cuando termine de ejecutarse una celda, verá un número entre corchetes, como este: `[1]`. El número corresponde al orden en que ejecuta las celdas; entonces, la primera celda que ejecute mostrará un 1 cuando termine de ejecutarse, la segunda mostrará un 2, y así sucesivamente.

Es posible que se encuentre con problemas en los que su kernel esté bloqueado por una cantidad excesiva de tiempo, su computadora portátil sea muy lenta y no responda, o su kernel pierda su conexión. Si esto sucede, intente los siguientes pasos:
1. En la parte superior de la pantalla, haga clic en **Kernel** y luego en **Interrumpir kernel**.
2. Si eso no ayuda, haga clic en **Kernel** y luego en **Reiniciar kernel**. Si hace esto, tendrá que ejecutar las celdas de código desde el inicio de su computadora portátil hasta donde detuvo su trabajo.

## 1.6. Enviar su trabajo
Todas las tareas del curso se distribuirán en cuadernos como este, y usted enviará su trabajo desde el cuaderno. Usaremos un sistema llamado Otter Grader para verificar su trabajo y ayudarlo a enviarlo. Si aún no lo ha hecho, asegúrese de haber ejecutado la celda en la parte superior de este cuaderno para inicializar OtterGrader.

# 2. Números

La información cuantitativa surge en todas partes de la ciencia de datos. Además de representar comandos para imprimir líneas, las expresiones pueden representar números y métodos de combinación de números. La expresión `3.2500` se evalúa como el número 3.25. (Ejecute la celda y vea).

```python
3.2500
```

Observe que no tuvimos que "imprimir". Cuando ejecuta una celda del cuaderno, si la última línea tiene un valor, Jupyter imprime ese valor de manera útil. Sin embargo, no imprimirá las líneas anteriores automáticamente.

```python
print(2)
3
4
```

Arriba, deberías ver que 4 es el valor de la última expresión, se imprime 2, pero 3 se pierde para siempre porque no se imprimió ni fue la última.

De todos modos, no querrás imprimir todo todo el tiempo. Pero si lo deseas, imprime el número 3, cambia la celda de arriba para imprimirla.

## 2.1. Aritmética
La línea en la siguiente celda realiza la resta entre dos números. Su valor es el que cabría esperar. Ejecutalo.

```python
3.25 - 1.5
```

Intenta realizar el cálculo de tu edad menos los últimos dos dígitos de tu documento:

```python
"Escribe la solucion aca"
```

Muchas operaciones aritméticas básicas están integradas en Python. La sección del libro de texto sobre [Expresiones](http://www.inferentialthinking.com/chapters/03/1/Expressions.html) describe todos los operadores aritméticos utilizados en el curso. El operador común que difiere de la notación matemática típica es `**`, que eleva un número a la potencia del otro. Entonces, `2**3` significa $2^3$ y se evalúa como 8.
El orden de las operaciones es el mismo que aprendiste en el colegio y Python también tiene paréntesis.  Por ejemplo, compare los resultados de las siguientes celdas.


### 2.1.1 Suma y Multiplicación

```python
2 + 3 * 4
```

```python
(2 + 3) * 4
```

### 2.1.2 Resta y División

```python
10 - 4 / 2
```

```python
(10 - 4) / 2
```

### 2.1.3  Exponentes y Multiplicación

```python
2 ** 3 * 4
```

```python
2 ** (3 * 4)
```

### 2.1.4 Múltiples Operaciones


```python
2 + 3 * 4 - 5 / 2
```

```python
(2 + 3) * (4 - 5) / 2
```

Note que todas las operaciones anteriores, las podemos resumir en una sola celda:

```python
print(2 + 3 * 4)  # El resultado es 14 porque la multiplicación se realiza antes que la suma
print((2 + 3) * 4)  # El resultado es 20 porque los paréntesis alteran el orden de las operaciones
print("---------------------")
print(10 - 4 / 2)  # El resultado es 8.0 porque la división se realiza antes que la resta
print((10 - 4) / 2)  # El resultado es 3.0 porque los paréntesis alteran el orden de las operaciones
print("---------------------")

print(2 ** 3 * 4)  # El resultado es 32 porque la exponenciación se realiza antes que la multiplicación
print(2 ** (3 * 4))  # El resultado es 4096 porque los paréntesis alteran el orden de las operaciones
print("---------------------")

print(2 + 3 * 4 - 5 / 2)  # El resultado es 10.5 porque la multiplicación y división se realizan antes que la suma y resta
print((2 + 3) * (4 - 5) / 2)  # El resultado es -2.5 porque los paréntesis alteran el orden de las operaciones

```

# Ejemplos:
Revisa las siguientes operaciones matemáticas:

```python
6+6*5-6*3**2*2**3/4*7
```

En notación matemática estándar, el primer ejemplo seria así:

$$6 + 6 × 5 - 6 × 3^2 \times \frac{2^3}{4} \times 7,$$

Revisemos el siguiente ejemplo:

```python
6 + (6 * 5 - (6 * 3) ** 2) * (2 ** 3 / 4 * 7)
```

En notación matemática, seria de esta manera:

$$6 + (6 \times 5 - (6 \times 3))^2 \times \left(\frac{\left(2^3\right)}{4} \times 7\right).$$

**Pregunta 2:**  Escribe una expresión de Python en la siguiente celda que sea igual a $\displaystyle 5 \times \left(3 \frac{10}{11}\right) - 50 \frac{1}{3 } + 2^{0.5 \times 22} - \frac{7}{33} + 7$. Eso es cinco por tres y diez onceavos, menos cincuenta y un tercio, más dos elevado a la mitad de veintidós, menos siete treinta y tres más seis.

Reemplace las elipses (`...`) con su expresión. Intente utilizar paréntesis sólo cuando sea necesario.

**Nota:** Por "$\displaystyle 3 \frac{10}{11}$", nos referimos a $\displaystyle 3+\frac{10}{11}$, no a $\displaystyle 3 \times \frac{ 10}{11}$.

**Pista:** El resultado correcto debe comenzar con un número familiar.

```python
equation = ...
```

```python
equation
```

```python
grader.check("q2")
```

# 3. Nombres
En lenguaje natural, tenemos terminología que nos permite hacer referencia rápidamente a conceptos muy complicados. No decimos: "¡Ese es un mamífero grande con pelaje marrón y dientes afilados!" En cambio, simplemente decimos: "¡Oso!"

En Python, hacemos esto con *declaraciones de asignación*. Una declaración de asignación tiene un nombre en el lado izquierdo de un signo `=` y una expresión que se evaluará en el lado derecho.

```python
diez = 3 * 2 + 4
```

Cuando ejecuta esa celda, Python primero calcula el valor de la expresión en el lado derecho, `3 * 2 + 4`, que es el número 10. Luego asigna ese valor al nombre `diez`. En ese momento, el código en la celda termina de ejecutarse.

Después de ejecutar esa celda, el valor 10 está vinculado al nombre "diez":

```python
diez
```

```python
print(diez)
```

La afirmación "diez = 3 * 2 + 4" no afirma que "diez" ya sea igual a "3 * 2 + 4", como podríamos esperar por analogía con la notación matemática. Más bien, esa línea de código cambia lo que significa "diez"; ahora se refiere al valor 10, mientras que antes no significaba nada en absoluto.

Si los diseñadores de Python hubieran sido despiadadamente pedantes, nos habrían hecho escribir

     define el nombre diez para que en adelante tenga el valor de 3 * 2 + 4

en cambio. ¡Probablemente apreciará la brevedad de "`=`"! Pero tenga en cuenta que este es el verdadero significado.

**Pregunta 3.1.** Ejecute la siguiente celda que usa un nombre de variable `eleven` que no ha sido asignado a nada. ¡Verás un error!

```python
eleven + 8
```

Un patrón común en los cuadernos de Jupyter es asignar un valor a un nombre y luego evaluar inmediatamente el nombre en la última línea de la celda para que el valor se muestre como salida.

```python
cerca_pi = 355/113
cerca_pi
```

Un patrón común en programación es utilizar varias líneas en una sola celda para construir un cálculo complejo paso a paso, nombrando los resultados intermedios en cada etapa."

```python
salario_quincenal = 650000
salario_mensual = 2 * salario_quincenal
numero_de_meses_en_un_ano = 12
salario_anual = numero_de_meses_en_un_ano * salario_mensual
salario_anual
```

Los nombres en Python pueden tener **letras** (las letras mayúsculas y minúsculas están bien y cuentan como letras diferentes), guiones bajos y números. El primer carácter no puede ser un número (de lo contrario, el nombre podría parecer un número). Y los nombres no pueden contener espacios, ya que los espacios se utilizan para separar fragmentos de código entre sí.

Aparte de esas reglas, el nombre que le pongas a algo no le importa *a Python*. Por ejemplo, las dos celdas abajo, hacen  lo mismo que la celda anterior, excepto que todo tiene un nombre diferente:

```python
a = 650000
b = 2 * a
c = 12
d = c * b
d
```

Igual que estas con nombres en inglés:

```python
semimonthly_salary = 650000
monthly_salary = 2 * semimonthly_salary
number_of_months_in_a_year = 12
yearly_salary = number_of_months_in_a_year * monthly_salary
yearly_salary
```

**Sin embargo**, los nombres son muy importantes para que su código sea *legible* para usted y los demás. La celda de arriba es más corta, pero es totalmente inútil sin una explicación de lo que hace.

## 3.1. Comprobando su código
Ahora que sabes cómo nombrar cosas, puedes comenzar a usar las *pruebas* integradas para comprobar si tu trabajo es correcto. A veces, hay varias pruebas para una sola pregunta y es necesario aprobarlas todas para recibir crédito por la pregunta. No cambie el contenido de las celdas de prueba.

Continúe e intente la **pregunta 3.1.2.** Al ejecutar la celda directamente después se comprobará si ha asignado "segundos_en_una_década" correctamente en la Pregunta 3.1.2. Si no lo has hecho, esta prueba te dirá la respuesta correcta. Resista la tentación de simplemente copiarlo y, en su lugar, intente ajustar su expresión. (A veces las pruebas darán pistas sobre lo que salió mal...)

**Pregunta 3.1.2.** Asigne el nombre `segundos_en_una_decada` al número de segundos entre la medianoche del 1 de enero de 2010 y la medianoche del 1 de enero de 2020. Tenga en cuenta que hay dos años bisiestos en este lapso de una década. Un año no bisiesto tiene 365 días y un año bisiesto tiene 366 días.

**Pista:** Si estás atascado, la siguiente sección te muestra cómo obtener pistas.

```python
# Cambiar la siguiente línea
# para que calcule el número de segundos en una década
# y asigna a ese número el nombre, segundos_en_una_década.
segundos_en_una_decada = ...
```

```python
grader.check("q3_1_2")
```

# 3.2. Comentarios
Es posible que hayas notado estas líneas en la celda en la que respondiste la pregunta 3.2:

     # Cambiar la siguiente línea
     # para que calcule el número de segundos en una década
     # y asigna a ese número el nombre, segundos_en_una_década.
    
Esto se llama *comentario*. No hace que suceda nada en Python; Python ignora cualquier cosa en una línea después de `#`. En cambio, está ahí para comunicarle algo sobre el código a usted, el lector humano. Los comentarios son extremadamente útiles.

<img src="http://imgs.xkcd.com/comics/future_self.png">

Fuente: [XKCD](http://imgs.xkcd.com/comics/future_self.png)

## 3.3. Aplicación: un experimento de física

En la misión Apolo 15 a la Luna, el astronauta David Scott replicó el famoso experimento de física de Galileo en el que demostró que la gravedad acelera objetos de diferente masa a la misma velocidad. Debido a que no hay resistencia del aire para un objeto que cae sobre la superficie de la Luna, incluso dos objetos con masas y densidades muy diferentes deberían caer al mismo ritmo. David Scott comparó una pluma y un martillo.

Puede ejecutar la siguiente celda para ver un vídeo del experimento.

```python
from IPython.display import YouTubeVideo
# The original URL is:
# https://www.youtube.com/watch?v=U7db6ZeLR5s

YouTubeVideo("U7db6ZeLR5s")
```

Aquí está la transcripción del vídeo:

**167:22:06 Scott**: Bueno, en mi mano izquierda tengo una pluma; en mi mano derecha, un martillo. Y supongo que una de las razones por las que llegamos hoy aquí fue porque un caballero llamado Galileo, hace mucho tiempo, hizo un descubrimiento bastante significativo sobre la caída de objetos en campos de gravedad. Y pensamos cuál sería un mejor lugar para confirmar sus hallazgos que la Luna. Y por eso pensamos en probarlo aquí para usted. La pluma resulta ser, apropiadamente, una pluma de halcón para nuestro Halcón. Y los dejaré aquí a los dos y, con suerte, caerán al suelo al mismo tiempo.

**167:22:43 Scott**: ¿Qué tal eso?

**167:22:45 Allen**: ¡Qué tal eso! (Aplausos en Houston)

**167:22:46 Scott**: Lo que demuestra que el Sr. Galileo tenía razón en sus hallazgos..

**Ley de Newton.** Usando este video, también podemos intentar confirmar otra leyes de la física: la ley de gravitación universal de Newton. Las leyes de Newton predicen que cualquier objeto que caiga cerca de la superficie de la Luna debería caer:

$$y = \frac{1}{2} G \frac{M}{R^2} t^2 \text{ metros}$$

después de $t$ segundos, donde $G$ es una constante universal, $M$ es la masa de la luna en kilogramos y $R$ es el radio de la luna en metros. Entonces, si conocemos $G$, $M$ y $R$, las leyes de Newton nos permiten predecir qué tan lejos caerá un objeto en cualquier  tiempo.

Para verificar la precisión de esta ley, calcularemos la diferencia entre la distancia prevista que cae el martillo y la distancia real. (Si son diferentes, podría deberse a que las leyes de Newton son erróneas, a que nuestras mediciones son imprecisas o a que hay otros factores que afectan al martillo y que no hemos tenido en cuenta).

Alguien estudió el vídeo y estimó que el martillo cayó a 113 cm de la superficie. Contando fotogramas en el vídeo, el martillo cae durante 1,2 segundos (36 fotogramas).

**Pregunta 3.3.1.** Complete el código en la siguiente celda para completar los datos del experimento.

*Pista:* No se requiere ningún cálculo; simplemente complete los datos del párrafo anterior. 

```python
# t, la duración de la caída en el experimento, en segundos.
# Complete esto.

time = ...
estimated_distance_m=...
# La distancia estimada que realmente cayó el martillo, en metros.
# Complete esto.
estimated_distance_m
```

```python
grader.check("q3_3_1")
```

**Pregunta 3.3.2.** Ahora, completa el código en la siguiente celda para calcular la diferencia entre las distancias previstas y estimadas (en metros) a las que cayó el martillo en este experimento.

Esto solo significa traducir la fórmula anterior ($\frac{1}{2}G\frac{M}{R^2}t^2$) al código Python. Tendrás que reemplazar cada variable en la fórmula matemática con el nombre que le dimos a ese número en el código Python.

*Pista:* Intente utilizar variables que ya definió en la pregunta 3.3.1

```python
# Primero, hemos anotado los valores de las 3 constantes universales.
# que aparecen en la fórmula de Newton.

# G, la constante universal que mide la fuerza de la gravedad.
gravity_constant = 6.674 * 10**-11

# M, la masa de la luna, en kilogramos.
moon_mass_kg = 7.34767309 * 10**22

# R, el radio de la luna, en metros.
moon_radius_m = 1.737 * 10**6


# La distancia que debería haber caído el martillo
# durante la duración de la caída, en metros,
# según la ley de gravedad de Newton.
# El texto de arriba describe la fórmula.
# para esta distancia dada por la ley de Newton.
# **USTED COMPLETA ESTA PARTE.**


predicted_distance_m = ...

# Aquí hemos calculado la diferencia.
# entre la distancia de caída prevista y la distancia que realmente medimos.
# Si ha completado el código anterior, esto debería funcionar.

difference = predicted_distance_m - estimated_distance_m
difference
```

```python
grader.check("q3_3_2")
```

## 4. Llamada a funciones

La forma más común de combinar o manipular valores en Python es llamando funciones. Python viene con muchas funciones integradas que realizan operaciones comunes.

Por ejemplo, la función `abs` toma un solo número como argumento y devuelve el valor absoluto de ese número. Ejecute las siguientes dos celdas y vea si comprende el resultado.

```python
abs(5)
```

## 4.1. Aplicación: Computación de distancias recorridas


María está en la esquina de la Calle 52 con la Carrera 49 (Edificio Coltejer) en Medellín, y quiere saber hasta dónde tendría que caminar para llegar a la Placita de Flores en la esquina de la Carrera 39 con la Calle 50.

No puede cortar bloques en diagonal, ya que hay edificios en el camino. Tiene que caminar por las aceras. Usando un mapa, ve que tendría que caminar 10 carreras (cuadras largas) y 2 calles (cuadras cortas). En términos de los números dados, el número  10 esta cálculado como la diferencia entre 49 y 39, en valor absoluto, y 2 de manera similar.

María también sabe que las cuadras en Medellín miden aproximadamente 80 m (calles) por 80 m (carreras). Entonces, en total, tendría que caminar $(80 \times |52 - 50| + 80 \times |49 - 39|)$ metros para llegar a la Placita de Flores.


<img src="https://github.com/fundatosudea/material_20242/blob/develop_otter/Imag/fig/fig_02caminata.png?raw=truee"/>

**Pregunta 4.1.1.** Complete la línea `num_carreras = ...` en la siguiente celda para que calcule la distancia que debe caminar Maria, dale el nombre de `distancia_medellin`. Todo lo demás ya lo hemos completado. **Utilice la función `abs`.** Además, asegúrese de ejecutar la celda de prueba después para probar su código.

```python

carrera_inicial = 49
calle_inicial = 52
carrera_final = 39
calle_final = 50

distancia_calle = 80  #80m  aproximadamente 
distancia_carrera = 80  #80m aproximadamente 

# Cálculo de la distancia en cuadras
num_calles = abs(calle_final - calle_inicial)
num_carreras = ...

# Cálculo de la distancia total en metros
distancia_medellin = ....
```

```python
grader.check("q4_1_1")
```

##### Múltiples argumentos
Algunas funciones toman múltiples argumentos, separados por comas. Por ejemplo, la función incorporada `max` devuelve el argumento máximo que se le pasa.

```python
max(2, -3, 4, -5)
```

# 5. Comprensión de las expresiones anidadas

Las llamadas a funciones y las expresiones aritméticas pueden contener expresiones, como viste en la pregunta anterior:

    num_calles = abs(calle_final - calle_inicial)

que indica que ```num_calles = abs(39 - 49)```. Las expresiones anidadas pueden convertirse en código de apariencia  complicada. Sin embargo, la forma en que se descomponen las expresiones complicadas es muy regular.

Supongamos que estamos interesados en longitudes de gatos que son muy inusuales(longitud desde la cabeza hasta la cola). Diremos que una longitud es inusual en la medida en que está muy alejada en la recta numérica de la longitud promedio de un gato. Una estimación de la longitud promedio de un gato (que esperamos que sea un promedio de todos los gatos de la Tierra hoy en día) es de **46.2cm**



Entonces, si los gatos Savannah miden en promedio 55.7cm  de largo, entonces su longitud está a $|46.2 - 55.7|$, o $9.5$cm  de distancia del promedio. Aquí hay una foto de eso:

<img src="https://github.com/fundatosudea/material_20242/blob/develop_otter/Imag/fig/fig_03recta.png?raw=true"/>


La fuente de la longitud promedio de los gatos es [Wikipedia](https://en.wikipedia.org/wiki/Cat#:~:text=The%20domestic%20cat%20has%20a,(9%20and%2011%20lb). ).

Y así es como escribiríamos esa expresión en una línea de código Python:

```python
abs(46.7-55.7)
```

¿Que está pasando aqui? `abs` toma solo un argumento, por lo que todo lo que está dentro del paréntesis es parte de ese *argumento único*. Específicamente, el argumento es el valor de la expresión `55.7 - 46.2`. El valor de esa expresión es "9.0". Ese valor es el argumento de "abs". El valor absoluto de eso es "9.0", por lo que "9.0" es el valor de la expresión completa `abs(55.7-46.2 )`.

Imagen simplificando la expresión en varios pasos:

1. `abs(55.7-46.2 )`
2. `abs(9.0)`
3. `9.0`

Lo anterior, es lo que hace Python para calcular el valor de la expresión.

**Pregunta 5.1.** Supongamos  que la longitud de una raza particular de  gatos es de 42.8cm. En la siguiente celda, use "abs" para calcular el valor absoluto de la diferencia entre la longitud promedio del gato y esta raza particular. Asigne a ese valor el nombre `distancia_del_promedio_cm`.

```python
# Aquí está el número de calles de distancia:
distancia_del_promedio_cm = abs(42.8-46.2)
distancia_del_promedio_cm = ....
```

```python
grader.check("q51")
```

**Pregunta 5.1.1.** Dadas las longitudes de los gatos de Yanay, Hummus, Gatkes y Zeepty, escribe una expresión que calcule la diferencia más pequeña entre cualquiera de las tres longitudes. Su expresión no debe tener números, solo llamadas a funciones y los nombres "hummus", "gatkes" y "zeepty". Asigne al valor de su expresión el nombre `dif_de_longitud_min`.*

```python
hummus =  24.5 # Hummus is 24.5 inches long
gatkes = 19.7 # Gatkes is 19.7 inches long
zeepty = 15.8 # Zeepty is 15.8 inches long


dif_de_longitud_min = ...
```

```python
dif_de_longitud_min
```

```python
grader.check("q5_1_1")
```

¡Ya terminaste con el Laboratorio 1!

**Información importante sobre el envío:** Asegúrese de ejecutar las pruebas y verificar que todas pasen, luego elija **Guardar Notebook** en el menú **Archivo**, luego **ejecute la celda final** y haga clic en enlace para descargar el archivo zip. Luego, vaya a [Gradescope](https://www.gradescope.com/courses/576157) y envíe el archivo zip a la tarea correspondiente. El nombre de esta tarea es "Lab 01 Autograder".

**Una vez que haya enviado, su tarea de Gradescope debería verse como la siguiente imagen si pasó todas las pruebas.** <span style="color: red">**Asegúrese de ver "Todos los casos de prueba aprobados" para asegurarse tu tarea se envió correctamente.**</span> Si tu tarea no se ve así, es probable que no hayas cargado tu cuaderno correctamente.

<img src="https://github.com/darioblandon/materials-fa23/blob/main/lab/lab01/gradescope.png?raw=1">

¡Dash quería felicitarte y darte la bienvenida al curso de  Fundamentos en ciencias de datos!

<img src="dash.JPG" alt="dibujo" width="200"/>

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

These are some submission instructions.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q0 (points: None)
```
>>> assert isinstance(secret_word, str)
>>> assert np.isclose(len(secret_word), 10)
```


Tests for q2 (points: None)
```
>>> answer = equation
>>> assert isinstance(answer, (int, float))
>>> assert np.isclose(answer, equation)
```


Tests for q3_1_2 (points: None)
```
>>> answer = segundos_en_una_decada
>>> assert isinstance(answer, int)
>>> assert np.isclose(answer, segundos_en_una_decada)
```


Tests for q3_3_1 (points: None)
```
>>> assert time != '....'
```
```
>>> assert isinstance(estimated_distance_m, float) == True
```
```
>>> import numpy as np
>>> assert np.isclose(estimated_distance_m, 1.13, rtol=100.0) == True
```


Tests for q3_3_2 (points: None)
```
>>> assert predicted_distance_m != '....'
```
```
>>> assert isinstance(difference, float) == True
```
```
>>> import numpy as np
>>> assert np.isclose(difference, 0.040223694659304865, rtol=0.01) == True
```


Tests for q4_1_1 (points: None)
```
>>> assert isinstance(num_calles, int) == True
>>> assert isinstance(distancia_medellin, int) == True
```
```
>>> assert np.isclose(num_carreras, 10, 0.01)
```
```
>>> assert np.isclose(distancia_medellin, 960.0, 0.01)
```


Tests for q51 (points: None)
```
>>> assert isinstance(distancia_del_promedio_cm, float) == True
>>> assert np.isclose(distancia_del_promedio_cm, abs(-3.4000000000000057), 0.002)
```


Tests for q5_1_1 (points: None)
```
>>> assert isinstance(dif_de_longitud_min, float) == True
```
```
>>> assert np.isclose(dif_de_longitud_min, 3.89, rtol=0.01) == True
```
