```python
# Initialize Otter
import otter
grader = otter.Notebook("lab10.ipynb")
```


<p><img alt="Colaborat" style="height: 220px;" src="https://github.com/CienciaDatosUdea/001_DataSetFundCienciaDatos/blob/main/semestre2024-2/lab/lab02/Escudo-UdeA.svg.png?raw=true" align="left" hspace="10px" vspace="0px"></p>


<h1> Fundamentos en ciencias de datos</h1>




# Laboratorio 10: Regresión

¡Bienvenido al Laboratorio 10!

Hoy practicaremos con  con la regresión lineal. Puedes encontrar más información sobre este tema en
[Capítulo 15.2](https://www.inferentialthinking.com/chapters/15/2/Regression_Line.html#the-regression-line).

Para este laboratorio, trabajaremos con un archivo llamado `faithful.csv`, que contiene las columnas `duration` y `wait`. Utilizaremos estas columnas para aplicar técnicas de regresión lineal y analizar los datos. ¡Empecemos!


```python
# Run this cell, but please don't change it.

# These lines import the Numpy and Datascience modules.
import numpy as np
from datascience import *

# These lines do some fancy plotting magic.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plots
plots.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore', FutureWarning)
```

<!-- BEGIN QUESTION -->

Ingresa tu nombre  con apellido y número de carnet en las variables dadas en la celda siguiente:
Ejemplo:
```python
NombreApellidos="Ada Lovelace"
NumeroCarnet="27182818"
Email="Ada.Lovelace@udea.edu.co"
```

```python
NombreApellidos="..."
NumeroCarnet=".."
Email="..."
```

```python
grader.check("qt")
```

<!-- END QUESTION -->

# ¿Qué tan predecible es Old Faithful?

Old Faithful es un [géiser](https://es.wikipedia.org/wiki/G%C3%A9iser) en el Parque Nacional Yellowstone, famoso por sus erupciones relativamente regulares.  
¡Ejecuta la celda a continuación para ver a Old Faithful en acción!

```python
# For the curious: this is how to display a YouTube video in a
# Jupyter notebook.  The argument to YouTubeVideo is the part
# of the URL (called a "query parameter") that identifies the
# video.  For example, the full URL for this video is:
#   https://www.youtube.com/watch?v=wE8NDuzt8eg
from IPython.display import YouTubeVideo
YouTubeVideo("wE8NDuzt8eg")
```

Algunas de las erupciones del géiser Old Faithful duran más que otras. Cuando hay una erupción prolongada, generalmente es seguida por un tiempo de espera aún más largo antes de la siguiente erupción. 

Hoy utilizaremos un conjunto de datos sobre la duración de las erupciones y los tiempos de espera para ver si podemos hacer tales predicciones con precisión mediante regresión lineal.

El conjunto de datos cuenta con una fila para cada erupción observada e incluye las siguientes columnas:
- `duration`: Duración de la erupción, en minutos
- `wait`: Tiempo entre esta erupción y la siguiente, también en minutos

Ejecuta la siguiente celda para cargar el conjunto de datos.

```python
faithful = Table.read_table("faithful.csv")
faithful
```

**Pregunta 1.0.** Las siguientes afirmaciones son los pasos desordenados de una regresión lineal.

1. Calcular los parámetros de la línea de regresión: la pendiente y la intersección.
2. Evaluar la línea de regresión calculando el RMSE de la línea y analizando el gráfico de residuos.
3. Usar la línea de regresión para generar predicciones para cada valor de x.
4. Determinar si la regresión lineal es un método razonable visualizando tus datos y calculando el coeficiente de correlación.

Crea un arreglo llamado `least_squares_order` que contenga el orden correcto de un análisis de regresión lineal, donde el primer elemento del arreglo sea el primer paso de un análisis de regresión lineal y el último elemento del arreglo sea el último paso de un análisis de regresión lineal.



```python
least_squares_order = make_array(....)
```

```python
grader.check("q0")
```

Nos gustaría utilizar la regresión lineal para realizar predicciones. Sin embargo, este método no será eficaz si los datos no presentan una relación aproximadamente lineal. Para verificar esto, debemos examinar los datos en detalle. Es fundamental explorar estos datos para entender cómo se relacionan entre sí antes de proceder con cualquier modelo de regresión lineal.

**Pregunta 1.1.** Realiza un diagrama de dispersión (scatter plot) de los datos. Es convencional colocar en el eje vertical la columna que deseamos predecir y en el eje horizontal la otra columna. 

```python
faithful = Table.read_table('faithful.csv')
faithful.scatter('...', '....')
plots.title('....')
plots.xlabel('...')
plots.ylabel('...')
plots.show()
```

**Pregunta 1.2.** ¿Están aproximadamente relacionadas de manera lineal la duración de la erupción y el tiempo de espera según el diagrama de dispersión anterior? ¿Es esta relación positiva?

Analiza el diagrama de dispersión proporcionado para evaluar si existe una tendencia lineal entre estas dos variables y si dicha relación puede considerarse positiva.

_Escribe tu respuesta aquí, reemplazando este texto._

Vamos a continuar con la suposición de que existe una relación lineal entre las variables, por lo que es razonable utilizar la regresión lineal para analizar estos datos.

A continuación, nos gustaría graficar los datos en unidades estándar. Si no recuerdas la definición de unidades estándar, la sección [14.2](https://www.inferentialthinking.com/chapters/14/2/Variability.html#standard-units) del libro de texto puede ser de ayuda.

**Pregunta 1.3.** Calcula la media y la desviación estándar de las duraciones de las erupciones y los tiempos de espera. **Luego**, crea una tabla llamada `faithful_standard` que contenga las duraciones de las erupciones y los tiempos de espera en unidades estándar. Las columnas deben llamarse `duration_standard_units` and `wait_standard_units`.


```python
duration_mean = ...
duration_std = ...
wait_mean = ...
wait_std = ...

faithful_standard = Table().with_columns(
    "duration_standard_units", ...,
    "wait_standard_units", ...)
faithful_standard
```

```python
grader.check("q3")
```

**Pregunta 1.4.** Grafica los datos nuevamente, pero esta vez en unidades estándar. Estamos analizando los datos del archivo `faithful.csv`, que contiene información sobre la duración (duration) y el tiempo de espera (wait) de las erupciones del géiser Old Faithful. Es importante transformar los datos a unidades estándar para tener una perspectiva unificada y facilitar el análisis comparativo.

```python
faithful_standard.scatter('....', '....')
plots.title('Pon el nombre adecuado ...')
plots.xlabel('Pon el nombre adecuado ...')
plots.ylabel('Pon el nombre adecuado ...')
plots.show()
```

Podrás notar que este gráfico luce igual que el anterior. Sin embargo, los datos y los ejes están escalados de manera diferente. Por lo tanto, es importante que prestes atención a las marcas en los ejes.

**Pregunta 1.5.** Entre los siguientes números, ¿cuál crees que se aproxima más a la correlación entre la duración de la erupción y el tiempo de espera en este conjunto de datos?

1. -1
2. 0
3. 1

Asigna a `correlation` el número que corresponde a tu estimación (ya sea 1, 2 o 3). 

```python
correlation = ...

```

```python
grader.check("q5")
```

**Pregunta 1.6.** Calcula el coeficiente de correlación: `r`. 



*Pista:* Utiliza `faithful_standard`. La sección [15.1](https://www.inferentialthinking.com/chapters/15/1/Correlation.html#calculating-r) explica cómo hacerlo.



```python
r = ...

```

```python
grader.check("q6")
```

## 2. La  regresión lineal 

Recordemos que la **correlación** es la **pendiente de la línea de regresión cuando los datos están en unidades estándar**.

La siguiente celda grafica la línea de regresión en unidades estándar:

$$\text{tiempo de espera en unidades estándar} = r \times \text{duración de la erupción en unidades estándar}$$

Con nuestras variables sería:

$$ \text{duration\_standard\_units} = r \times \text{wait\_standard\_units} $$

A continuación, se grafica nuevamente los datos en unidades estándar para su comparación.


```python
def plot_data_and_line(dataset, x, y, point_0, point_1):
    """Makes a scatter plot of the dataset, along with a line passing through two points."""
    dataset.scatter(x, y, label="data")
    xs, ys = zip(point_0, point_1)
    plots.plot(xs, ys, label="regression line")
    plots.legend(bbox_to_anchor=(1.5,.8))

plot_data_and_line(faithful_standard, 
                   "duration_standard_units", 
                   "wait_standard_units", 
                   [-2, -2*r], 
                   [2, 2*r])
```

¿Cómo convertirías un punto en unidades estándar de nuevo a sus unidades originales? Necesitaríamos "estirar" su posición horizontal multiplicando por `duration_std` y su posición vertical multiplicando por `wait_std`. Esto significa que lo mismo ocurriría con la pendiente de la línea.

Estirar una línea horizontalmente la hace menos inclinada, por lo que dividimos la pendiente entre el factor de estiramiento horizontal. Estirar una línea verticalmente la hace más inclinada, por lo que multiplicamos la pendiente por el factor de estiramiento vertical.

**Pregunta 2.1.** Calcula la pendiente de la línea de regresión en unidades originales y asígnala a la variable `slope`.

(Si la explicación sobre el "estiramiento" no es intuitiva, consulta la sección [15.2](https://www.inferentialthinking.com/chapters/15/2/Regression_Line.html#the-equation-of-the-regression-line) en el libro de texto.)


```python

slope = ...
slope
```

```python
grader.check("q7")
```

Sabemos que la línea de regresión pasa por el punto `(duration_mean, wait_mean)`. Recordemos que la ecuación de la línea de regresión en las unidades originales es:

$$ (y-\mu_{y}) = r \frac{\sigma_{y}}{\sigma_{x}} (x-\mu_{x})$$

donde: 

- y es el tiempo de espera en las unidades  **no** estandar

- x es la duracion en unidades escaladas **no** estandar 

- $\mu_y$ es la media de y (duration_mean).

- $\sigma_{y}$ es la desviacion estandar de y (duration_std).

- $\mu_x$ es la media de x (wait_mean).

- $\sigma_{x}$ es la desviacion estandar de x (wait_std).
    
es decir, que el intercepto puede ser computado como:

$$ y = r \frac{\sigma_{y}}{\sigma_{x}} \mu_{x}+\mu_{y}$$


$$ \text{tiempo de espera} = \text{pendiente} \cdot \text{duración de erupción}  + \text{pendiente} \cdot \text{promedio de duración} +\text{promedio de espera} $$


**Pregunta 2.2.** Calcula el término de intersección en unidades originales y asígnalo a `intercept`. La [Sección 15.2.5](https://inferentialthinking.com/chapters/15/2/Regression_Line.html#the-equation-of-the-regression-line) puede ser útil.



```python
intercept= ...
```

```python
grader.check("q8")
```

## 3. Investigando la línea de regresión

La pendiente y la intersección de la línea de regresión definen su forma. Para predecir el tiempo de espera de una erupción, multiplica la duración de la erupción por `pendiente` y luego suma `intersección`.

**Pregunta 3.1.** Calcula el tiempo de espera previsto para una erupción que dura 2 minutos y para una erupción que dura 5 minutos.



```python
two_minute_predicted_waiting_time = ...
five_minute_predicted_waiting_time = ...

def print_prediction(duration, predicted_waiting_time):
    print("After an eruption lasting", duration,
          "minutes, we predict you'll wait", predicted_waiting_time,
          "minutes until the next eruption.")

print_prediction(2, two_minute_predicted_waiting_time)
print_prediction(5, five_minute_predicted_waiting_time)

```

```python
grader.check("q9")
```

La siguiente celda traza la línea que conecta esos dos puntos, lo cual representa un segmento de la línea de regresión. Esta línea sirve para modelar la relación entre la duración de las erupciones y el tiempo de espera en el archivo `faithful.csv`, que contiene las columnas `duration` y `wait`.

```python
plot_data_and_line(faithful, "duration", "wait", 
                   [2, two_minute_predicted_waiting_time], 
                   [5, five_minute_predicted_waiting_time])
```

**Pregunta 3.2.** Realiza predicciones sobre el tiempo de espera después de cada erupción en la tabla `faithful`. (Por supuesto, ya conocemos exactamente cuáles son los tiempos de espera. Hacemos esto para ver qué tan precisas son nuestras predicciones). Coloca estos números en una columna de una nueva tabla llamada `faithful_predictions`. La primera fila debería verse así:

|duración|espera|espera pronosticada|
|-|-|-|
|3.6|79|72.1011|

*Sugerencia:* Tu respuesta puede consistir en solo una línea, aunque no estás limitado a una sola línea. No es necesario utilizar un bucle `for`; emplea aritmética de arreglos en su lugar.

```python
faithful_predictions = ...

```

```python
np.mean(y)
```

```python
grader.check("q10")
```

**Pregunta 3.3.** ¿Qué tan cerca estuvimos? Calcula el *residuo* para cada erupción en el conjunto de datos. El residuo es el tiempo de espera real menos el tiempo de espera predicho. Añade los residuos a `faithful_predictions` como una nueva columna llamada `residual` y nombra la tabla resultante `faithful_residuals`.

Como se menciona en el [Capítulo 15.6](https://inferentialthinking.com/chapters/15/6/Numerical_Diagnostics.html?highlight=residuals#average-of-residuals), una propiedad útil de los residuos producidos por la regresión lineal con un intercepto es que la suma de los residuos siempre será cero. Siempre puedes verificar esto para asegurarte de que has calculado correctamente tus residuos. Nota que puede estar ligeramente desviado de cero debido a errores de redondeo.

¡Siéntete libre de copiar y pegar esta línea de código para verificar tu trabajo! Ten en cuenta que la respuesta puede ser *muy ligeramente* diferente de cero debido a errores de redondeo en Python.

```python
sum(faithful_residuals.column('residual'))
```

*Pista:* Nuevamente, tu código será mucho más simple si no utilizas un bucle `for`.

```python
faithful_predictions
```

```python
residuals = ...
faithful_residuals = ...
faithful_residuals
```

```python
grader.check("q11")
```

Aquí tienes un gráfico de los residuos que calculaste. Cada punto representa una erupción del géiser Old Faithful. El gráfico muestra cuánto nuestra predicción sobrestimó o subestimó el tiempo de espera. 

```python
faithful_residuals.scatter("duration", "residual", color="r")
```

No hay realmente un patrón en los residuos, lo cual confirma que fue razonable intentar una regresión lineal. Es cierto que hay dos grupos separados; las duraciones de las erupciones parecen dividirse en dos grupos distintos. Sin embargo, eso es solo un patrón en las duraciones de las erupciones, no un patrón en la relación entre las duraciones de las erupciones y los tiempos de espera.

## 4. ¿Qué tan precisas son diferentes predicciones?

Anteriormente, deberías haber encontrado que la correlación es bastante cercana a 1, lo que indica que la línea se ajusta bastante bien a los datos de entrenamiento. Esto significa que, en general, los residuos son pequeños (cercanos a 0) en comparación con los tiempos de espera.

Podemos observar esto visualmente al graficar los tiempos de espera y los residuos juntos:



```python
# Just run this cell.
faithful_residuals.scatter("duration", "wait", label="actual waiting time", color="blue")
plots.scatter(faithful_residuals.column("duration"), faithful_residuals.column("residual"), label="residual", color="r")
plots.plot([2, 5], [two_minute_predicted_waiting_time, five_minute_predicted_waiting_time], label="regression line")
plots.legend(bbox_to_anchor=(1.7,.8));
```

A menos que tengas una razón convincente para creer que el modelo de regresión lineal es preciso, debes tener cuidado al aplicar tu modelo de predicción a datos que son muy diferentes de los datos de entrenamiento.

**Pregunta 4.1.** En el conjunto de datos `faithful`, no hay erupciones que duren exactamente 0, 2.5 o 60 minutos. Basado en esta información, ¿cuál sería el tiempo de espera predicho para una erupción que dure 0 minutos? ¿Y para una que dure 2.5 minutos? ¿Y para una hora?



```python
zero_minute_predicted_waiting_time = ...
two_point_five_minute_predicted_waiting_time = ...
hour_predicted_waiting_time = ...

print_prediction(0, zero_minute_predicted_waiting_time)
print_prediction(2.5, two_point_five_minute_predicted_waiting_time)
print_prediction(60, hour_predicted_waiting_time)

```

```python
grader.check("q12")
```

**Pregunta 4.2.** Para cada predicción, indica si consideras que es confiable y explica tu razonamiento.

Basado en los resultados de la predicción, analiza la consistencia y validez de cada estimación. Considera aspectos como el ajuste del modelo, la variabilidad de los datos y otros factores que puedan influir en la precisión de las predicciones.

_Escribe tu respuesta aquí, reemplazando este texto._


## 5. Divide y vencerás  

En esta sección, exploraremos la estrategia de "divide y vencerás". 

El objetivo es comprender cómo resolver problemas complejos dividiéndolos en partes más manejables y luego combinando sus resultados. Este enfoque es particularmente útil en regresión lineal, ya que nos permite desglosar los datos para obtener predicciones más precisas.  


El diagrama de dispersión a continuación muestra la relación entre la duración de las erupciones en el eje x y el tiempo de espera hasta la siguiente erupción en el eje y. 

Del diagrama de dispersión parece que hay dos grupos de puntos: uno para duraciones alrededor de 2 minutos y otro para duraciones entre 3.5 y 5 minutos. Una línea vertical en 3 minutos divide estos dos grupos. 



```python
faithful.scatter("duration", "wait", label="actual waiting time", color="blue")
plots.plot([3, 3], [40, 100]);
```

La función `standardize`, presentada en clase, se muestra a continuación. Esta función recibe una tabla con columnas numéricas y devuelve la misma tabla, pero con cada columna transformada a su versión estandarizada.  


```python
# Just run this cell.

def standard_units(any_numbers):
    "Convert any array of numbers to standard units."
    return (any_numbers - np.mean(any_numbers)) / np.std(any_numbers)  

def standardize(table):
    """Return a table in which all columns of t are converted to standard units."""
    t_su = Table()
    for label in table.labels:
        t_su = t_su.with_column(label + '_su', standard_units(table.column(label)))
    return t_su
```

**Pregunta 5.1.** Calcula por separado el coeficiente de correlación *r* para todos los puntos con una duración inferior a 3 **y luego** para todos los puntos con una duración superior a 3. Para ello, crea una función que calcule `r` a partir de una tabla y luego pásale dos tablas diferentes de puntos, llamadas `below_3` y `above_3`.

*Sugerencia:* Puedes asumir que la tabla no tiene valores de duración exactamente de 3. 

```python
faithful = Table.read_table('faithful.csv')
t_su = standardize(faithful)
t_su
```

```python
def corr_coeff(table):
    '''Return the regression coefficient for columns 0 & 1.'''
    t_su = standardize(table)
    ...



below_3 = ...
above_3 = ...
below_3_r = corr_coeff(above_3_r)
above_3_r = corr_coeff(above_3)
print("For points below 3, r is", below_3_r, "; for points above 3, r is", above_3_r)

```

```python
grader.check("q14")
```

**Pregunta 5.2.** Completa las funciones `slope_of` e `intercept_of` a continuación.  

Una vez finalizadas, las funciones `wait_below_3` y `wait_above_3` deberán utilizar líneas de regresión distintas para predecir el tiempo de espera en función de la duración.  

- `wait_below_3` usará la línea de regresión para todos los puntos cuya duración sea menor a 3.  
- `wait_above_3` usará la línea de regresión para todos los puntos cuya duración sea mayor a 3.  


```python
faithful
```

```python
above_3_r
```

```python
def slope_of(table, r):
    '''Return the slope of the regression line for table in original units.
    
    Assume that column 0 contains x values and column 1 contains y values.
    r is the regression coefficient for x and y.
    '''
    ...

def intercept_of(table, r):
    '''Return the intercept of the regression line for table in original units.
    
    Assume that column 0 contains x values and column 1 contains y values.
    r is the regression coefficient for x and y.
    '''
    slope = slope_of(table, r)
    ...
    
below_3_slope = slope_of(below_3, below_3_r)
below_3_intercept = intercept_of(below_3, below_3_r)
above_3_slope = slope_of(above_3, above_3_r)
above_3_intercept = intercept_of(above_3, above_3_r)

def wait_below_3(duration):
    return below_3_slope * duration + below_3_intercept

def wait_above_3(duration):
    return above_3_slope * duration + above_3_intercept
```

```python
print(wait_below_3(1))
print(wait_below_3(3))
print(wait_above_3(3))
print(wait_above_3(6))
```

```python
grader.check("q15")
```

El gráfico a continuación muestra las dos líneas de regresión diferentes, una para cada grupo, junto con la línea de regresión original.  

```python
faithful.scatter(0, 1)
plots.plot([2, 5], [two_minute_predicted_waiting_time, five_minute_predicted_waiting_time])
plots.plot([1, 3], [wait_below_3(1), wait_below_3(3)])
plots.plot([3, 6], [wait_above_3(3), wait_above_3(6)]);
```

**Pregunta 5.3.** Escribe una función llamada `predict_wait` que reciba un parámetro `duration` y devuelva el tiempo de espera previsto utilizando la línea de regresión adecuada, dependiendo de si la duración es menor o mayor a 3.

*Consejo:* Considera usar declaraciones condicionales. 


```python
def predict_wait(duration):
    ...

```

```python
predict_wait(2)
```

```python
grader.check("q16")
```

A continuación, se presentan los tiempos de espera predichos para cada punto.

```python
faithful_pred_split = faithful.with_column('predicted', faithful.apply(predict_wait, 'duration'))
faithful_pred_split.scatter(0)
```

**Pregunta 5.4.** ¿Crees que las predicciones generadas por la función `predict_wait` serían más o menos precisas que las predicciones de la línea de regresión que creaste en la sección 2? ¿Cómo podrías determinarlo?

**SOLUCIÓN:** Más precisa, ya que cada línea es específica para los valores en su clúster. Para verificarlo, podríamos medir la magnitud promedio de los valores residuales. 

La siguiente celda generará un gráfico de los residuos para cada erupción en el conjunto de datos cuando utilizamos una línea de regresión simple y cuando utilizamos dos líneas de regresión separadas. También se mostrará la magnitud promedio de los valores residuales. 

```python
# Just run this cell
faithful_pred_split_residuals = faithful_pred_split.with_column('residual', faithful_pred_split.column(1) - faithful_pred_split.column(2))
plots.scatter(faithful_residuals.column('duration'), faithful_residuals.column('residual'), label='one regression line')
plots.scatter(faithful_pred_split_residuals.column('duration'), faithful_pred_split_residuals.column('residual'), label='two regression lines');
plots.axis([1, 6, -15, 15])
plots.legend(bbox_to_anchor=(1.5,.8));
print("Average Magnitude of Residual Values for One Regression Line: ", np.mean(abs(faithful_residuals.column('residual'))))
print("Average Magnitude of Residual Values for Two Regression Lines: ", np.mean(abs(faithful_pred_split_residuals.column('residual'))))
```

El gráfico de residuos para los tiempos de espera, cuando son predichos por dos líneas de regresión (en rojo), no presenta realmente un patrón, lo que confirma que también fue apropiado usar regresión lineal en nuestro escenario de "Dividir y Conquistar". ¿Cómo se comparan los dos gráficos de residuos?

## ¡Terminado!

<div style="display: flex; justify-content: center; gap: 5px;">
    <img src="Manchas.png" alt="Manchas" width="200"/>
    <img src="Nina.jpg" alt="Nina" width="200"/>
    <img src="Nashira.jpeg" alt="Nashira" width="200"/>
</div>

¡**Manchas, Nina y Nashira** te felicitan. Haz terminado el último laboratorio.

¡Buen trabajo!

¡Felicitaciones por completar la Laboratorio 10!

¡Terminaste con el laboratorio!

**Información importante sobre el envío:**
- **Ejecuta todas las pruebas** y verifica que todas pasan
- **Guardar** desde el menú **Archivo**
- **Ejecute la celda final **
- Luego, ve a este [enlace](https://cienciadatosudea.github.io/001_SitioWebFundCienciaDatos/04entregas.html) y envía el archivo ipynb a la tarea correspondiente. El nombre de esta tarea es "Lab XX Autograder", donde XX es el número de laboratorio: 01, 02, 03, etc.

- Si terminas temprano en el laboratorio regular, **registra tu asistencia con el instructor**.

**Es su responsabilidad asegurarse de que su trabajo esté guardado antes de ejecutar la última celda.**

`Referencia`: Este material esta basado en [Data 8: The Foundations of Data Science](https://www.data8.org/) - University of California, Berkeley.


```python

```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

These are some submission instructions.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q0 (points: 1)
```
>>> assert np.isclose(least_squares_order[0] * 3 + least_squares_order[1] * 12 + least_squares_order[2] * 3 + +least_squares_order[3] * 6, 4 * 12, 0.1)
```


Tests for q10 (points: 1)
```
>>> assert np.isclose(np.mean(faithful_predictions.column(1)), 70.89, 0.01)
```


Tests for q11 (points: 1)
```
>>> assert np.isclose(sum(faithful_residuals.column('residual')), 0, atol=0.01)
```


Tests for q12 (points: 1)
```
>>> assert np.isclose(zero_minute_predicted_waiting_time, 33.47, atol=0.01)
>>> assert np.isclose(two_point_five_minute_predicted_waiting_time, 60.29, atol=0.01)
>>> assert np.isclose(hour_predicted_waiting_time, 677.25, atol=0.01)
```


Tests for q14 (points: 1)
```
>>> assert np.isclose(below_3_r, 0.29, atol=0.01)
>>> assert np.isclose(above_3_r, 0.372, atol=0.001)
```


Tests for q15 (points: 1)
```
>>> assert np.isclose(wait_below_3(1), 47.9, 0.01)
>>> assert np.isclose(wait_above_3(6), 89.18, 0.01)
```


Tests for q16 (points: 1)
```
>>> assert np.isclose(predict_wait(2), 54.25, 0.001)
>>> assert np.isclose(predict_wait(10), 111.03, 0.0001)
```


Tests for q3 (points: 1)
```
>>> assert np.isclose(mean_duration, 3.4, 0.1)
>>> assert np.isclose(std_duration, 1.1, 0.1)
>>> assert np.isclose(mean_wait, 70.8, 0.1)
>>> assert np.isclose(std_wait, 13.56, 0.1)
>>> assert np.isclose(np.mean(duration_standard_units), 0.0, 0.1)
>>> assert np.isclose(np.std(wait_standard_units), 1.0, 0.1)
```


Tests for q5 (points: 1)
```
>>> assert np.isclose(correlation * 4, 12, 0.1)
```


Tests for q6 (points: 1)
```
>>> assert np.isclose(r * 5, 4.5, atol=0.1)
```


Tests for q7 (points: 1)
```
>>> assert np.isclose(slope * 5, 53.6, 0.001)
```


Tests for q8 (points: 1)
```
>>> assert np.isclose(round(intercept, 2) + 0.53 + 20, 54, 0.1)
```


Tests for q9 (points: 1)
```
>>> assert isinstance(two_minute_predicted_waiting_time, (int, float)), 'El resultado debe ser un número'
>>> assert isinstance(five_minute_predicted_waiting_time, (int, float)), 'El resultado debe ser un número'
>>> assert np.isclose(two_minute_predicted_waiting_time * 5, 274.66, 0.01)
>>> assert np.isclose(five_minute_predicted_waiting_time * 5, 435.61, 0.01)
```


Tests for qt (points: 0)
```
>>> assert isinstance(NombreApellidos, str)
>>> assert isinstance(NumeroCarnet, str)
>>> assert isinstance(Email, str)
>>> assert len(NombreApellidos) > 3
>>> assert len(NumeroCarnet) > 3
>>> assert len(Email) > 3
```
