```python
# Initialize Otter
import otter
grader = otter.Notebook("proyecto_01_parte2.ipynb")
```

```python
# Run this cell to set up the notebook, but please don't change it.
# These lines import the NumPy and Datascience modules.
from datascience import *
import numpy as np
# These lines do some fancy plotting magic.
%matplotlib inline
import matplotlib.pyplot as plots
plots.style.use('fivethirtyeight')
from ipywidgets import interact, interactive, fixed, interact_manual
import ipywidgets as widgets

```

_Type your answer here, replacing this text._

```python
Nombre = "..."
Carnet = "..."
```

<!-- END QUESTION -->

## 1. Crecimiento de la Población Mundial

La población mundial de seres humanos alcanzó los 1000 millones alrededor del año 1800; 3000 millones alrededor de 1960, y 7000 millones alrededor de 2011. El impacto potencial del crecimiento exponencial de la población ha preocupado a científicos, economistas y políticos por igual. La División de Población de las Naciones Unidas, estima que es probable que la población mundial continúe creciendo a lo largo del siglo XXI, pero a un ritmo más lento, quizás alcanzando los 11000 millones para el año 2100. Sin embargo, la ONU no descarta escenarios de un crecimiento más extremo.<a href="http://www.pewresearch.org/fact-tank/2015/06/08/scientists-more-worried-than-public-about-worlds-growing-population/ft_15-06-04_popcount/">  <img src="pew_population_projection.png"/> </a>

En esta parte del proyecto, examinaremos algunos de los factores que influyen en el crecimiento de la población y cómo han ido cambiando a lo largo de los años y en todo el mundo. Hay dos subpartes principales en este análisis.

- Primero, examinaremos los datos de un país, Bangladesh. Veremos cómo han cambiado factores como la esperanza de vida, la tasa de fertilidad y la mortalidad infantil a lo largo del tiempo en Bangladesh, y cómo se relacionan con la tasa de crecimiento de la población.

- A continuación, examinaremos si los cambios que hemos observado para Bangladesh son particulares de ese país o si reflejan patrones generales observables en otros países también. Estudiaremos aspectos del crecimiento de la población mundial y veremos cómo han estado cambiando.La primera tabla que consideraremos contiene la población total de cada país a lo largo del tiempo. Ejecuta la celda a continuación.

```python
population = Table.read_table('population.csv')
fertility = Table.read_table('fertility.csv')
child_mortality = Table.read_table('child_mortality.csv')

```

**Nota:** El archivo csv de población también se puede encontrar [aquí](https://github.com/open-numbers/ddf--gapminder--gapminder_world/blob/master/ddf--datapoints--population_total--by--geo--time.csv). Los datos para este proyecto fueron descargados en febrero de 2017.

# Segunda parte del proyecto. 


### El Mundo


Los cambios demográficos observados en Polonia, como en muchos otros países, reflejan una tendencia común que se puede analizar con datos globales. A medida que los países experimentan mejoras en sus servicios de salud, la esperanza de vida tiende a aumentar y la mortalidad infantil disminuye. Sin embargo, estas mejoras suelen estar acompañadas por una caída en la tasa de fertilidad, lo que genera una desaceleración en el crecimiento de la población. Este fenómeno,  es clave para entender los retos y oportunidades del desarrollo global. A pesar de que las personas viven más tiempo, el número de nacimientos no suele ser suficiente para sostener el crecimiento poblacional, lo que plantea desafíos económicos y sociales que pueden estudiarse con técnicas de ciencia de datos.

Ejecuta la celda a continuación para generar dos histogramas superpuestos, uno para 1962 y otro para 2010, que muestran las distribuciones de las tasas de fertilidad total para estos dos años entre los 201 países en la tabla `fertility`.

```python
Table().with_columns(
    '1962', fertility.where('time', 1962).column(2),
    '2010', fertility.where('time', 2010).column(2)
).hist(bins=np.arange(0, 10.5, 0.5), unit='child per woman')
_ = plots.xlabel('Children per woman')
_ = plots.ylabel('Percent per children per woman')
_ = plots.xticks(np.arange(10))
```

```python
q=Table().with_columns(
    '1962', fertility.where('time', 1962).column(2),
    '2010', fertility.where('time', 2010).column(2)
).bin(bins=np.arange(0, 10.5, 0.5))
q.show()
```

```python
q.where("bin", are.between(0,3))
```

```python
a=3
b=10
q.where("bin", are.between(a,b))
```

```python
q.where("bin", are.between(a,b)).to_df().sum(axis=0)
```

```python
q.to_df().sum(axis=0)
```

```python
82/202
```


- Eje X : El eje X está etiquetado como "Children per woman", es decir, el número de hijos por mujer. Indica el número promedio de hijos que una mujer tendría en su vida en cada país, que esta relacionado con la tasa de fertilidad, que se refiere generalmente a la cantidad promedio de hijos por mujer en un país.


- Eje Y : El eje Y está etiquetado como "Percent per children per woman", Indica el porcentaje de países que tienen ese número específico de hijos por mujer, para los años 1962 y 2010.


**Pregunta 9.** Observa los dos histogramas que comparan las tasas de fertilidad de los países en 1962 y 2010. A continuación, se proporcionan varias afirmaciones que podrían inferirse de estos gráficos.
Tu tarea es asignar a `fertility_statements` un arreglo que contenga valores booleanos, asociados a cada  afirmacion,  que se pueden inferir correctamente a partir de los histogramas.

Ejemplo: Si las afirmaciones correctas son la 0, la 3 y la 5,  y las falsas son la 1, 2, 4,  entonces el arreglo se puede definir así:

`fertility_statements = make_array(False, False, True, True, True, True)`

Estas son las afirmaciones:

0. Aproximadamente el mismo número de países tuvo una tasa de fertilidad entre 3.5 y 4.5, tanto en 1962 como en 2010.

1. En 1962, menos del 20% de los países tenían una tasa de fertilidad inferior a 3.

2.  Al menos la mitad de los países tuvieron una tasa de fertilidad entre 5 y 8 en 1962.

3. En 2010, alrededor del 40% de los países tenían una tasa de fertilidad entre 1.5 y 2.

4.  Al menos la mitad de los países tuvieron una tasa de fertilidad inferior a 3 en 2010.

5. Más países tenían una tasa de fertilidad superior a 3 en 1962 que en 2010.

```python
fertility_statements = make_array(....)
```

```python
grader.check("q9")
```

**Pregunta 10.**  La **poblacion mundial** es la suma de las poblaciones de todos los países, cálcule esa suma en una tabla llamada `world_population` y después realice un gráfico de líneas de la poblacion mundial desde 1800 hasta 2020 (inclusive de ambos extremos). Debes usar la tabla `population` definida anteriormente en el proyecto, el nombre de las columnas de `world_population` deberá ser: 
`time`, `population_total sum`



```python

population = Table.read_table('population.csv')
world_population = ...
...
...
```

```python
grader.check("q10")
```

**Pregunta 11.** Crea una función `stats_for_year` que tome un `year` (año) como argumento y devuelva una tabla de estadísticas:

- La tabla que devuelva debe tener cuatro columnas: `geo` (geografía), `population_total` (población total), `children_per_woman_total_fertility` (número de hijos por mujer - tasa de fecundidad total), y `child_mortality_under_5_per_1000_born` (mortalidad infantil bajo los 5 años por cada 1000 nacidos).

- Cada fila debe contener un código de país único en formato Alpha-3 y tres estadísticas: población, tasa de fecundidad y mortalidad infantil para ese `year` (año) de las tablas `population` (población), `fertility` (fecundidad) y `child_mortality` (mortalidad infantil). Solo incluye filas para las cuales las tres estadísticas estén disponibles para el país y año. Además, restringe el resultado a los códigos de país que aparecen en `big_50`, un arreglo de los 50 países más poblados en 2020. Esta restricción acelerará los cálculos más adelante en el proyecto.

- Después de escribir `stats_for_year`, intenta llamar a `stats_for_year` en cualquier año entre 1960 y 2020. Trata de entender el resultado de `stats_for_year`.


*Pista 1*: Las pruebas para esta pregunta son bastante completas, así que si pasas las pruebas, tu función probablemente sea correcta. Sin embargo, sin llamar a tu función tú mismo y observar el resultado, será muy difícil entender cualquier problema que tengas, así que haz tu mejor esfuerzo para escribir la función correctamente y verifica que funcione antes de confiar en las pruebas del `grader` para confirmar tu trabajo.

*Pista 2*: ¿Qué tienen en común las tres tablas (presta atención a los nombres de las columnas)?

*Pista 3*: Crea celdas adicionales antes de escribir directamente la función.

```python


#Primero creamos una tabla de población que solo incluye los 
#50 países con las mayores poblaciones en 2020. Nos enfocamos 
#solo en estos 50 países para que los gráficos más adelante se ejecuten más rápido.


big_50 = population.where('time', are.equal_to(2020)).sort("population_total", descending=True).take(np.arange(50)).column('geo')
population_of_big_50 = population.where('time', are.above(1959)).where('geo', are.contained_in(big_50))

def stats_for_year(year):
    p = population_of_big_50.where('time', are.equal_to(year)).drop('time')
    f = fertility.where('time', are.equal_to(year)).drop('time')
    c = child_mortality.where('time', are.equal_to(year)).drop('time')
    ...
    ...

```

```python
grader.check("q11")
```

**Pregunta 12.** Crea una tabla llamada `pop_por_decada` con dos columnas: `decada` y `poblacion`, en ese orden. 

- La tabla debe tener una fila para cada año que inicia una década, en orden ascendente desde 1960 hasta 2020. Por ejemplo, 1960 marca el comienzo de la década de los 60.

- La columna `poblacion` debe contener la población total de todos los países incluidos en el resultado de `stats_for_year(year)` para el primer año de cada década. Deberías notar que estos países representan la mayor parte de la población mundial.

**Consejo:** Un enfoque útil sería definir una función llamada `poblacion_para_ano` que calcule la población total, y luego aplicar esa función a la columna `decada`. **Piensa en cómo podrías usar la función `stats_for_year`** de la pregunta anterior si decides implementar `poblacion_para_ano`. Esta primera prueba es simplemente una verificación de coherencia para tu función auxiliar si decides utilizarla. **No perderás puntos por no implementar la función `poblacion_para_ano`.**

**Nota:** La celda en la que generarás la tabla `pop_por_decada` está justo debajo de la celda donde puedes optar por definir la función auxiliar `poblacion_para_ano`. Debes definir tu tabla `pop_por_decada` en la celda que comienza con la tabla `decadas` ya definida.


```python
def pop_for_year(year):
    '''Return the total population for the specified year.'''
    ...
    
# Ahora que has definido tu función auxiliar (si has decidido hacerlo), define la tabla `pop_by_decade`.
pop_por_decada = ...

pop_por_decada.set_format(1, NumberFormatter)
```

```python
pop_por_decada
```

```python
grader.check("q12")
```

Ahora que has definido tu función auxiliar (si has decidido hacerlo), define la tabla `pop_by_decade`.

La tabla `countries` describe diversas características de los países. La columna `country` contiene los mismos códigos que la columna `geo` en cada una de las otras tablas de datos (`population`, `fertility` y `child_mortality`). La columna `world_6region` clasifica a cada país en una región del mundo. Ejecuta la celda de abajo para inspeccionar los datos.

```python
countries = Table.read_table('countries.csv').where('country', are.contained_in(population.group('geo').column('geo')))
countries.select('country', 'name', 'world_6region')
```

**Pregunta 13.** Crea una tabla llamada `region_counts`. Debe contener dos columnas llamadas `region` y `count`. La columna `region` debe contener las regiones del mundo, y la columna `count` debe contener el número de países en cada región que aparece en el resultado de `stats_for_year(2020)`.Por ejemplo, una fila tendría `south_asia` como su valor de `region` y un número entero como su valor de `count`: la cantidad de países grandes de Asia del Sur para los cuales tenemos cifras de población, fertilidad y mortalidad infantil del año 2020.


*Sugerencia*: Es posible que tengas que renombrar una columna para llamarla `region`.

```python
year_2020_data = stats_for_year(2020)
region_counts = ...
region_counts
```

```python
grader.check("q13")
```

El siguiente diagrama de dispersión compara la tasa de fecundidad total y la tasa de mortalidad infantil para cada país en 1960. El área de cada punto representa la población del país, y el color representa su región del mundo. Ejecuta la celda. ¿Crees que puedes identificar alguno de los puntos?

```python
from functools import lru_cache as cache

# This cache annotation makes sure that if the same year
# is passed as an argument twice, the work of computing
# the result is only carried out once.
@cache(None)
def stats_relabeled(year):
    """Relabeled and cached version of stats_for_year."""
    return stats_for_year(year).relabel(2, 'Children per woman').relabel(3, 'Child deaths per 1000 born')

def fertility_vs_child_mortality(year):
    """Draw a color scatter diagram comparing child mortality and fertility."""
    with_region = stats_relabeled(year).join('geo', countries.select('country', 'world_6region'), 'country')
    with_region.scatter(2, 3, sizes=1, group=4, s=500)
    plots.xlim(0,10)
    plots.ylim(-50, 500)
    plots.title(year)
    plots.show()

fertility_vs_child_mortality(1960)
```

```python
with_region = stats_relabeled(1960).join('geo', countries.select('country', 'world_6region'), 'country')

with_region.sort("population_total", descending=True)
```

```python
with_region.sort("population_total", descending=True).sort("Child deaths per 1000 born", descending=True)
```

**Pregunta 14.** Asigna `scatter_statements` a un arreglo con los valores booleanos,  de cada una de las afirmaciones que se pueden inferir de este diagrama de dispersión para 1960, asigna True o False al arreglo, ejemplo :

make_array(True, True, True, False, False)` :

1. Todos los países en `europe_central_asia` tenían tasas de fertilidad uniformemente bajas.
2. La tasa de mortalidad infantil más baja de cualquier país fue en un país de `east_asia_pacific`.
3. La mayoría de los países tiene una tasa de fertilidad superior a 5.
4. Hay una asociación entre la mortalidad infantil y la fertilidad.
5. Los dos países más grandes por población también tuvieron las dos tasas de mortalidad infantil más altas.



```python
scatter_statements = make_array(...)
```

```python
grader.check("q14")
```

El resultado de la celda a continuación es interactivo. Arrastra el deslizador hacia la derecha para ver cómo han cambiado los países a lo largo del tiempo. Descubrirás que, en términos de crecimiento poblacional, la división entre los países del Norte global y el Sur global que existía en los años 1960 se ha reducido significativamente.Este cambio en las tasas de fertilidad es la razón por la cual se espera que la población global crezca más lentamente en el siglo XXI de lo que lo hizo en los siglos XIX y XX. Las tasas de fertilidad cambian por razones que incluyen patrones culturales, mejores perspectivas de supervivencia infantil hasta la adultez y la planificación familiar (como el uso de anticonceptivos y el mayor control de las mujeres sobre su reproducción).**Nota:** No te preocupes si aparece una advertencia roja al ejecutar la celda a continuación. ¡Aún podrás ejecutarla!

```python
_ = widgets.interact(fertility_vs_child_mortality, 
                     year=widgets.IntSlider(min=1960, max=2020, value=1960))
```

Ahora es un buen momento para tomar un descanso y ver los mismos datos presentados por [Hans Rosling en una charla TEDx de 2010](https://www.ted.com/talks/hans_rosling_the_good_news_of_the_decade_we_re_winning_the_war_against_child_mortality?subtitle=en&geo=es) con una animación y algunos comentarios importantes.

**Berli, Kiriku y Rocky quiere decirte, ¡has terminado!** Felicitaciones por descubrir muchos hechos importantes a base de ciencia de datos y demostrar tu dominio en la manipulación de tablas y visualización de datos. Es momento de enviar tu trabajo.**<span style="color: #BC412B"></span>** <u></u>


<div style="display: flex; justify-content: space-between; align-items: center;">
    <div style="flex: 1; text-align: center;">
        <img src="Rocky.jpeg" alt="Rocky el golden retriever" width="300"/>
        <p><b>Rocky: <strong>¡¡¡WOOOHOO!!!</strong> </b></p>
    </div>
    <div style="flex: 1; text-align: center;">
        <img src="Kiriku.jpg" alt="Kiriku el golden retriever" width="300"/>
        <p><b>Kiriku: <strong>¡¡¡MIIIAAAUU!!!</strong></b></p>
    </div>
    <div style="flex: 1; text-align: center;">
        <img src="Berlin.jpg" alt="Berlin el golden retriever" width="300"/>
        <p><b>Berlin: <strong>¡¡¡MIIIAAAUU!!!</strong></p></b></p>
    </div>
</div>

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

These are some submission instructions.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q10 (points: 1)
```
>>> assert 'time' in world_population.labels, "La columna 'time' no está presente en world_population."
>>> assert 'population_total sum' in world_population.labels, "La columna 'population_total sum' no está presente en world_population."
>>> assert np.isclose(world_population.column('time').mean(), 1910, 0.001)
>>> assert np.isclose(world_population.column('population_total sum').mean(), 2626917081.678733, 0.001)
```


Tests for q11 (points: 1)
```
>>> assert np.isclose(np.mean(stats_for_year(2018).column('population_total')), 132609584.28, 0.01)
>>> assert np.isclose(np.mean(stats_for_year(2010).column('population_total')), 121293482.64, 0.01)
>>> assert np.isclose(np.mean(stats_for_year(2000).column('population_total')), 107541243.38, 0.01)
>>> assert np.isclose(np.mean(stats_for_year(2010).column('children_per_woman_total_fertility')), 2.9691999999999994, 0.01)
>>> assert np.isclose(np.mean(stats_for_year(2018).column('children_per_woman_total_fertility')), 2.7053999999999996, 0.01)
>>> assert np.isclose(np.mean(stats_for_year(2000).column('children_per_woman_total_fertility')), 3.33559999999, 0.01)
>>> a = []
>>> a.append(stats_for_year(2000).to_df().describe()['population_total'].iloc[0])
>>> a.append(stats_for_year(2000).to_df().describe()['population_total'].iloc[1])
>>> a.append(stats_for_year(2000).to_df().describe()['population_total'].iloc[2])
>>> a.append(stats_for_year(2000).to_df().describe()['population_total'].iloc[3])
>>> assert a[0] == 50
>>> assert a[1] == 107541243.38
>>> assert a[2] == 227538018.45596266
>>> assert a[3] == 16395477.0
```


Tests for q12 (points: 1)
```
>>> assert pop_por_decada.to_df().shape == (7, 2), 'La tabla resultante no tiene las dimensiones correctas'
>>> assert np.isclose(pop_por_decada.to_df().loc[0, 'poblacion'], 2635123897, 0.01)
>>> assert np.isclose(pop_por_decada.to_df().loc[6, 'poblacion'], 6765161289, 0.01)
```


Tests for q13 (points: 1)
```
>>> df = region_counts.to_df()
>>> assert np.isclose(df['count'].sum(), 50, 1e-06)
```


Tests for q14 (points: 1)
```
>>> u = make_array(False, False, True, True, False)
>>> assert scatter_statements[0] == u[0]
>>> assert scatter_statements[1] == u[1]
>>> assert scatter_statements[2] == u[2]
>>> assert scatter_statements[3] == u[3]
>>> assert scatter_statements[4] == u[4]
```


Tests for q9 (points: 1)
```
>>> u = make_array(False, False, True, True, True, True)
>>> assert fertility_statements.sum() == u.sum()
>>> assert fertility_statements[0] == u[0]
>>> assert fertility_statements[1] == u[1]
>>> assert fertility_statements[2] == u[2]
>>> assert fertility_statements[3] == u[3]
>>> assert fertility_statements[4] == u[4]
>>> assert fertility_statements[5] == u[5]
```
