```python
# Initialize Otter
import otter
grader = otter.Notebook("lab09.ipynb")
```

<p><img alt="Colaborat" style="height: 220px;" src="https://github.com/CienciaDatosUdea/001_DataSetFundCienciaDatos/blob/main/semestre2024-2/lab/lab02/Escudo-UdeA.svg.png?raw=true" align="left" hspace="10px" vspace="0px"></p>


<h1> Fundamentos en ciencias de datos</h1>

# Laboratorio 9: Distribución normal y variabilidad de medias muestrales

¡Bienvenidos al Laboratorio 0!

En el laboratorio de hoy, aprenderemos sobre [la variabilidad de las medias muestrales](https://inferentialthinking.com/chapters/14/5/Variability_of_the_Sample_Mean.html), así como [la distribución normal](https://www.inferentialthinking.com/chapters/14/3/SD_and_the_Normal_Curve.html).



```python
# Ejecute esta celda, pero no la cambie.

# Estas líneas importan los módulos Numpy y Datascience.
import numpy as np
from datascience import *

# Estas líneas hacen algo de magia en la trama.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plots
plots.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore')
```

<!-- BEGIN QUESTION -->

Ingresa tu nombre  con apellido y número de carnet en las variables dadas en la celda siguiente:
Ejemplo:
```python
NombreApellidos="Ada Lovelace"
NumeroCarnet="27182818"
Email="Ada.Lovelace@udea.edu.co"
```

```python
NombreApellidos="..."
NumeroCarnet=".."
Email="..."
```

```python
grader.check("qt")
```

<!-- END QUESTION -->

# 1. Distribuciones normales

Cuando visualizamos la distribución de una muestra, a menudo nos interesa la media y la desviación estándar de la muestra (durante el resto de esta práctica de laboratorio, abreviaremos “desviación estándar” como “SD” -standar desviation-). Estas dos estadísticas de muestra pueden darnos una visión general de la distribución, permitiéndonos saber dónde se ubica la distribución en la recta numérica y qué tan extendida está, respectivamente.

Queremos comprobar si los datos están relacionados linealmente, por lo que deberíamos mirar los datos.

Primero, carguemos la tabla de "baby.csv"  que es una muestra aleatoria grande de nacimientos e incluye información sobre las parejas de madre e hijo.

```python
births = Table.read_table('baby.csv')
births
```

**Pregunta 1.1.** **Ahora, escriba una línea de código para trazar la distribución de las edades de la madre a partir de la tabla `births`.** No cambie la última línea, que trazará la media de la muestra en el distribución misma.


```python
#...
# No cambies nada debajo de esta línea
plots.scatter(np.mean(births.column("Maternal Age")), 0, color='red', zorder=2, s=50);
plots.ylim(-0.005, 0.07);
```

En el gráfico anterior, podemos observar que la media representa el centro de gravedad o el punto de equilibrio de la distribución. Si imaginamos que la distribución está recortada en cartón y la apoyamos sobre un dedo en la posición de la media, esta se mantendría perfectamente equilibrada. Dado que la distribución mostrada está sesgada hacia la derecha (es decir, tiene una cola larga en ese sentido), podemos deducir que la media es mayor que la mediana, la cual representa el punto "medio" de los datos. Por el contrario, si la distribución estuviera sesgada hacia la izquierda, sabríamos que la media sería menor que la mediana.

**Pregunta 1.2.** Ejecute la siguiente celda para comparar la media (rojo) y la mediana (azul) de la distribución de edades de las madres.

```python
births.hist("Maternal Age")
plots.vlines(np.mean(births.column("Maternal Age")), ymin=0, ymax=9 ,color="red", label="Mean")
plots.vlines(np.median(births.column("Maternal Age")), ymin=0, ymax=9,color="blue",label="Median")
plots.legend()
plots.ylim(-0.005, 0.07);
```

También nos interesa la desviación estándar de las edades de la madre. La SD nos da una idea de cuán variables son las edades de las madres en torno a la edad promedio de las madres. Si la SD es grande, entonces las edades de las madres deberían extenderse en un rango amplio con respecto a la media. Si la SD es pequeña, entonces las edades de las madres deberían agruparse estrechamente en torno a la edad materna promedio.

**La SD de una arreglo se define como la "media cuadrática de las desviaciones", donde las desviaciones son la diferencia entre los elementos originales y su promedio.**

Puede leer más sobre Variabilidad y SD en el [Capítulo 14.2](https://inferentialthinking.com/chapters/14/2/Variability.html). Este capítulo también profundiza un poco más en cómo se calcula la "media cuadrática de las desviaciones". 

Comenatario: $\sigma$ (letra griega sigma) se usa para representar la SD y $\mu$ (letra griega mu) se usa para la media.

**Pregunta 1.3.** Asigne `age_mean` a la media de las edades maternas y `age_sd` a la SD de las edades maternas. Luego ejecute la celda a continuación para ver el ancho de una SD (azul) de la media de la muestra (rojo) trazada en el histograma de edades maternas.

*Pista*: `np.std` podría resultar útil aquí.

```python
age_mean=...
age_sd=...

births.hist("Maternal Age")

plots.scatter(age_mean, 0, color='red', zorder=2, s=50);
plots.scatter(age_mean+age_sd, 0, marker='^', color='blue', zorder=2, s=50);
plots.scatter(age_mean-age_sd, 0, marker='^', color='blue', zorder=2, s=50);
plots.ylim(-0.005, 0.07);

```

```python
grader.check("q2")
```

En este histograma, la desviación estándar no es fácil de identificar con solo mirar el gráfico.

Sin embargo, las distribuciones de algunas variables nos permiten detectar fácilmente la desviación estándar en el gráfico. Por ejemplo, si una muestra sigue una **[distribución normal](https://inferentialthinking.com/chapters/14/3/SD_and_the_Normal_Curve.html)**, la desviación estándar se detecta fácilmente en el punto de inflexión (el punto donde la curva comienza a cambiar la dirección de su curvatura) de la distribución.

**Pregunta 1.4.** Complete el siguiente código para examinar la distribución de las alturas maternas, que se distribuye aproximadamente normalmente. Graficaremos la desviación estándar en el histograma, como antes; observe dónde cae en el gráfico una desviación estándar (azul) de la media (rojo).


```python
height_mean = ...
height_sd = ...
births.hist("Maternal Height", bins=np.arange(55,75,1))

plots.scatter((height_mean), 0, color='red', zorder=2, s=50);
plots.scatter(height_mean+height_sd, 0, marker='^', color='blue', zorder=2, s=50);
plots.scatter(height_mean-height_sd, 0, marker='^', color='blue', zorder=2, s=50);
plots.ylim(-0.01, 0.16)
```

```python
grader.check("q3")
```

No siempre sabemos cómo se distribuirá una variable, y hacer suposiciones sobre si una variable seguirá o no una distribución normal es peligroso. Sin embargo, el **Teorema del límite central** define una distribución que siempre sigue una distribución normal. **La distribución de las sumas o medias de los valores dentro de grandes muestras aleatorias, extraídas con reemplazo de una única distribución, será normal, sin importar la forma original de dicha distribución.** Recuerde que el teorema del límite central se refiere a la distribución de una *estadístico* calculado a partir de una distribución, no de la distribución de la muestra o población original. Reflexiona sobre esto, y discute con el profesor del laboratorio.

La siguiente sección explorará las distribuciones de medias muestrales y verá cómo la desviación estándar de estas distribuciones depende del tamaño de las muestras.

# 2. Variabilidad de la media muestral

Según el [Teorema del límite central](https://www.inferentialthinking.com/chapters/14/4/Central_Limit_Theorem.html), la distribución de probabilidad de la media o suma de una muestra aleatoria grande es aproximadamente normal. La curva de campana está centrada en la media poblacional. Algunas de las medias muestrales son más altas y otras más bajas, pero las desviaciones de la media poblacional son aproximadamente simétricas en ambos lados, como hemos visto repetidamente.

En nuestras simulaciones, también notamos que las medias de muestras más grandes tienden a agruparse más estrechamente alrededor de la media poblacional que las medias de muestras más pequeñas. En esta sección, cuantificaremos la [variabilidad de la media muestral](https://www.inferentialthinking.com/chapters/14/5/Variability_of_the_Sample_Mean.html) y desarrollaremos una relación entre la variabilidad y el tamaño de la muestra.

Echemos un vistazo a los salarios de los empleados de la ciudad de San Francisco en 2014. El salario medio informado por el gobierno de la ciudad fue de aproximadamente $75,463.92.

*Nota: Si se queda atascado en alguna parte de esta práctica de laboratorio, consulte el [Capítulo 14 del libro de texto](https://www.inferentialthinking.com/chapters/14/Why_the_Mean_Matters.html).*

```python
salaries = Table.read_table('sf_salaries_2014.csv').select("salary")
salaries.show(5)
```

```python
salary_mean = np.mean(salaries.column('salary'))
print('Salario promedio de los empleados de la ciudad de San Francisco en 2014:',\
      salary_mean.round(2))
```

```python
salaries.hist('salary', bins=np.arange(0, 250000+10000*2, 10000))
plots.scatter(salary_mean, 0, color='red', zorder=2, s=50)
plots.title('Salario en 2014 de los empleados de la ciudad de SF')
plots.ylim(-0.0000005, 0.000014);
```

Claramente, la población *no* sigue una distribución normal. Tenga esto en cuenta a medida que avancemos en estos ejercicios.

Tomemos muestras aleatorias **con reemplazo** y observemos la distribución de probabilidad de la media muestral. Como es habitual, utilizaremos la simulación para obtener una aproximación empírica a esta distribución.

**Pregunta 2.1**  Define una función llamada `one_sample_mean` con los siguientes argumentos:  

- `tabla`: el nombre de una tabla.  
- `etiqueta`: el nombre de la columna que contiene la variable de interés.  
- `tamano_muestra`: el número de empleados en la muestra.  

La función debe extraer una muestra de la tabla con reemplazo y devolver la media de los valores en la columna especificada por la `etiqueta`.  

**Nota:** La función debe ser genérica, haciendo referencia a los argumentos, para que pueda trabajar con cualquier tabla.  



Ejemplo de ejecución:**  

```python
# Supongamos que tenemos una tabla llamada datos
one_sample_mean(datos, "Salario", 3)


```python
def one_sample_mean(table, label, sample_size):
    new_sample = ...
    new_sample_mean = ...
    ...
```

```python
grader.check("q4")
```

**Pregunta 2.2**  
Usa la función `one_sample_mean` para definir `simulate_sample_mean` con los argumentos:  

- `tabla`: nombre de la tabla.  
- `etiqueta`: nombre de la columna que contiene la variable de interés.  
- `tamano_muestra`: número de empleados en cada muestra.  
- `repeticiones`: cantidad de veces que se repetirá la simulación.  

La función debe :  
1. Extraer muestras con reemplazo de la tabla.  
2. Calcular la media de cada muestra.  
3. Almacenar las medias en una lista llamada `means`.  
4. Generar un histograma empírico de las medias muestrales.  

**Nota:** Todos los valores se redondean a 2 decimales, ya que trabajamos con salarios.


```python

means = simulate_sample_mean(salaries, 'salary', 100, 10000)
Table().with_column('Medias de Muestra', means).hist()
plots.xlim(50000, 100000)
```

```python
grader.check("q5")
```

**Verifique** con un compañero o  el profesor que la **función** este implementada correctamente. Si no lo ha implementado correctamente, el resto del laboratorio no funcionará, por lo que este paso es crucial.

En la siguiente celda, crearemos una muestra de tamaño 100 a partir de "salarios" y la graficaremos usando nuestra nueva función "simulate_sample_mean".

  *Pista: Deberías ver una distribución similar a algo de lo que hemos estado hablando. Si no, verifique su función.*

```python
means = simulate_sample_mean(salaries, 'salary', 100, 10000)
Table().with_column('Medias de Muestra', means).hist()
plots.xlim(50000, 100000);
```

¡Observa que nuestra distribución de medias muestrales es aproximadamente normal! ¿Importaba si nuestra distribución original seguía una distribución normal?

**Pregunta 2.3.** Simule dos medias muestrales, una para una muestra de 400 salarios y otra para una muestra de 625 salarios. En cada caso, realiza 10.000 repeticiones. No se preocupe por las líneas `plots.xlim` o `plots.ylim`; solo se aseguran de que todos los gráficos tengan el mismo eje x e y, respectivamente.

```python
simulate_sample_mean(..., ..., ..., ...)
plots.xlim(50000, 100000)
plots.ylim(0, 0.0002)
plots.show();
print('\n')
simulate_sample_mean(..., ..., ..., ...)
plots.xlim(50000, 100000)
plots.ylim(0, 0.0002)
plots.show();
```

```python
grader.check("q6")
```

**Pregunta 2.4**

Asigna a la variable `q2_4` a una matriz de números que correspondan a **las afirmaciones** verdaderas sobre las gráficas de la sección 3.3.  

**Pista:** Si una distribución tiene forma de campana, ¿qué tipo de curva sigue?  

1. Observamos el Teorema del Límite Central (CLT) en acción, ya que las distribuciones de las medias muestrales tienen forma de campana.  
2. Observamos la Ley de los Promedios en acción, ya que las distribuciones de las medias muestrales se asemejan a la distribución de la población.  
3. Una de las condiciones para el CLT es extraer una pequeña muestra aleatoria con reemplazo de la población.  
4. Una de las condiciones para el CLT es extraer una muestra aleatoria grande con reemplazo de la población.  
5. Una de las condiciones para el CLT es que la población debe seguir una distribución normal.  


```python
q2_4 = ...
```

```python
grader.check("q7")
```

**Pregunta 2.5**

Asigna `q2_5` a una matriz de números que correspondan a las afirmaciones verdaderas sobre las gráficas de la sección 2.3.  

1. Ambos gráficos en 2.3 están aproximadamente centrados en la media poblacional.  
2. Ambos gráficos en 2.3 están aproximadamente centrados en la media de una muestra específica.  
3. La distribución de medias muestrales para un tamaño de muestra de 625 tiene menos variabilidad que la distribución de medias muestrales para un tamaño de muestra de 400.  
4. La distribución de medias muestrales para un tamaño de muestra de 625 tiene más variabilidad que la distribución de medias muestrales para un tamaño de muestra de 400.  


```python
q2_5 = ...
```

```python
grader.check("q8")
```

A continuación, veremos lo que sucede cuando tomamos un **número creciente de remuestreos de un tamaño de muestra fijo**. Observe qué número en el código cambia y qué permanece igual. ¿Cómo cambia la distribución de las medias remuestreadas?

```python
means_0=simulate_sample_mean(salaries, 'salary', 100, 500)
Table().with_column('Medias de Muestra', means_0).hist()
plots.xlim(50000, 100000)
plots.ylim(0, 0.00009);
```

```python
means_1=simulate_sample_mean(salaries, 'salary', 100, 1000)
Table().with_column('Medias de Muestra', means_1).hist()
plots.xlim(50000, 100000);
plots.ylim(0, 0.00009);
```

```python
means_2=simulate_sample_mean(salaries, 'salary', 100, 5000)
Table().with_column('Medias de Muestra', means_2).hist()

plots.xlim(50000, 100000);
plots.ylim(0, 0.00009);
```

```python
means_3=simulate_sample_mean(salaries, 'salary', 100, 10000)
Table().with_column('Medias de Muestra', means_3).hist()
plots.xlim(50000, 100000);
plots.ylim(0, 0.00009);
```

¿Qué notaste acerca de las distribuciones de medias muestrales en los cuatro histogramas anteriores? Discuta con sus compañeros o pregúntele a un miembro del personal.

**Pregunta 2.6**

Asigna la variable `SD_of_sample_means` al número entero correspondiente a la respuesta correcta para la siguiente pregunta:  

Cuando aumento el número de *nuevas muestras*, manteniendo un tamaño de muestra *fijo*, la desviación estándar (SD) de las medias muestrales:  

1. Aumentará.  
2. Disminuirá.  
3. Se mantendrá igual.  
4. Variará ampliamente.  


```python
SD_of_sample_means=...
```

```python
grader.check("q9")
```

**Pregunta 2.7.** Pensemos en cómo cambian las relaciones entre la SD de la población, la SD de la muestra y la SD de las medias de la muestra al variar el tamaño de la muestra. ¿Cual de los siguientes enunciados  es verdadero? Asigne la variable `pop_vs_sample` a una matriz de números enteros que correspondan a declaraciones verdaderas.

*Pista 1:* La SD de la muestra es diferente de la SD de las medias de la muestra.

*Pista 2:* [Capítulo 14.5](https://inferentialthinking.com/chapters/14/5/Variability_of_the_Sample_Mean.html) podría ser útil para responder esta pregunta.

1. La SD de la muestra se hace más pequeña a medida que aumenta el tamaño de la muestra.
2. La SD de la muestra aumenta a medida que aumenta el tamaño de la muestra.
3. La SD de la muestra se vuelve más consistente con la SD de la población a medida que aumenta el tamaño de la muestra.
4. La SD de las medias muestrales se reduce al aumentar el tamaño de la muestra.
5. La SD de las medias muestrales aumenta a medida que aumenta el tamaño de la muestra.
6. La SD de las medias muestrales permanece igual al aumentar el tamaño de la muestra.

```python
pop_vs_sample=....
```

```python
grader.check("q10")
```

**Pregunta 2.8.** ¿Existe una relación entre el tamaño de la muestra y la desviación estándar de las medias muestrales? Asigne `q2_8` al número correspondiente a la declaración que responde a esta pregunta.

*Pista:* [Capítulo 14.5](https://inferentialthinking.com/chapters/14/5/Variability_of_the_Sample_Mean.html) del libro de texto puede resultar útil.

1. La SD de las medias muestrales es inversamente proporcional a la raíz cuadrada del tamaño de la muestra.
2. La SD de las medias muestrales es directamente proporcional a la raíz cuadrada del tamaño de la muestra.

```python
q2_8=...

```

```python
grader.check("q11")
```

A lo largo de este laboratorio, hemos estado tomando muchas muestras aleatorias de una población. Sin embargo, todos estos principios son válidos para remuestreos con arranque de una sola muestra. El bootstrap funciona porque es como extraer datos de la población original, asumiendo que la muestra es representativa. Si su muestra original es relativamente grande, todas sus remuestras también serán relativamente grandes, por lo que la SD de las medias remuestreadas será relativamente pequeña.

**Para cambiar la variabilidad de la media de su muestra, tendría que cambiar el tamaño de la muestra original de la cual está tomando remuestras con arranque.**

-----
## ¡Hecho!


<img src="./Zaira.jpg" width="30%" alt="Imagen en primer plano de un perro sonriendo, sentado en el césped."/>

**Zaira** quiere recordarte que ¡Puedes hacerlo!

¡Felicitaciones por completar la Laboratorio 8!

¡Terminaste con el laboratorio!

**Información importante sobre el envío:**
- **Ejecuta todas las pruebas** y verifica que todas pasan
- **Guardar** desde el menú **Archivo**
- **Ejecute la celda final **
- Luego, ve a este [enlace](https://cienciadatosudea.github.io/001_SitioWebFundCienciaDatos/04entregas.html) y envía el archivo ipynb a la tarea correspondiente. El nombre de esta tarea es "Lab XX Autograder", donde XX es el número de laboratorio: 01, 02, 03, etc.

- Si terminas temprano en el laboratorio regular, **registra tu asistencia con el instructor**.

**Es su responsabilidad asegurarse de que su trabajo esté guardado antes de ejecutar la última celda.**

`Referencia`: Este material esta basado en [Data 8: The Foundations of Data Science](https://www.data8.org/) - University of California, Berkeley.


## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

These are some submission instructions.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q10 (points: 1)
```
>>> assert (pop_vs_sample, make_array(3, 4))
```


Tests for q11 (points: 1)
```
>>> assert q2_8 == 1
```


Tests for q2 (points: 1)
```
>>> assert np.isclose(age_mean, 27.228, 0.01)
>>> assert np.isclose(age_sd, 5.815, 0.01)
```


Tests for q3 (points: 1)
```
>>> assert np.isclose(height_mean, 64.049, 0.001)
>>> assert np.isclose(height_sd, 2.525, 0.001)
```


Tests for q4 (points: 1)
```
>>> baby = Table.read_table('baby.csv')
>>> resultado = one_sample_mean(baby, 'Birth Weight', 100)
>>> assert isinstance(resultado, float), 'El resultado debe ser un valor float.'
>>> assert np.isclose(abs(resultado), 120, 5 * 16), 'El resultado debe ser un valor float.'
```


Tests for q5 (points: 1)
```
>>> u = np.mean(simulate_sample_mean(salaries, 'salary', 1000, 1000))
>>> assert np.isclose(u, 75497, 3000)
```


Tests for q6 (points: 1)
```
>>> u = simulate_sample_mean(salaries, 'salary', 400, 10000)
>>> v = simulate_sample_mean(salaries, 'salary', 625, 10000)
>>> assert np.isclose(np.mean(u), 75461.9, 60)
>>> assert np.isclose(np.mean(v), 75480, 60)
```


Tests for q7 (points: 1)
```
>>> assert (q2_4 == make_array(1, 4)).sum() == 2
```


Tests for q8 (points: 1)
```
>>> assert (q2_5 == make_array(1, 3)).sum() == 2
```


Tests for q9 (points: 1)
```
>>> assert SD_of_sample_means == 3
```


Tests for qt (points: 0)
```
>>> assert isinstance(NombreApellidos, str)
>>> assert isinstance(NumeroCarnet, str)
>>> assert isinstance(Email, str)
>>> assert len(NombreApellidos) > 3
>>> assert len(NumeroCarnet) > 3
>>> assert len(Email) > 3
```
