```python
# Initialize Otter
import otter
grader = otter.Notebook("proyecto_01.ipynb")
```

<p><img alt="Colaborat" style="height: 220px;" src="https://github.com/CienciaDatosUdea/001_DataSetFundCienciaDatos/blob/main/semestre2024-2/lab/lab02/Escudo-UdeA.svg.png?raw=true" align="left" hspace="10px" vspace="0px"></p>


# Fundamentos en ciencias de datos


## Proyecto 01: Población Mundial y Pobreza


En este proyecto, explorarás datos de [Gapminder.org](http://gapminder.org), un sitio web dedicado a ofrecer una visión del mundo basada en hechos y cómo ha cambiado. Ese sitio incluye varias visualizaciones y presentaciones de datos, pero también publica los datos en bruto que utilizaremos en este proyecto para recrear y ampliar algunas de sus visualizaciones más famosas. El sitio web de Gapminder recopila datos de muchas fuentes y los compila en tablas que describen muchos países alrededor del mundo. Todos los datos que agregan se publican en el [Systema Globalis](https://github.com/open-numbers/ddf--gapminder--systema_globalis/blob/master/README.md). Su objetivo es compilar todas las estadísticas públicas; sociales, económicas y ambientales; en un conjunto de datos . Todos los conjuntos de datos de este proyecto se copian directamente del Systema Globalis sin ningún cambio. Este proyecto está dedicado a [Hans Rosling](https://en.wikipedia.org/wiki/Hans_Rosling) (1948-2017), quien promovió el uso de datos para entender y priorizar los desafíos del desarrollo global.

### Logística **Fecha límite.** 



**Punto de control.** Para obtener la nota  completa en el punto de control, debes completar las preguntas hasta el punto de control, **aprobar todas las pruebas de autocalificación.** 

**Compañeros.** Puedes trabajar con un compañero; tu compañero debe ser de tu sección de laboratorio asignada.  

**Reglas.** No compartas tu código con nadie más que tu compañero. Puedes discutir las preguntas con otros estudiantes, pero no compartas las respuestas. La experiencia de resolver los problemas en este proyecto te preparará para los exámenes (y la vida). Si alguien te pide la respuesta, ¡resiste! En su lugar, puedes demostrar cómo resolverías un problema similar.

**Pruebas.** <span style="color: #BC412B">Las pruebas que se proporcionan **no son exhaustivas** y pasar las pruebas de una pregunta **no** significa que respondiste la pregunta correctamente.</span> Las pruebas generalmente solo verifican que tu tabla tenga las etiquetas de columnas correctas. Sin embargo, se aplicarán más pruebas para verificar la exactitud de tu envío con el fin de asignar tu puntuación final, así que ten cuidado y verifica tu trabajo. Podrías querer crear tus propias verificaciones a lo largo del camino para ver si tus respuestas tienen sentido. Además, antes de enviar, asegúrate de que ninguna de tus celdas tome mucho tiempo en ejecutarse (varios minutos).

**Preguntas de respuesta libre:** Asegúrate de poner las respuestas a las preguntas escritas en la celda indicada que proporcionamos. **Cada pregunta de respuesta libre debe incluir una explicación** que responda adecuadamente a la pregunta.

**Consejo.** Desarrolla tus respuestas de manera incremental. Para realizar una manipulación de tabla complicada, divídela en pasos, realiza cada paso en una línea diferente, da un nuevo nombre a cada resultado y verifica que cada resultado intermedio sea lo que esperas. Puedes agregar los nombres o funciones adicionales que desees en las celdas proporcionadas. Asegúrate de estar usando nombres de variables distintos y significativos en todo el cuaderno. En esa línea, **NO** reutilices los nombres de variables que usamos cuando calificamos tus respuestas. Por ejemplo, en la Pregunta 1 de la sección de Pobreza Global te pedimos que asignes una respuesta a `b_pop`. No reasignes el nombre de variable `b_pop` a nada más en tu cuaderno, de lo contrario existen posibilidades de que nuestras pruebas califiquen contra lo que `b_pop` fue reasignado. **Nunca** estás restringido a utilizar solo una línea de código para resolver una pregunta en este proyecto o en cualquier otro. ¡Siéntete libre de usar variables intermedias y múltiples líneas tanto como desees!

Para comenzar, carga `datascience`, `numpy`, `plots` y `otter`.

```python
# Run this cell to set up the notebook, but please don't change it.
# These lines import the NumPy and Datascience modules.
from datascience import *
import numpy as np
# These lines do some fancy plotting magic.
%matplotlib inline
import matplotlib.pyplot as plots
plots.style.use('fivethirtyeight')
from ipywidgets import interact, interactive, fixed, interact_manual
import ipywidgets as widgets

```

_Type your answer here, replacing this text._

```python
Nombre = "..."
Carnet = "..."
Grupo = "..."
```

<!-- END QUESTION -->

## 1. Crecimiento de la Población Mundial

La población mundial de seres humanos alcanzó los 1000 millones alrededor del año 1800; 3000 millones alrededor de 1960, y 7000 millones alrededor de 2011. El impacto potencial del crecimiento exponencial de la población ha preocupado a científicos, economistas y políticos por igual. La División de Población de las Naciones Unidas, estima que es probable que la población mundial continúe creciendo a lo largo del siglo XXI, pero a un ritmo más lento, quizás alcanzando los 11000 millones para el año 2100. Sin embargo, la ONU no descarta escenarios de un crecimiento más extremo.<a href="http://www.pewresearch.org/fact-tank/2015/06/08/scientists-more-worried-than-public-about-worlds-growing-population/ft_15-06-04_popcount/">  <img src="pew_population_projection.png"/> </a>

En esta parte del proyecto, examinaremos algunos de los factores que influyen en el crecimiento de la población y cómo han ido cambiando a lo largo de los años y en todo el mundo. Hay dos subpartes principales en este análisis.

- Primero, examinaremos los datos de un país, Bangladesh. Veremos cómo han cambiado factores como la esperanza de vida, la tasa de fertilidad y la mortalidad infantil a lo largo del tiempo en Bangladesh, y cómo se relacionan con la tasa de crecimiento de la población.

- A continuación, examinaremos si los cambios que hemos observado para Bangladesh son particulares de ese país o si reflejan patrones generales observables en otros países también. Estudiaremos aspectos del crecimiento de la población mundial y veremos cómo han estado cambiando.La primera tabla que consideraremos contiene la población total de cada país a lo largo del tiempo. Ejecuta la celda a continuación.

```python
population = Table.read_table('population.csv').where("time", are.below(2021))
population.show(3)
```

Output:
```
<IPython.core.display.HTML object>
```

**Nota:** El archivo csv de población también se puede encontrar [aquí](https://github.com/open-numbers/ddf--gapminder--gapminder_world/blob/master/ddf--datapoints--population_total--by--geo--time.csv). Los datos para este proyecto fueron descargados en febrero de 2017.

### Bangladesh 

La nación de [Bangladesh](https://es.wikipedia.org/wiki/Bangladesh) se estableció como una democracia parlamentaria después de que terminara la Guerra de Liberación de Bangladesh en 1971. La nación naciente, devastada por la guerra, se enfrentó casi de inmediato a inundaciones y hambrunas.  En esta sección del proyecto, examinaremos aspectos del desarrollo de Bangladesh desde ese momento. En la tabla `población`, la columna `geo` contiene códigos de tres letras establecidos por la [Organización Internacional de Normalización](https://es.wikipedia.org/wiki/Organizaci%C3%B3n_Internacional_de_Normalizaci%C3%B3n) (ISO) en el estándar [Alpha-3](https://en.wikipedia.org/wiki/ISO_3166-1_alpha-3). Comenzaremos examinando de cerca a Bangladesh. Utiliza el enlace Alpha-3 para encontrar el código de tres letras para Bangladesh.

**Pregunta 1.** Crea una tabla llamada `b_pop` que tenga dos columnas etiquetadas como `time` y `population_total` desde la información dada en la tabla `population.csv`. La primera columna debe contener los años desde 1970 hasta 2020 (incluyendo tanto 1970 como 2020) y la segunda debe contener la población de Bangladesh en cada uno de esos años.

```python
b_pop = ...
b_pop
```

Output:
```
Ellipsis
```

```python
grader.check("q1")
```

Output:
```
q1 results:
    q1 - 1 result:
        ❌ Test case failed
        Trying:
            import pandas as pd
        Expecting nothing
        ok
        Trying:
            assert b_pop.to_df().time.mean() == 1995
        Expecting nothing
        **********************************************************************
        Line 2, in q1 0
        Failed example:
            assert b_pop.to_df().time.mean() == 1995
        Exception raised:
            Traceback (most recent call last):
              File "/home/gustavo/anaconda3/envs/cienciaDeDatos/lib/python3.11/doctest.py", line 1355, in __run
                exec(compile(example.source, filename, "single",
              File "<doctest q1 0[1]>", line 1, in <module>
                assert b_pop.to_df().time.mean() == 1995
                       ^^^^^^^^^^^
            AttributeError: 'ellipsis' object has no attribute 'to_df'
        Trying:
            assert np.isclose(b_pop.to_df().population_total.mean(), 114373388.50980392, 0.00
```

Ejecuta la siguiente celda para crear una tabla llamada `b_five` que contenga la población de Bangladesh cada cinco años. A simple vista, ¡parece que la población de Bangladesh ha estado creciendo rápidamente!

```python
b_pop.set_format('population_total', NumberFormatter)
fives = np.arange(1970, 2021, 5) # 1970, 1975, 1980, ...
b_five = b_pop.sort('time').where('time', are.contained_in(fives))
b_five.show()
```

Output:
```
AttributeError                            Traceback (most recent call last)
Cell In[7], line 1
----> 1 b_pop.set_format('population_total', NumberFormatter)
      2 fives = np.arange(1970, 2021, 5) # 1970, 1975, 1980, ...
      3 b_five = b_pop.sort('time').where('time', are.contained_in(fives))

AttributeError: 'ellipsis' object has no attribute 'set_format'
```


**Pregunta 2.** Asigna `initial` a un arreglo que contenga la población para cada intervalo de cinco años desde 1970 hasta 2015 (inclusive). Luego, asigna `changed` a un arreglo que contenga la población para cada intervalo de cinco años desde 1975 hasta 2020 (inclusive). El primer arreglo debe incluir tanto 1970 como 2015, y el segundo arreglo debe incluir tanto 1975 como 2020. Debes usar la tabla `b_five` para crear ambos arreglos, filtrando primero la tabla para que solo contenga los años relevantes. La tasa de crecimiento anual para un período de tiempo es igual a:$$\left(\left(\frac{\text{Población al final del período}}{\text{Población al inicio del período}}\right)^{\displaystyle\frac{1}{\text{número de años}}}\right) -1$$ 

Hemos proporcionado el siguiente código, que utiliza los arreglos `initial` y `changed` para construir la tabla `b_five_growth`. Esta tabla incluye una nueva columna llamada `annual_growth`, junto con las columnas originales de la tabla `b_five`.

No te preocupes por el cálculo de las tasas de crecimiento; ejecuta la prueba a continuación para probar tu solución. Si estás interesado en cómo llegamos a la fórmula para las tasas de crecimiento, consulta la sección [tasas de crecimiento](https://inferentialthinking.com/chapters/03/2/1/Growth.html) del libro de texto.

```python
initial = ...
changed = ...

b_1970_through_2015 = b_five.where('time', are.below_or_equal_to(2015)) 
b_five_growth = b_1970_through_2015.with_column('annual_growth', (changed/initial)**0.2-1)
b_five_growth.set_format('annual_growth', PercentFormatter)

```

```python
grader.check("q2")
```

Mientras que la población ha crecido cada cinco años desde 1970, la tasa de crecimiento anual disminuyó drásticamente de 1985 a 2015. Veamos alguna otra información para desarrollar una posible explicación. Ejecuta la siguiente celda para cargar tres tablas adicionales de mediciones sobre países a lo largo del tiempo.

```python
life_expectancy = Table.read_table('life_expectancy.csv').where('time', are.below(2021))
child_mortality = Table.read_table('child_mortality.csv').relabel(2, 'child_mortality_under_5_per_1000_born').where('time', are.below(2021))
fertility = Table.read_table('fertility.csv').where('time', are.below(2021))

```

La tabla `life_expectancy` contiene una estadística que a menudo se utiliza para medir cuánto tiempo viven las personas, llamada *esperanza de vida al nacer*. Este número, para un país en un año dado, [no mide cuánto tiempo se espera que vivan los bebés nacidos ese año](http://blogs.worldbank.org/opendata/what-does-life-expectancy-birth-really-mean). En cambio, mide cuánto tiempo viviría alguien, en promedio, si las *condiciones de mortalidad* de ese año persistieran durante toda su vida. Estas "condiciones de mortalidad" describen qué fracción de personas de cada edad sobrevivió al año. Por lo tanto, es una forma de medir la proporción de personas que permanecen vivas, agregada en diferentes grupos de edad de la población.

Ejecuta las siguientes celdas a continuación para ver `esperanza_de_vida`, `mortalidad_infantil` y `fertilidad`. Refiérete a estas tablas ya que serán útiles para responder preguntas adicionales.

```python
life_expectancy.show(3)
```

```python
child_mortality.show(3)
```

```python
fertility.show(3)
```

**Pregunta 3.** Tal vez la población está creciendo más lentamente porque las personas no viven tanto tiempo. Utiliza la tabla `life_expectancy` para dibujar un gráfico de líneas con los años 1970 y posteriores en el eje horizontal que muestre cómo ha cambiado la *esperanza de vida al nacer* en Bangladesh.

```python
#Escribir el codigo 
...
```

```python
grader.check("q3")
```

**Pregunta 4.** Suponiendo que todo lo demás permanece igual, ¿los cambios en la esperanza de vida en el gráfico anterior explican directamente por qué la tasa de crecimiento de la población disminuyó de 1985 a 2015 en Bangladesh? ¿Por qué sí o por qué no?

Pista: Reflexiona sobre cómo los cambios en la esperanza de vida podrían o no estar relacionados con la tasa de crecimiento de la población. ¿Qué otros factores importantes podrían haber influido en esta disminución? Además, incorpora un análisis de los eventos significativos ocurridos en Bangladesh en 1991 y cómo pudieron haber impactado en la dinámica poblacional. Si utilizas información de fuentes externas, asegúrate de citarlas adecuadamente

```python
explicacion = "Escribe tu respuesta aqui"
```

La tabla `fertility` contiene una estadística que a menudo se utiliza para medir cuántos bebés están naciendo, la *tasa de fecundidad total*. Este número describe el [número de hijos que tendría una mujer a lo largo de su vida](https://www.measureevaluation.org/prh/rh_indicators/specific/fertility/total-fertility-rate), en promedio, si las tasas actuales de nacimiento por edad de la madre se mantuvieran durante todos sus años fértiles, asumiendo que sobreviva hasta los 49 años.

**Pregunta 5.** Completa la función `fertility_over_time`. Esta función recibe el código Alpha-3 de un país como `country_code` y un año de `start`. Devuelve una tabla de dos columnas con las etiquetas `Anio` y `Ninos por mujer` que puede ser utilizada para generar un gráfico de líneas de la tasa de fecundidad del país cada año, comenzando desde el año inicio `start`. La gráfica debe incluir el año `start` y todos los años posteriores que aparezcan en la tabla `fertility`. Luego, determina el código Alpha-3 para Bangladesh. El código para Bangladesh y el año 1970 se utilizan en la llamada a tu función `fertility_over_time` para graficar cómo ha cambiado la tasa de fecundidad de Bangladesh desde 1970. Ten en cuenta que la función `fertility_over_time` no debe devolver la gráfica en sí. **La expresión que dibuja el gráfico de líneas se proporciona para ti; por favor, no la cambies.**

```python
countries = Table.read_table('countries.csv')
fertility = Table.read_table("fertility.csv")

def fertility_over_time(country_code, start):
    country_fertility = ...
    country_fertility_after_start = ...
    cleaned_table = ...
    ...
    
bangladesh_code = ...

# No borre estas lineas de codigo 
fertility_1970_ = fertility_over_time(bangladesh_code, 1970)
fertility_over_time(bangladesh_code, 1970).plot(0, 1)
```

```python
grader.check("q5")
```

**Pregunta 6.** Asumiendo que todos los demás factores permanecen constantes, ¿las tendencias en fertilidad en el gráfico anterior ayudan a explicar directamente por qué la tasa de crecimiento de la población disminuyó de 1980 a 2020 en Bangladesh? ¿Por qué sí o por qué no?

```python
explicacion="Escribe tu respuesta aqui"
```

Se ha [observado](https://www.ageing.ox.ac.uk/download/143) que las tasas de fertilidad más bajas a menudo están asociadas con tasas de mortalidad infantil más bajas. Esta conexión se ha atribuido a la planificación familiar: si los padres pueden esperar que todos sus hijos sobrevivan hasta la adultez, entonces elegirán tener menos hijos. En la dirección inversa, tener menos hijos, puede permitir a las familias dedicar más recursos a cada niño, reduciendo así la mortalidad infantil. Podemos ver si esta asociación es evidente en Bangladesh al graficar la relación entre la tasa de fertilidad total y la [tasa de mortalidad infantil por cada 1000 niños](https://es.wikipedia.org/wiki/Mortalidad_infantil).

**Pregunta 7.** Usando tanto las tablas `fertility` como `child_mortality`, dibuja un diagrama de dispersión que tenga la fertilidad total de Bangladesh en el eje horizontal y su mortalidad infantil en el eje vertical, con un punto para cada año, comenzando en 1970.**Se proporciona el código que dibuja el diagrama de dispersión; por favor, no lo cambies.** En su lugar, crea una tabla llamada `post_1969_fertility_and_child_mortality` con las etiquetas de columna y los datos apropiados para generar el gráfico correctamente. Utiliza la etiqueta `Ninos_por_mujer` para describir la fertilidad total y la etiqueta `Muertes_infantiles_por_1000_nacidos` para describir la mortalidad infantil.

*Consejo*: No elimines la columna `time` o recibirás un error en el diagrama de dispersión en la siguiente celda.

```python
fertility=Table().read_table("fertility.csv")

bgd_fertility = ...
bgd_child_mortality = ...
fertility_and_child_mortality = ...
post_1969_fertility_and_child_mortality = ...

# No cambiar las siguiente lineas de codigo por favor
x_births = post_1969_fertility_and_child_mortality.column("Children per woman")
y_deaths = post_1969_fertility_and_child_mortality.column("Child deaths per 1000 born")
time_colors = post_1969_fertility_and_child_mortality.column("time")

plots.figure(figsize=(6,6))
plots.scatter(x_births, y_deaths, c=time_colors, cmap="Blues_r")
plots.colorbar()                  
plots.xlabel("Children per woman")
plots.ylabel("Child deaths per 1000 born");
```

```python
grader.check("q7")
```

El gráfico anterior utiliza **color** para codificar datos sobre la columna `time` de la tabla `post_1969_fertility_and_child_mortality`. Los colores, que van desde azul oscuro hasta blanco, representan el paso del tiempo desde los años 1970 hasta los años 2020. Por ejemplo, un punto en el diagrama de dispersión que representa datos de los años 1970 aparecería como **azul oscuro**, y un punto de los años 2010 aparecería como **azul claro**.

<!-- BEGIN QUESTION -->

**Pregunta 8.** En una o dos oraciones, describe la asociación (si la hay) que ilustra este diagrama de dispersión. ¿Muestra el diagrama que la reducción de la mortalidad infantil *causa* que los padres elijan tener menos hijos?

```python
descripcion="Escribe tu respuesta aqui"
```

<!-- END QUESTION -->

### Punto de control 

<div style="display: flex; justify-content: space-between; align-items: center;">
    <div style="flex: 1; text-align: center;">
        <img src="Rocky.jpeg" alt="Rocky el golden retriever" width="300"/>
        <p><b>Rocky: <strong>¡¡¡WOOOHOO!!!</strong> </b></p>
    </div>
    <div style="flex: 1; text-align: center;">
        <img src="Kiriku.jpg" alt="Kiriku el golden retriever" width="300"/>
        <p><b>Kiriku: <strong>¡¡¡MIIIAAAUU!!!</strong></b></p>
    </div>
    <div style="flex: 1; text-align: center;">
        <img src="Berlin.jpg" alt="Berlin el golden retriever" width="300"/>
        <p><b>Berlin: <strong>¡¡¡MIIIAAAUU!!!</strong></p></b></p>
    </div>
</div>

Berlin, Kiriku y Rocky quieren felicitarte por alcanzar el punto de control!

**Ejecuta las siguientes celdas y envía la tarea a la asignación en el correspondiente al punto de control: <u>Punto de control del Proyecto 1</u>****<span style="color: #BC412B">

Para verificar tu trabajo, la celda a continuación volverá a ejecutar todas las pruebas del autograder para la Sección 1.

```python
checkpoint_tests = ["q1", "q2", "q3","q5","q7"]

for test in checkpoint_tests:
    display(grader.check(test))
```

## Envío 

Asegúrate de haber ejecutado todas las celdas en tu cuaderno en orden antes de ejecutar la celda de abajo, de modo que todas las imágenes/gráficas aparezcan en el resultado. La celda de abajo generará un archivo zip para que lo envíes. 

**¡Por favor guarda antes de exportar! **

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

These are some submission instructions.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1 (points: 1)
```
>>> import pandas as pd
>>> assert b_pop.to_df().time.mean() == 1995
>>> assert np.isclose(b_pop.to_df().population_total.mean(), 114373388.50980392, 0.001)
>>> assert 'time' in b_pop.labels, "La columna 'time' no está presente en b_pop."
>>> assert 'population_total' in b_pop.labels, "La columna 'population_total' no está presente en b_pop."
```


Tests for q2 (points: 1)
```
>>> import pandas as pd
>>> assert np.isclose(b_five_growth.to_df().time.mean(), 1992.5, 0.001)
>>> assert np.isclose(b_five_growth.to_df().population_total.mean(), 109356945.1, 0.001)
>>> assert np.isclose(b_five_growth.to_df().annual_growth[0], 0.01753862351142521, 0.001)
>>> assert 'time' in b_five_growth.labels, "La columna 'time' no está presente en b_pop."
>>> assert 'population_total' in b_five_growth.labels, "La columna 'population_total' no está presente en b_pop."
```


Tests for q3 (points: 1)
```
>>> assert np.isclose(bangladesh_life_expectancy.to_df().time.mean(), 2035.0, 1e-06)
>>> assert np.isclose(bangladesh_life_expectancy.to_df().life_expectancy_years.mean(), 74.0656822, 1e-06)
>>> assert 'geo' in bangladesh_life_expectancy.labels, "La columna 'geo' no esta presente"
>>> assert 'time' in bangladesh_life_expectancy.labels, 'La columna time no esta presente'
```


Tests for q5 (points: 1)
```
>>> assert 'Anio' in fertility_1970_.labels, "La columna 'Anio' no esta presente"
>>> assert 'Ninos_por_mujer' in fertility_1970_.labels, 'La columna Ninos_por_mujer no esta presente'
>>> assert bangladesh_code == 'bgd'
>>> assert np.isclose(fertility_1970_.to_df().Anio.mean(), 2035.0, 0.1)
>>> assert np.isclose(fertility_1970_.to_df().Ninos_por_mujer.mean(), 2.6768708, 0.001)
```


Tests for q7 (points: 1)
```
>>> assert 'Ninos_por_mujer' in post_1969_fertility_and_child_mortality.labels, "La columna 'Ninos_por_mujer' no esta presente"
>>> assert 'Muertes_infantiles_por_1000_nacidos' in post_1969_fertility_and_child_mortality.labels, "La columna 'Muertes_infantiles_por_1000_nacidos' no esta presente"
>>> assert 'Anios' in post_1969_fertility_and_child_mortality, "La columna 'Anio' no esta presente"
>>> assert np.isclose(time_colors.mean(), 1995.0, 1e-06)
>>> assert np.isclose(y_deaths.mean(), 121.15352941176468, 0.01)
>>> assert np.isclose(x_births.mean(), 4.166274509803921, 1e-06)
```
