```python
# Initialize Otter
import otter
grader = otter.Notebook("lab08.ipynb")
```

**Recurso Útil:**

- [Referencia de Python](http://data8.org/fa23/reference/): Esta es una guía rápida que contiene métodos útiles para trabajar con arreglos y tablas en el curso de Data 8.

**Lectura Recomendada**:

* [Estimación](https://www.inferentialthinking.com/chapters/13/Estimation): Este capítulo abarca los conceptos fundamentales de la estimación en estadística, proporcionando las bases necesarias para comprender cómo, a partir de una muestra, podemos realizar inferencias acerca de una población.

<p><img alt="Colaborat" style="height: 220px;" src="https://github.com/CienciaDatosUdea/001_DataSetFundCienciaDatos/blob/main/semestre2024-2/lab/lab02/Escudo-UdeA.svg.png?raw=true" align="left" hspace="10px" vspace="0px"></p>

# Fundamentos en Ciencias de Datos

## Laboratorio 8: Intervalos de Confianza

¡Bienvenidos al Laboratorio 8!

En este laboratorio, aprenderemos y aplicaremos técnicas de intervalos de confianza para estimar parámetros poblacionales, utilizando ejemplos concretos que nos permitirán comprender mejor su aplicación práctica. Imagina que estás evaluando la popularidad de restaurantes de comida paisa en tu ciudad. 

```python
# Don't change this cell; just run it. 

import numpy as np
from datascience import *

# These lines do some fancy plotting magic.",
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt 
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore', FutureWarning)
```

<!-- BEGIN QUESTION -->

Ingresa tu nombre  con apellido y número de carnet en las variables dadas en la celda siguiente:
Ejemplo:
```python
NombreApellidos="Ada Lovelace"
NumeroCarnet="27182818"
Email="Ada.Lovelace@udea.edu.co"
```

```python
NombreApellidos="..."
NumeroCarnet=".."
Email="..."
```

```python
grader.check("qt")
```

<!-- END QUESTION -->

## 1. Encuesta sobre el Mejor Restaurante de Frijoles en Medellín

Estamos llevando a cabo una encuesta para descubrir cuál es el restaurante de frijoles favorito en Medellín. Los resultados actuales son los siguientes:

- **Don Frijolón**: 8%
- **Los Frijolitos**: 53%
- **Sabor a Frijol**: 25%
- **El Fogón del Frijol**: 14%

Tu opinión es importante para nosotros. Comparte tu experiencia y ayúdanos a encontrar el mejor lugar para disfrutar de unos deliciosos frijoles en la ciudad.

# Encuesta sobre el Mejor Restaurante de Frijoles en Medellín

Jessica y Ciara están tratando de identificar cuál es el restaurante de frijoles más popular en Medellín. Han llevado a cabo una encuesta en la que participaron **1,500** habitantes de la ciudad, seleccionados al azar de manera uniforme. A los encuestados se les pidió que escogieran el restaurante que, a su juicio, sirve los mejores frijoles.

Las opciones de restaurantes que se les presentaron fueron:

- **Don Frijolón** 
- **Los Frijolitos**
- **Sabor a Frijol**
- **El Fogón del Frijol**

Tras analizar las respuestas, Jessica y Ciara publicaron los siguientes porcentajes que reflejan el resultado de la encuesta:

| Restaurante de Frijoles | Porcentaje |
|:----------------------:|:----------:|
| Don Frijolón | 8% |
| Los Frijolitos | 53% |
| Sabor a Frijol | 25% |
| El Fogón del Frijol | 14% |

Estos datos proporcionan una visión general de las preferencias de los participantes sobre en qué restaurante se pueden encontrar los mejores frijoles de Medellín. Cabe destacar que este es un experimento hipotético y tanto los resultados como los nombres de los restaurantes son ficticios.

```python
# Just run this cell
votes = Table.read_table('Votos_frijoles.csv')
votes
```

**Pregunta 1.1.** Completa la función `one_resampled_percentage` a continuación. Debe devolver el ***porcentaje*** de votos de Los Frijolitos después de tomar la tabla original (`tbl`) y realizar una muestra bootstrap de ella. Recuerda que un porcentaje está entre 0 y 100. **(8 Puntos)**

*Nota 1:* `tbl` siempre tendrá el mismo formato que `votes`.

*Nota 2:* Esta función debe completarse **sin** `.group` o `.pivot`. Usar estas funciones causará que tu código se agote en el tiempo de ejecución.

*Pista:* Dada una tabla de votos, ¿cómo puedes determinar qué porcentaje de los votos son para un cierto restaurante? ¡Asegúrate de usar porcentajes, no proporciones, para esta pregunta!

```python
def one_resampled_percentage(tbl):
    # Escribe la solucion acá
    pass

one_resampled_percentage(votes)
one_resampled_percentage(votes)
```

```python
grader.check("q0")
```

**Pregunta 1.2.** Completa la función `percentages_in_resamples` de manera que simule y devuelva un arreglo de **2023 elementos**, donde cada elemento representa una estimación remuestreada del porcentaje de votantes que votarán por Don Frijolón. Debes utilizar la función `one_resampled_percentage` que escribiste anteriormente. **(8 Puntos)**, la función debe retornar un arreglo de numpy.

_Nota:_ Realizamos nuestra simulación con **solo 2023 pruebas** en este problema para reducir el tiempo de ejecución, pero generalmente deberíamos usar más repeticiones.

```python

def percentages_in_resamples():
    resampled_percentages = make_array()
    pass

percentages_in_resamples()
```

```python
grader.check("q1")
```

En la celda siguiente, ejecutaremos la función `percentages_in_resamples` que acabas de definir. Esta función nos permitirá crear un histograma del estadístico calculado para las estimaciones bootstrap del año 2023, específicamente del porcentaje de votantes que eligieron a Don Frijolón. 

*Nota:* La ejecución de este proceso puede tardar unos segundos.

```python
resampled_percentages = percentages_in_resamples()
Table().with_column('Estimated Percentage', resampled_percentages).hist("Estimated Percentage")
```

**Pregunta 1.3.** En esta laboratorio, vamos a trabajar con el arreglo `resampled_percentages` para identificar los valores que se encuentran en los extremos de un intervalo que abarca el 95% de las estimaciones porcentuales re-muestreadas. El objetivo es calcular tanto el límite inferior como el superior de este intervalo de confianza, y asignar estos valores a las variables `lower_bound` e `upper_bound`, respectivamente. **(8 Puntos)**, define un array llamado interval con el limte inferior y superior.

Toma en cuenta que este proceso es fundamental para entender el intervalo de confianza de nuestras estimaciones, lo cual es un concepto clave en el análisis inferencial de datos. Asegúrate de calcular correctamente estos extremos para obtener resultados precisos.

*Sugerencia:* Si encuentras dificultades al abordar esta pregunta, considera revisar el [Capítulo 13.1](https://inferentialthinking.com/chapters/13/1/Percentiles.html) del libro de texto, donde se exploran los percentiles en mayor detalle.

```python
lower_bound = ...
upper_bound = ...
interval=make_array(lower_bound, upper_bound)
print(f"Intervalo de confianza del 95% para el porcentaje de votos de Los Frijolitos en la población: [{lower_bound:.2f}, {upper_bound:.2f}]")

```

```python
grader.check("q2")
```

**Pregunta 1.4.** Los resultados de la encuesta parecen indicar que Los Frijolitos está superando a todos los demás restaurantes de comida paisa entre los votantes. Nos gustaría utilizar intervalos de confianza para determinar un rango de valores probables para la verdadera ventaja de Los Frijolitos sobre todos los demás restaurantes combinados. El cálculo para la ventaja de Los Frijolitos sobre Don Frijolon, Sabor a Frijol y El Fogón del Frijol combinados es:

$$\text{porcentaje de votos de Los Frijolitos} - (\text{100 por ciento} - \text{porcentaje de votos de Los Frijolitos})$$

Define la función `one_resampled_difference` que devuelve **exactamente un valor** de la ventaja porcentual de Los Frijolitos sobre Don frijolon, Sabor a Frijol y El Fogón del Frijol combinados a partir de una muestra bootstrap de `tbl`. **(8 Puntos)**

*Sugerencia 1:* La ventaja de Los Frijolitos puede ser negativa.

*Sugerencia 2:* Dada una tabla de votos, ¿cómo puedes calcular qué porcentaje de los votos son para un cierto restaurante? **Asegúrate de usar porcentajes, no proporciones, para esta pregunta.**

*Nota:* Si el código base proporcionado dentro de la función no es útil para ti, siéntete libre de abordar la pregunta utilizando tus propias variables.

```python
def one_resampled_difference():
    bootstrap = ...
    los_frijolitos_percentage = ...
    ...
    ...
```

```python
grader.check("q3")
```

**Pregunta 1.5.** Escribe una función llamada `leads_in_resamples` que devuelva un arreglo de 2023 elementos que represente las estimaciones por remuestreo (el resultado de llamar a `one_resampled_difference`) de la ventaja de Los Frijolitos sobre los otros restaurantes combinados. Luego, ejecuta la celda para graficar un histograma de las muestras resultantes. **(8 Puntos)**

*Pista:* Si ves un error relacionado con `NoneType`, considera qué componentes necesita tener una función.

```python
def leads_in_resamples():
    ...

sampled_leads = leads_in_resamples()
Table().with_column('Estimated Lead', sampled_leads).hist("Estimated Lead")

#sampled_leads = leads_in_resamples()
#Table().with_column('Estimated Lead', sampled_leads).hist("Estimated Lead")
```

```python
grader.check("q4")
```

**Pregunta 1.6.** Utiliza los datos simulados en `sampled_leads` de la Pregunta 1.5 para calcular un intervalo de confianza aproximado del 95% para la verdadera ventaja de Los Frijolitos sobre don Frijolon, Sabor a Frijol, y El Fogón del Frijol combinados. **(10 Puntos)**

```python
diff_lower_bound = ...
diff_upper_bound = ...
make_array = (diff_lower_bound, diff_upper_bound)
print("Intervalo de confianza del 95% con bootstrap para la verdadera ventaja de Los Frijolitos sobre los demas es: [{:.2f}%, {:.2f}%]".format(diff_lower_bound, diff_upper_bound))

```

```python
grader.check("q5")
```

## 2. Interpretación de los Intervalos de Confianza

Los intervalos de confianza son herramientas estadísticas esenciales que nos permiten estimar un rango plausible para un parámetro poblacional desconocido, basado en una muestra de datos. En este caso, estamos analizando las preferencias de los clientes por diferentes restaurantes de comida paisa en la ciudad. Nuestra muestra indica la popularidad de los siguientes establecimientos, basándonos en encuestas realizadas:

- **Don Frijolón**: 8% de preferencia
- **Los Frijolitos**: 53% de preferencia
- **Sabor a Frijol**: 25% de preferencia
- **El Fogón del Frijol**: 14% de preferencia

Usando estos datos, calculamos intervalos de confianza que nos ayudarán a inferir con mayor confianza las preferencias reales de la población general por cada restaurante. La interpretación de estos intervalos es crucial para tomar decisiones informadas sobre futuras estrategias de marketing y expansión de mercado.


Hemos calculado el siguiente intervalo de confianza del 95% para el porcentaje de votantes de Los Frijolitos:

$$[50.53, 55.53]$$

(Tu respuesta en el punto 1.3 podría haber sido un poco diferente debido al azar; ¡eso no significa que esté equivocada!)

**Pregunta 2.1.** También creamos intervalos de confianza del 70%, 90% y 99% a partir de la misma muestra, ¡pero olvidamos etiquetar qué intervalo de confianza representaba qué porcentajes! ***Primero***, **asocia cada nivel de confianza (70%, 90%, 99%) con su intervalo correspondiente en la celda a continuación** (por ejemplo, __ % IC: [52.1, 54] $\rightarrow$ reemplaza el espacio en blanco con uno de los tres niveles de confianza). ***Luego***, **explica tu proceso de pensamiento** y cómo llegaste a tus respuestas. **(10 Puntos)**

Los intervalos son los siguientes:

* [50.03, 55.94]
* [52.1, 54]
* [50.97, 54.99]

*Pista:* Si estás atascado en esta pregunta, intenta revisar los Capítulos [13.3](https://inferentialthinking.com/chapters/13/3/Confidence_Intervals.html) y [13.4](https://inferentialthinking.com/chapters/13/4/Using_Confidence_Intervals.html) del libro de texto.

_Escribe aquí tu respuesta, proporcionando contexto adicional según sea necesario:_ ...

**Pregunta 2.2.** Supongamos que producimos **6,000 nuevas muestras** (cada una de ellas es una muestra aleatoria uniforme de 1,500 estudiantes) de la población y **calculamos un intervalo de confianza del 95% para cada una**. **Aproximadamente, ¿cuántos de esos 6,000 intervalos esperas que realmente incluyan el porcentaje verdadero de la población?** **(10 Puntos)**

Asigna tu respuesta a `true_percentage_intervals`.

```python
true_percentage_intervals = ...
```

```python
grader.check("q7")
```

Recuerda el segundo intervalo de confianza de bootstrap que creaste, el cual estimaba la ventaja de Los Frijolitos sobre Don Frijolón, Sabor a Frijol y El Fogón del Frijol combinados. Entre los votantes en la muestra, la ventaja de Los Frijolitos era del 6%. El intervalo de confianza del 95% de nuestro calculo para la ventaja verdadera (en la población de todos los votantes) fue:

$$[1.2, 11.2]$$

Supongamos que estamos interesados en probar una pregunta sencilla de sí o no:

> "¿Es el porcentaje de votos para Los frijolitos  igual al porcentaje de votos para don Frijolon, Sabor a Frijol y El Fogón del Frijol combinados?"

Nuestra hipótesis nula es: que los porcentajes son iguales, o equivalentemente, que la ventaja de Los frijolitos es exactamente 0. 

Nuestra hipótesis alternativa es: que la ventaja de los frijolitos no es igual a 0. 

En las preguntas siguientes, no calcules ningún intervalo de confianza por tu cuenta; usa solo el intervalo de confianza del 95% dados en este enunciado:

**Pregunta 2.3.** En un estudio estadístico, hemos decidido utilizar un nivel de significancia del 5%, lo que significa que nuestro punto de corte para el valor p es del 5%. Con base en este nivel de significancia, ¿debemos rechazar la hipótesis nula, mantenerla o es imposible determinarlo utilizando nuestro intervalo de confianza personalizado? **(10 Puntos)**

Debes asignar el valor numérico correcto a `cutoff_five_percent` según la respuesta adecuada:

1. Rechazar la hipótesis nula / Los datos son congruentes con la hipótesis alternativa
2. No rechazar la hipótesis nula / Los datos son congruentes con la hipótesis nula
3. Es imposible determinarlo usando nuestro intervalo de confianza personalizado

*Pista:* Es crucial considerar cómo se relacionan el nivel de significancia y el intervalo de confianza. Si necesitas más información, revisa [este capítulo](https://inferentialthinking.com/chapters/13/4/Using_Confidence_Intervals.html) del libro de texto.

```python
cutoff_five_percent = 

```

```python
grader.check("q8")
```

**Pregunta 2.4.** En esta pregunta, estamos explorando el impacto de cambiar el nivel de significancia en un análisis estadístico. Específicamente, vamos a considerar qué ocurre si ajustamos nuestro valor de corte (nivel de significancia) al 1%. Con un nivel de significancia más estricto, el criterio para rechazar la hipótesis nula se vuelve más riguroso. ¿Qué implica esto en términos de nuestro análisis de hipótesis? ¿Podemos aún rechazar la hipótesis nula, o los datos ahora sugieren que no deberíamos hacerlo? Analiza esto utilizando el intervalo de confianza que hemos calculado según nuestro criterio. **(10 Puntos)**

Asigna `cutoff_one_percent` al número correspondiente a la respuesta correcta.

1. Rechazar la hipótesis nula / Los datos son consistentes con la hipótesis alternativa
2. No rechazar la hipótesis nula / Los datos son consistentes con la hipótesis nula
3. No podemos determinarlo usando nuestro intervalo de confianza del personal


```python
cutoff_one_percent =

```

```python
grader.check("q9")
```

**Pregunta 2.5.** Imagina que decides utilizar un nivel de significancia del 10% en tu análisis estadístico. ¿Cómo afectaría esto a tu decisión respecto a la hipótesis nula? Con este nuevo nivel de significancia, ¿tu decisión sería rechazar la hipótesis nula, no rechazarla, o es que el intervalo de confianza no nos permite llegar a una conclusión definida? **(10 Puntos)**

Asigna `cutoff_ten_percent` al número que corresponde a la respuesta correcta.

1. Rechazar la hipótesis nula / Los datos son consistentes con la hipótesis alternativa
2. No rechazar la hipótesis nula / Los datos son consistentes con la hipótesis nula
3. No se puede determinar usando nuestro intervalo de confianza

```python
cutoff_ten_percent = 

```

```python
grader.check("q10")
```

¡Felicitaciones por completar la Laboratorio 8!

¡Terminaste con el laboratorio!

**Información importante sobre el envío:**
- **Ejecuta todas las pruebas** y verifica que todas pasan
- **Guardar** desde el menú **Archivo**
- **Ejecute la celda final **
- Luego, ve a este [enlace](https://cienciadatosudea.github.io/001_SitioWebFundCienciaDatos/04entregas.html) y envía el archivo ipynb a la tarea correspondiente. El nombre de esta tarea es "Lab XX Autograder", donde XX es el número de laboratorio: 01, 02, 03, etc.

- Si terminas temprano en el laboratorio regular, **registra tu asistencia con el instructor**.

**Es su responsabilidad asegurarse de que su trabajo esté guardado antes de ejecutar la última celda.**

`Referencia`: Este material esta basado en [Data 8: The Foundations of Data Science](https://www.data8.org/) - University of California, Berkeley.


## Mascotas de Data 8

**Tigrillo** te felicita por haber terminado la Tarea 8. ¡Solo te quedan 3 tareas más para completar el curso exitosamente! 

<img src="./Tigrillo.jpg" width="30%" alt="Imagen en primer plano de un perro sonriendo, sentado en el césped."/>

Mascota de la semana: **Chipper**

## Entrega

Antes de continuar con el siguiente paso, asegúrate de haber ejecutado todas las celdas de tu cuaderno en el orden adecuado para garantizar que todas las imágenes y gráficos se muestren correctamente en la salida final. La celda a continuación generará un archivo zip que deberás enviar. **¡Por favor, guarda todos los cambios antes de exportar el archivo!**

```python

```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

These are some submission instructions.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q0 (points: 1)
```
>>> def test_one_resampled_percentage():
...     votes_tbl = Table().with_columns('Votos', ['Don Frijolón', 'Los Frijolitos', 'Don Frijolón', 'Sabor a Frijol', 'El Fogón del Frijol'])
...     results = [one_resampled_percentage(votes_tbl) for _ in range(100)]
...     assert all((0 <= res <= 100 for res in results)), 'El porcentaje debe estar entre 0 y 100'
...     all_don_frijolon = Table().with_columns('Votos', ['Los Frijolitos'] * 10)
...     no_don_frijolon = Table().with_columns('Votos', ['Don Frijolon', 'Sabor a Frijol', 'El Fogón del Frijol'] * 10)
...     assert one_resampled_percentage(no_don_frijolon) == 0, 'Debe devolver 0% cuando no hay votos para Los Frijolitos'
...     print('Todas las pruebas pasaron exitosamente.')
>>> test_one_resampled_percentage()
Todas las pruebas pasaron exitosamente.
```


Tests for q1 (points: 1)
```
>>> z = percentages_in_resamples()
>>> u = np.mean(z)
>>> longitud = len(z)
>>> assert np.isclose(u, 52.9, 0.1)
>>> assert np.isclose(longitud, 2023, 0.1)
```


Tests for q10 (points: 1)
```
>>> assert cutoff_ten_percent == 2, 'Revisa el valor de cutoff_one_percent. Considera el intervalo de confianza y el p-valor.'
```


Tests for q2 (points: 1)
```
>>> assert np.isclose(lower_bound, 50.4, 0.3)
>>> assert np.isclose(upper_bound, 55.5, 0.3)
```


Tests for q3 (points: 1)
```
>>> ventaja = one_resampled_difference()
>>> assert np.count_nonzero(ventaja < 20)
```


Tests for q4 (points: 1)
```
>>> assert np.isclose(np.mean(sampled_leads), 6, 0.5)
```


Tests for q5 (points: 1)
```
>>> assert np.isclose(diff_lower_bound, 0.93, 0.5)
>>> assert np.isclose(diff_upper_bound, 11.2, 0.5)
```


Tests for q7 (points: 1)
```
>>> assert true_percentage_intervals == 5700
```


Tests for q8 (points: 1)
```
>>> assert cutoff_five_percent == 1
```


Tests for q9 (points: 1)
```
>>> assert cutoff_one_percent == 2, 'Revisa el valor de cutoff_one_percent. Considera el intervalo de confianza y el p-valor.'
```


Tests for qt (points: 0)
```
>>> assert isinstance(NombreApellidos, str)
>>> assert isinstance(NumeroCarnet, str)
>>> assert isinstance(Email, str)
>>> assert len(NombreApellidos) > 3
>>> assert len(NumeroCarnet) > 3
>>> assert len(Email) > 3
```
