```python
# Initialize Otter
import otter
grader = otter.Notebook("lab07_.ipynb")
```



<p><img alt="Colaborat" style="height: 220px;" src="https://github.com/CienciaDatosUdea/001_DataSetFundCienciaDatos/blob/main/semestre2024-2/lab/lab02/Escudo-UdeA.svg.png?raw=true" align="left" hspace="10px" vspace="0px"></p>


# Fundamentos en ciencias de datos



# Laboratorio 7: Great British Bake Off (Prueba A/B)

¡Bienvenido al Laboratorio 7 de Data 8! El laboratorio de esta semana se centrará en las pruebas A/B utilizando datos del  programa de televisión británico, [*The Great British Bake Off*](https://en.wikipedia.org/wiki/The_Great_British_Bake_Off).



Como recordatorio, estas son las políticas para obtener una nota completa:

1. Recuerden que  recibiran el 80% del laboratorio solo con **asistir** a la sección de los viernes laboratorio y lograr un progreso significativo en ```notebook de jupyter```.



**Envío**: Una vez que haya terminado, ejecute todas las celdas además de la última, seleccione Archivo > Guardar cuaderno y luego ejecute la celda final. El resultado contendrá un archivo zip que podrá utilizar para enviarlo.

Comencemos configurando las pruebas e importaciones ejecutando la siguiente celda.

```python
# Ejecute esta celda para configurar el cuaderno, pero no la cambie.

# Estas líneas importan los módulos Numpy y Datascience.
import numpy as np
from datascience import *

# Estas líneas hacen algo de magia en la trama.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plots
plots.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore', (FutureWarning, DeprecationWarning))

```

<!-- BEGIN QUESTION -->

Ingresa tu nombre  con apellido y número de carnet en las variables dadas en la celda siguiente:
Ejemplo:
```python
NombreApellidos="Alan Turing"
NumeroCarnet="27182818"
Email="alan.turing@udea.edu.co"
```

```python
NombreApellidos="..."
NumeroCarnet=".."
Email="..."
```

```python
grader.check("qt")
```

<!-- END QUESTION -->

## 1. Pruebas A/B

La prueba A/B es una forma de prueba de hipótesis que le permite hacer comparaciones entre dos distribuciones. También podemos referirnos a una prueba A/B como prueba de permutación.

Casi nunca se te pedirá explícitamente que realices una prueba A/B. Asegúrese de poder identificar situaciones en las que la prueba sea apropiada y sepa cómo implementar correctamente cada paso. A menudo, utilizamos una prueba A/B para determinar si dos muestras provienen o no de la misma distribución subyacente.

**Pregunta 1.1.** Las siguientes afirmaciones son los pasos de una prueba de hipótesis A/B presentada en *orden aleatorio*:

1. Elija un estadístico de prueba (normalmente la diferencia de medias entre dos categorías)

2. Mezcle las etiquetas de la muestra original, encuentre su estadístico de prueba simulada y repita muchas veces.

3. Encuentre el valor del estadístico de prueba observado.

4. Calcule el valor p basándose en los estadísticos de prueba observadas y simuladas.

5. Definir la hipotesis nulo y alternativa.

6. Utilice el valor p y el límite del valor p para sacar una conclusión sobre la hipótesis nula.

Para responder la pregunta, asigne `ab_test_order` a una matriz de números enteros con los números correctos de una prueba A/B, donde el primer elemento de la matriz es el primer paso de una prueba A/B y el último elemento de la matriz es el último paso de una prueba A/B.

```python
ab_test_order = ...
```

```python
grader.check("q0")
```

**Pregunta 1.2.** 
Si la hipótesis nula de una prueba A/B es correcta, ¿debería influir el orden de las etiquetas (es decir, la asignación de participantes a los grupos A o B) en las diferencias de medias entre ambos grupos? Asigna False or True a la variable respuesta y justifica más abajo. 

Explica por qué es importante aleatorizar las etiquetas en una prueba A/B. 

Si estás en una sesión de laboratorio, confirma tu respuesta con un compañero o con el profesor antes de continuar.


```python
respuesta=...
```

```python
grader.check("q1")
```

_Escriba su justificación aquí, reemplazando este texto._

## 2. *The Great British Bake Off*

>"The Great British Bake Off (a menudo abreviado como Bake Off o GBBO, el gran pastelero en español) es una competencia de repostería de la televisión británica, en la que un grupo de panaderos aficionados compiten entre sí en una serie de rondas, intentando impresionar a un grupo de jueces con sus habilidades para hornear" [Wikipedia](https://en.wikipedia.org/wiki/The_Great_British_Bake_Off)

En cada semana de la competencia, los jueces eligen a un concursante como **"*Star Baker*"** (*mejor panadero de la semana*). Este título se otorga al participante que tuvo el mejor desempeño en los desafíos de esa semana. 

Al final de la temporada, se selecciona a un único **ganador del programa**.

Con estos datos, nuestro objetivo es **analizar la relación entre ganar premios *Star Baker* y la probabilidad de convertirse en el campeón de la temporada**. Es decir, queremos saber si **obtener más títulos de *Star Baker* realmente aumenta las probabilidades de ganar** o si es solo una coincidencia.


**Pregunta 2.1:**  
Queremos investigar si obtener más premios *Star Baker* **causa** un aumento en la probabilidad para ganar la temporada. 

Sin embargo, **¿por qué no es suficiente con comparar la frecuencia de premios *Star Baker* entre quienes ganaron la temporada y quienes no?**

```python

```

_Escriba su respuesta aquí, reemplazando este texto._

...

### 🧪 **Ejecución del experimento**  

Realizaremos la siguiente prueba de hipótesis para determinar la asociación entre ganar y la cantidad de premios Star Baker. La población que estamos examinando son todos los concursantes de las temporadas 2 a 11 de *The Great British Bake Off* GBBO. Vamos a utilizar las siguientes hipótesis nula y alternativa:

- 📌 **Hipótesis nula (H₀):**  La distribución de los premios *Star Baker* entre los concursantes que ganaron su temporada y los concursantes que no ganaron su temporada es la misma.

- 💡 **Hipótesis alternativa (H₁):** Los concursantes que ganen su temporada del programa ganarán más premios Star Baker en promedio.

Nuestra hipótesis alternativa está relacionada con nuestra sospecha de que los concursantes que ganan más premios Star Baker son más hábiles, por lo que tienen más probabilidades de ganar la temporada.



**Pregunta 2.2.** ¿Deberíamos utilizar una prueba A/B para probar estas hipótesis? En caso afirmativo, ¿cuál es nuestro grupo "A" y cuál es nuestro grupo "B"?

```python
grupo_a = "Escriba acá quien pueden ser el grupo A"
grupo_b = "Escriba acá quien pueden ser del grupo B"
```

_Justifique su respuesta aquí, reemplazando este texto._

Verifique sus respuestas con sus compañeros o el profesor de clase.

La siguiente tabla de "panaderos" describe la cantidad de premios de panadero estrella que ganó cada concurso y si ganaron o no su temporada ("1" si ganaron, "0" si no ganaron). Los datos se agregaron manualmente de Wikipedia para las temporadas 2 a 11 del programa. Aleatorizamos el orden de las filas para no estropear el resultado del espectáculo.

```python
bakers = Table.read_table("star_bakers.csv")
bakers.show(3)
```

**Pregunta 2.3.** Cree una nueva tabla llamada `medias` que contenga el número medio de premios de panadero estrella para los panaderos que no ganaron (`ganaron==0`) y los panaderos que ganaron (`ganaron==1`). ). La tabla debe tener los nombres de las columnas "won" y "star baker awards average".

```python
medias = ...
```

```python
grader.check("q4")
```

**Pregunta 2.4.** Visualice la distribución de los premios Star Baker para los ganadores y no ganadores como histogramas superpuestos. Debe utilizar los contenedores que le proporcionamos.

Sugerencia: querrá utilizar el argumento de grupo de `tbl.hist`. Para producir varios histogramas superpuestos basados en valores únicos en una columna determinada, podemos hacer algo como `tbl.hist(..., group=<col_name>, bins=...)`. Esto graficará un histograma para cada valor único en la columna especificada, todo en un solo gráfico.

```python
useful_bins = np.arange(0, 7)
...
```

**Pregunta 2.5.** Queremos determinar si existe una diferencia entre la distribución de los premios Star Baker entre ganadores y no ganadores.

¿Cuál debería ser el estadístico de prueba? ¿Qué valores de este estadístico de prueba respaldan la nula y qué valores respaldan la alternativa? **Asigne `test_option` al número correspondiente a la estadística de prueba correcta.**

1. Valor absoluto de la diferencia de medias entre ambos grupos; los valores altos apoyan el nulo
2. Valor absoluto de la diferencia entre las medias entre ambos grupos; los valores bajos apoyan el nulo
3. Premios Star Baker promedio para los ganadores menos premios Star Baker promedio para los no ganadores; los valores altos apoyan el nulo
4. Premios Star Baker promedio para los ganadores menos premios Star Baker promedio para los no ganadores; los valores bajos apoyan el nulo

Antes de continuar, confirma tu respuesta con un compañero o en los foros de discusión.

_Pista:_ Deberías pensar en qué medidas utilizamos para describir una distribución.


```python
test_option = ...
```

```python
grader.check("q6")
```

```python
medias
```

**Pregunta 2.6.** Establezca la "diferencia observada" en el estadístico de prueba observado utilizando la tabla de "medias".

```python
diferencia_observada=...
diferencia_observada
```

```python
grader.check("q7")
```

**Pregunta 2.7.** 
A partir de una tabla llamada `bakers`, que contiene:  
- Una columna con el número de veces que un concursante ganó *Star Baker* (`star_col`),  
- Una columna que indica si ganó la temporada (`won_col`),  

**¿Cómo escribirías una función que calcule el estadística de prueba adecuado para evaluar si existe una relación entre ganar *Star Baker* y ganar la temporada?**
*Pista:* Asegúrese de tener en cuenta la direccionalidad de nuestra hipótesis alternativa.

```python
def prueba_estadistica(panaderos, label_col, val_col):
    ...
    return estadistica_prueba

prueba_estadistica(bakers,"won", "star baker awards mean")
```

```python
grader.check("q8")
```

Cuando ejecutamos una simulación para pruebas A/B, volvemos a muestrear **mezclando las etiquetas** de la muestra original. Si la hipótesis nula es cierta y las distribuciones de premios del panadero estrella son las mismas, esperamos que la diferencia en la media de los premios del panadero estrella no cambie cuando se cambian las etiquetas en "won".

**Pregunta 2.8.** Escriba una función `simulate_and_test_statistic` para calcular una prueba de nuestra prueba A/B. Su función debería ejecutar una simulación y devolver una estadística de prueba.

```python
def simulate_and_test_statistic(table, label_column):
    shuffled_labels=...
    
    ...
    ...
    return ...
simulate_and_test_statistic(bakers, "won")
```

```python
grader.check("q9")
```

**Pregunta 2.9.** Simule 5000 pruebas de nuestra prueba A/B y almacene las estadísticas de la prueba en una matriz llamada "differences".

```python
differences = ...
for ...
differences
```

```python
grader.check("q10")
```

Ejecute la celda a continuación para ver un histograma de sus estadísticas de prueba simuladas trazadas con su estadística de prueba observada.

```python
Table().with_column('Difference Between Group Means', differences).hist(bins=20)
plots.scatter(diferencia_observada, 0, color='red', s=30, zorder=2)
plots.ylim(-0.1, 1.35);
```

**Pregunta 2.10.** Encuentre el valor p para su prueba y asígnelo a `empirical_p`.

```python
empirical_p = ...
empirical_p
```

```python
grader.check("q11")
```

**Pregunta 2.11.** Utilizando un límite de valor P del 5%, saque una conclusión sobre las hipótesis nula y alternativa. Describe tus hallazgos usando un lenguaje simple y no técnico. ¿Qué le dice su análisis sobre la asociación entre los premios Star Baker y ganar? ¿Qué puede afirmar sobre la causalidad a partir de su análisis estadístico? Confirme su respuesta con un compañero o el instructor.

_Escriba su respuesta aquí, reemplazando este texto._

## ¡Todo listo!

**Polaris** te felicita,  lo has hecho bien, has terminado el laboratorio número 7.

<img src="Polaris.jpg" alt="Foto en primer plano de un golden retriever mirando a la cámara" width="300"/>



---

¡Terminaste con el laboratorio!

**Información importante sobre el envío:**
- **Ejecuta todas las pruebas** y verifica que todas pasan
- **Guardar** desde el menú **Archivo**
- **Ejecute la celda final **
- Luego, ve a este [enlace](https://cienciadatosudea.github.io/001_SitioWebFundCienciaDatos/04entregas.html) y envía el archivo ipynb a la tarea correspondiente. El nombre de esta tarea es "Lab XX Autograder", donde XX es el número de laboratorio: 01, 02, 03, etc.

- Si terminas temprano en el laboratorio regular, **registra tu asistencia con el instructor**.

**Es su responsabilidad asegurarse de que su trabajo esté guardado antes de ejecutar la última celda.**

`Referencia`: Este material esta basado en [Data 8: The Foundations of Data Science](https://www.data8.org/) - University of California, Berkeley.


## Envío

Asegúrese de haber ejecutado todas las celdas de su cuaderno en orden antes de ejecutar la siguiente celda, para que todas las imágenes/gráficos aparezcan en el resultado. La siguiente celda generará un archivo zip para que lo envíe. **¡Guarde antes de exportar!**

```python

```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

These are some submission instructions.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q0 (points: 1)
```
>>> ab_test_order = make_array(5, 1, 3, 2, 4, 6)
>>> assert len(ab_test_order) == 6
>>> assert ab_test_order.item(0) == 5
>>> assert ab_test_order.item(1) == 1
>>> assert ab_test_order.item(2) == 3
>>> assert ab_test_order.item(3) == 2
>>> assert ab_test_order.item(4) == 4
>>> assert ab_test_order.item(5) == 6
```


Tests for q1 (points: 1)
```
>>> assert respuesta == False
```


Tests for q10 (points: 1)
```
>>> assert len(differences) == 5000
>>> assert np.isclose(np.mean(differences), 0.26, 0.1)
```


Tests for q11 (points: 1)
```
>>> assert isinstance(empirical_p, float), 'empirical_p debe ser un número decimal'
>>> assert 0 <= empirical_p <= 0.05, 'empirical_p debe estar entre 0 y 1'
```


Tests for q4 (points: 1)
```
>>> assert set(medias.labels) == {'won', 'star baker awards average'}
>>> assert medias.num_rows == 2
>>> assert all(medias.column('won') == [0, 1])
>>> assert np.isclose(medias.column('star baker awards average').item(0), 0.651786, 0.01)
>>> assert np.isclose(medias.column('star baker awards average').item(1), 1.5, 0.01)
```


Tests for q6 (points: 1)
```
>>> assert test_option in [1, 2, 3, 4], 'El valor de test_option debe ser uno de los números proporcionados en las opciones.'
>>> assert test_option == 2, 'Recuerda que estamos buscando los valores bajos que apoyan la nula.'
```


Tests for q7 (points: 1)
```
>>> assert isinstance(diferencia_observada, (int, float)) and diferencia_observada > 0
>>> assert np.isclose(diferencia_observada, 0.848, 0.01)
```


Tests for q8 (points: 1)
```
>>> panaderos_test = Table().with_columns('Categoria', ['A', 'A', 'B', 'B'], 'Valores', [1, 2, 3, 4])
>>> assert np.isclose(prueba_estadistica(bakers, 'won', 'star baker awards mean'), 0.848, 0.01)
```


Tests for q9 (points: 1)
```
>>> panaderos_test = Table().with_columns('Categoria', ['A', 'A', 'B', 'B'], 'Valores', [2, 2, 2, 2])
>>> assert isinstance(simulate_and_test_statistic(panaderos_test, 'Categoria'), (int, float)) and simulate_and_test_statistic(panaderos_test, 'Categoria') >= 0
```


Tests for qt (points: 0)
```
>>> assert isinstance(NombreApellidos, str)
>>> assert isinstance(NumeroCarnet, str)
>>> assert isinstance(Email, str)
>>> assert len(NombreApellidos) > 3
>>> assert len(NumeroCarnet) > 3
>>> assert len(Email) > 3
```
