```python
# Initialize Otter
import otter
grader = otter.Notebook("proyecto_01.ipynb")
```

<p><img alt="Colaborat" style="height: 220px;" src="https://github.com/CienciaDatosUdea/001_DataSetFundCienciaDatos/blob/main/semestre2024-2/lab/lab02/Escudo-UdeA.svg.png?raw=true" align="left" hspace="10px" vspace="0px"></p>


# Fundamentos en ciencias de datos


## Proyecto 01: Población Mundial y Pobreza


En este proyecto, explorarás datos de [Gapminder.org](http://gapminder.org), un sitio web dedicado a ofrecer una visión del mundo basada en hechos y cómo ha cambiado. Ese sitio incluye varias visualizaciones y presentaciones de datos, pero también publica los datos en bruto que utilizaremos en este proyecto para recrear y ampliar algunas de sus visualizaciones más famosas. El sitio web de Gapminder recopila datos de muchas fuentes y los compila en tablas que describen muchos países alrededor del mundo. Todos los datos que agregan se publican en el [Systema Globalis](https://github.com/open-numbers/ddf--gapminder--systema_globalis/blob/master/README.md). Su objetivo es compilar todas las estadísticas públicas; sociales, económicas y ambientales; en un conjunto de datos . Todos los conjuntos de datos de este proyecto se copian directamente del Systema Globalis sin ningún cambio. Este proyecto está dedicado a [Hans Rosling](https://en.wikipedia.org/wiki/Hans_Rosling) (1948-2017), quien promovió el uso de datos para entender y priorizar los desafíos del desarrollo global.

### Logística **Fecha límite.** 


Este proyecto se debe entregar a más tardar a las **11:00pm del 6 de Mayo**. 


**Punto de control.** Para obtener la nota  completa en el punto de control, debes completar las preguntas hasta el punto de control, **aprobar todas las pruebas de autocalificación** <span style="color: #BC412B">**El punto de control vale el 5% de la calificación total de tu proyecto `proyecto01.ipynb`**</span>. 

Después de haber enviado el punto de control, aún puedes cambiar tus respuestas del proyecto antes de la fecha límite final del proyecto; solo tu envío final, a la asignación "Proyecto 1", se calificará completo `projecto01.ipynb`. 

**Compañeros.** Puedes trabajar con un compañero; tu compañero debe ser de tu sección de laboratorio asignada.  

**Reglas.** No compartas tu código con nadie más que tu compañero. Puedes discutir las preguntas con otros estudiantes, pero no compartas las respuestas. La experiencia de resolver los problemas en este proyecto te preparará para los exámenes (y la vida). Si alguien te pide la respuesta, ¡resiste! En su lugar, puedes demostrar cómo resolverías un problema similar.

**Pruebas.** <span style="color: #BC412B">Las pruebas que se proporcionan **no son exhaustivas** y pasar las pruebas de una pregunta **no** significa que respondiste la pregunta correctamente.</span> Las pruebas generalmente solo verifican que tu tabla tenga las etiquetas de columnas correctas. Sin embargo, se aplicarán más pruebas para verificar la exactitud de tu envío con el fin de asignar tu puntuación final, así que ten cuidado y verifica tu trabajo. Podrías querer crear tus propias verificaciones a lo largo del camino para ver si tus respuestas tienen sentido. Además, antes de enviar, asegúrate de que ninguna de tus celdas tome mucho tiempo en ejecutarse (varios minutos).

**Preguntas de respuesta libre:** Asegúrate de poner las respuestas a las preguntas escritas en la celda indicada que proporcionamos. **Cada pregunta de respuesta libre debe incluir una explicación** que responda adecuadamente a la pregunta.

**Consejo.** Desarrolla tus respuestas de manera incremental. Para realizar una manipulación de tabla complicada, divídela en pasos, realiza cada paso en una línea diferente, da un nuevo nombre a cada resultado y verifica que cada resultado intermedio sea lo que esperas. Puedes agregar los nombres o funciones adicionales que desees en las celdas proporcionadas. Asegúrate de estar usando nombres de variables distintos y significativos en todo el cuaderno. En esa línea, **NO** reutilices los nombres de variables que usamos cuando calificamos tus respuestas. Por ejemplo, en la Pregunta 1 de la sección de Pobreza Global te pedimos que asignes una respuesta a `b_pop`. No reasignes el nombre de variable `b_pop` a nada más en tu cuaderno, de lo contrario existen posibilidades de que nuestras pruebas califiquen contra lo que `b_pop` fue reasignado. **Nunca** estás restringido a utilizar solo una línea de código para resolver una pregunta en este proyecto o en cualquier otro. ¡Siéntete libre de usar variables intermedias y múltiples líneas tanto como desees!

Para comenzar, carga `datascience`, `numpy`, `plots` y `otter`.

```python
# Run this cell to set up the notebook, but please don't change it.
# These lines import the NumPy and Datascience modules.
from datascience import *
import numpy as np
# These lines do some fancy plotting magic.
%matplotlib inline
import matplotlib.pyplot as plots
plots.style.use('fivethirtyeight')
from ipywidgets import interact, interactive, fixed, interact_manual
import ipywidgets as widgets

```

_Type your answer here, replacing this text._

```python
Nombre = "..."
Carnet = "..."
```

<!-- END QUESTION -->

## 1. Crecimiento de la Población Mundial

La población mundial de seres humanos alcanzó los 1000 millones alrededor del año 1800; 3000 millones alrededor de 1960, y 7000 millones alrededor de 2011. El impacto potencial del crecimiento exponencial de la población ha preocupado a científicos, economistas y políticos por igual. La División de Población de las Naciones Unidas, estima que es probable que la población mundial continúe creciendo a lo largo del siglo XXI, pero a un ritmo más lento, quizás alcanzando los 11000 millones para el año 2100. Sin embargo, la ONU no descarta escenarios de un crecimiento más extremo.<a href="http://www.pewresearch.org/fact-tank/2015/06/08/scientists-more-worried-than-public-about-worlds-growing-population/ft_15-06-04_popcount/">  <img src="pew_population_projection.png"/> </a>

En esta parte del proyecto, examinaremos algunos de los factores que influyen en el crecimiento de la población y cómo han ido cambiando a lo largo de los años y en todo el mundo. Hay dos subpartes principales en este análisis.

- Primero, examinaremos los datos de un país, Bangladesh. Veremos cómo han cambiado factores como la esperanza de vida, la tasa de fertilidad y la mortalidad infantil a lo largo del tiempo en Bangladesh, y cómo se relacionan con la tasa de crecimiento de la población.

- A continuación, examinaremos si los cambios que hemos observado para Bangladesh son particulares de ese país o si reflejan patrones generales observables en otros países también. Estudiaremos aspectos del crecimiento de la población mundial y veremos cómo han estado cambiando.La primera tabla que consideraremos contiene la población total de cada país a lo largo del tiempo. Ejecuta la celda a continuación.

```python
population = Table.read_table('population.csv').where("time", are.below(2021))
population.show(3)
```

**Nota:** El archivo csv de población también se puede encontrar [aquí](https://github.com/open-numbers/ddf--gapminder--gapminder_world/blob/master/ddf--datapoints--population_total--by--geo--time.csv). Los datos para este proyecto fueron descargados en febrero de 2017.

### Bangladesh 

La nación de [Bangladesh](https://es.wikipedia.org/wiki/Bangladesh) se estableció como una democracia parlamentaria después de que terminara la Guerra de Liberación de Bangladesh en 1971. La nación naciente, devastada por la guerra, se enfrentó casi de inmediato a inundaciones y hambrunas.  En esta sección del proyecto, examinaremos aspectos del desarrollo de Bangladesh desde ese momento. En la tabla `población`, la columna `geo` contiene códigos de tres letras establecidos por la [Organización Internacional de Normalización](https://es.wikipedia.org/wiki/Organizaci%C3%B3n_Internacional_de_Normalizaci%C3%B3n) (ISO) en el estándar [Alpha-3](https://en.wikipedia.org/wiki/ISO_3166-1_alpha-3). Comenzaremos examinando de cerca a Bangladesh. Utiliza el enlace Alpha-3 para encontrar el código de tres letras para Bangladesh.

**Pregunta 1.** Crea una tabla llamada `b_pop` que tenga dos columnas etiquetadas como `time` y `population_total` desde la información dada en la tabla `population.csv`. La primera columna debe contener los años desde 1970 hasta 2020 (incluyendo tanto 1970 como 2020) y la segunda debe contener la población de Bangladesh en cada uno de esos años.

```python
b_pop = ...
b_pop
```

```python
grader.check("q1")
```

Ejecuta la siguiente celda para crear una tabla llamada `b_five` que contenga la población de Bangladesh cada cinco años. A simple vista, ¡parece que la población de Bangladesh ha estado creciendo rápidamente!

```python
b_pop.set_format('population_total', NumberFormatter)
fives = np.arange(1970, 2021, 5) # 1970, 1975, 1980, ...
b_five = b_pop.sort('time').where('time', are.contained_in(fives))
b_five.show()
```


**Pregunta 2.** Asigna `initial` a un arreglo que contenga la población para cada intervalo de cinco años desde 1970 hasta 2015 (inclusive). Luego, asigna `changed` a un arreglo que contenga la población para cada intervalo de cinco años desde 1975 hasta 2020 (inclusive). El primer arreglo debe incluir tanto 1970 como 2015, y el segundo arreglo debe incluir tanto 1975 como 2020. Debes usar la tabla `b_five` para crear ambos arreglos, filtrando primero la tabla para que solo contenga los años relevantes. La tasa de crecimiento anual para un período de tiempo es igual a:$$\left(\left(\frac{\text{Población al final del período}}{\text{Población al inicio del período}}\right)^{\displaystyle\frac{1}{\text{número de años}}}\right) -1$$ 

Hemos proporcionado el siguiente código, que utiliza los arreglos `initial` y `changed` para construir la tabla `b_five_growth`. Esta tabla incluye una nueva columna llamada `annual_growth`, junto con las columnas originales de la tabla `b_five`.

No te preocupes por el cálculo de las tasas de crecimiento; ejecuta la prueba a continuación para probar tu solución. Si estás interesado en cómo llegamos a la fórmula para las tasas de crecimiento, consulta la sección [tasas de crecimiento](https://inferentialthinking.com/chapters/03/2/1/Growth.html) del libro de texto.

```python
initial = ...
changed = ...

b_1970_through_2015 = b_five.where('time', are.below_or_equal_to(2015)) 
b_five_growth = b_1970_through_2015.with_column('annual_growth', (changed/initial)**0.2-1)
b_five_growth.set_format('annual_growth', PercentFormatter)

```

```python
grader.check("q2")
```

Mientras que la población ha crecido cada cinco años desde 1970, la tasa de crecimiento anual disminuyó drásticamente de 1985 a 2015. Veamos alguna otra información para desarrollar una posible explicación. Ejecuta la siguiente celda para cargar tres tablas adicionales de mediciones sobre países a lo largo del tiempo.

```python
life_expectancy = Table.read_table('life_expectancy.csv').where('time', are.below(2021))
child_mortality = Table.read_table('child_mortality.csv').relabel(2, 'child_mortality_under_5_per_1000_born').where('time', are.below(2021))
fertility = Table.read_table('fertility.csv').where('time', are.below(2021))

```

La tabla `life_expectancy` contiene una estadística que a menudo se utiliza para medir cuánto tiempo viven las personas, llamada *esperanza de vida al nacer*. Este número, para un país en un año dado, [no mide cuánto tiempo se espera que vivan los bebés nacidos ese año](http://blogs.worldbank.org/opendata/what-does-life-expectancy-birth-really-mean). En cambio, mide cuánto tiempo viviría alguien, en promedio, si las *condiciones de mortalidad* de ese año persistieran durante toda su vida. Estas "condiciones de mortalidad" describen qué fracción de personas de cada edad sobrevivió al año. Por lo tanto, es una forma de medir la proporción de personas que permanecen vivas, agregada en diferentes grupos de edad de la población.

Ejecuta las siguientes celdas a continuación para ver `esperanza_de_vida`, `mortalidad_infantil` y `fertilidad`. Refiérete a estas tablas ya que serán útiles para responder preguntas adicionales.

```python
life_expectancy.show(3)
```

```python
child_mortality.show(3)
```

```python
fertility.show(3)
```

**Pregunta 3.** Tal vez la población está creciendo más lentamente porque las personas no viven tanto tiempo. Utiliza la tabla `life_expectancy` para dibujar un gráfico de líneas con los años 1970 y posteriores en el eje horizontal que muestre cómo ha cambiado la *esperanza de vida al nacer* en Bangladesh.

```python
#Escribir el codigo 
...
```

```python
grader.check("q3")
```

**Pregunta 4.** Suponiendo que todo lo demás permanece igual, ¿los cambios en la esperanza de vida en el gráfico anterior explican directamente por qué la tasa de crecimiento de la población disminuyó de 1985 a 2015 en Bangladesh? ¿Por qué sí o por qué no?

Pista: Reflexiona sobre cómo los cambios en la esperanza de vida podrían o no estar relacionados con la tasa de crecimiento de la población. ¿Qué otros factores importantes podrían haber influido en esta disminución? Además, incorpora un análisis de los eventos significativos ocurridos en Bangladesh en 1991 y cómo pudieron haber impactado en la dinámica poblacional. Si utilizas información de fuentes externas, asegúrate de citarlas adecuadamente

```python
explicacion = "Escribe tu respuesta aqui"
```

La tabla `fertility` contiene una estadística que a menudo se utiliza para medir cuántos bebés están naciendo, la *tasa de fecundidad total*. Este número describe el [número de hijos que tendría una mujer a lo largo de su vida](https://www.measureevaluation.org/prh/rh_indicators/specific/fertility/total-fertility-rate), en promedio, si las tasas actuales de nacimiento por edad de la madre se mantuvieran durante todos sus años fértiles, asumiendo que sobreviva hasta los 49 años.

**Pregunta 5.** Completa la función `fertility_over_time`. Esta función recibe el código Alpha-3 de un país como `country_code` y un año de `start`. Devuelve una tabla de dos columnas con las etiquetas `Anio` y `Ninos por mujer` que puede ser utilizada para generar un gráfico de líneas de la tasa de fecundidad del país cada año, comenzando desde el año inicio `start`. La gráfica debe incluir el año `start` y todos los años posteriores que aparezcan en la tabla `fertility`. Luego, determina el código Alpha-3 para Bangladesh. El código para Bangladesh y el año 1970 se utilizan en la llamada a tu función `fertility_over_time` para graficar cómo ha cambiado la tasa de fecundidad de Bangladesh desde 1970. Ten en cuenta que la función `fertility_over_time` no debe devolver la gráfica en sí. **La expresión que dibuja el gráfico de líneas se proporciona para ti; por favor, no la cambies.**

```python
countries = Table.read_table('countries.csv')
fertility = Table.read_table("fertility.csv")

def fertility_over_time(country_code, start):
    country_fertility = ...
    country_fertility_after_start = ...
    cleaned_table = ...
    ...
    
bangladesh_code = ...

# No borre estas lineas de codigo 
fertility_1970_ = fertility_over_time(bangladesh_code, 1970)
fertility_over_time(bangladesh_code, 1970).plot(0, 1)
```

```python
grader.check("q5")
```

**Pregunta 6.** Asumiendo que todos los demás factores permanecen constantes, ¿las tendencias en fertilidad en el gráfico anterior ayudan a explicar directamente por qué la tasa de crecimiento de la población disminuyó de 1980 a 2020 en Bangladesh? ¿Por qué sí o por qué no?

```python
explicacion="Escribe tu respuesta aqui"
```

Se ha [observado](https://www.ageing.ox.ac.uk/download/143) que las tasas de fertilidad más bajas a menudo están asociadas con tasas de mortalidad infantil más bajas. Esta conexión se ha atribuido a la planificación familiar: si los padres pueden esperar que todos sus hijos sobrevivan hasta la adultez, entonces elegirán tener menos hijos. En la dirección inversa, tener menos hijos, puede permitir a las familias dedicar más recursos a cada niño, reduciendo así la mortalidad infantil. Podemos ver si esta asociación es evidente en Bangladesh al graficar la relación entre la tasa de fertilidad total y la [tasa de mortalidad infantil por cada 1000 niños](https://es.wikipedia.org/wiki/Mortalidad_infantil).

**Pregunta 7.** Usando tanto las tablas `fertility` como `child_mortality`, dibuja un diagrama de dispersión que tenga la fertilidad total de Bangladesh en el eje horizontal y su mortalidad infantil en el eje vertical, con un punto para cada año, comenzando en 1970.**Se proporciona el código que dibuja el diagrama de dispersión; por favor, no lo cambies.** En su lugar, crea una tabla llamada `post_1969_fertility_and_child_mortality` con las etiquetas de columna y los datos apropiados para generar el gráfico correctamente. Utiliza la etiqueta `Ninos_por_mujer` para describir la fertilidad total y la etiqueta `Muertes_infantiles_por_1000_nacidos` para describir la mortalidad infantil.

*Consejo*: No elimines la columna `time` o recibirás un error en el diagrama de dispersión en la siguiente celda.

```python
fertility=Table().read_table("fertility.csv")

bgd_fertility = ...
bgd_child_mortality = ...
fertility_and_child_mortality = ...
post_1969_fertility_and_child_mortality = ...

# No cambiar las siguiente lineas de codigo por favor
x_births = post_1969_fertility_and_child_mortality.column("Children per woman")
y_deaths = post_1969_fertility_and_child_mortality.column("Child deaths per 1000 born")
time_colors = post_1969_fertility_and_child_mortality.column("time")

plots.figure(figsize=(6,6))
plots.scatter(x_births, y_deaths, c=time_colors, cmap="Blues_r")
plots.colorbar()                  
plots.xlabel("Children per woman")
plots.ylabel("Child deaths per 1000 born");
```

```python
grader.check("q7")
```

El gráfico anterior utiliza **color** para codificar datos sobre la columna `time` de la tabla `post_1969_fertility_and_child_mortality`. Los colores, que van desde azul oscuro hasta blanco, representan el paso del tiempo desde los años 1970 hasta los años 2020. Por ejemplo, un punto en el diagrama de dispersión que representa datos de los años 1970 aparecería como **azul oscuro**, y un punto de los años 2010 aparecería como **azul claro**.

<!-- BEGIN QUESTION -->

**Pregunta 8.** En una o dos oraciones, describe la asociación (si la hay) que ilustra este diagrama de dispersión. ¿Muestra el diagrama que la reducción de la mortalidad infantil *causa* que los padres elijan tener menos hijos?

```python
descripcion="Escribe tu respuesta aqui"
```

<!-- END QUESTION -->

### Punto de control 


<div style="display: flex; justify-content: space-between;">
  <div style="flex: 1; text-align: center;">
    <img src="lukas.jpeg" alt="Foto en primer plano de un golden retriever mirando a la cámara" width="300"/>
  </div>
  
</div>


<p style="text-align: center;"><strong>¡¡¡WOOOHOO!!!</strong></p>


¡Otto y Cloe Mafalda. Quieren felicitarte por alcanzar el punto de control! 

**Ejecuta las siguientes celdas y envía la tarea a la asignación en el correspondiente al punto de control: <u>Punto de control del Proyecto 1</u>****<span style="color: #BC412B">

Para verificar tu trabajo, la celda a continuación volverá a ejecutar todas las pruebas del autograder para la Sección 1.

```python
checkpoint_tests = ["q1", "q2", "q3","q5","q7"]

for test in checkpoint_tests:
    display(grader.check(test))
```

## Envío 

Asegúrate de haber ejecutado todas las celdas en tu cuaderno en orden antes de ejecutar la celda de abajo, de modo que todas las imágenes/gráficas aparezcan en el resultado. La celda de abajo generará un archivo zip para que lo envíes. 

**¡Por favor guarda antes de exportar! **


---

¡Terminaste La primera parte del proyecto!

**Información importante sobre el envío:**
- **Ejecuta todas las pruebas** y verifica que todas pasan
- **Guardar** desde el menú **Archivo**
- **Ejecute la celda final **
- Luego, ve a este [enlace](https://cienciadatosudea.github.io/001_SitioWebFundCienciaDatos/04entregas.html) y envía el archivo ipynb . El nombre de esta parte es "Proyecto01_parte01 Autograder",

**Es su responsabilidad asegurarse de que su trabajo esté guardado antes de ejecutar la última celda.**

`Referencia`: Este material esta basado en [Data 8: The Foundations of Data Science](https://www.data8.org/) - University of California, Berkeley.


# Segunda parte del proyecto. 


### El Mundo


Los cambios demográficos observados en Polonia, como en muchos otros países, reflejan una tendencia común que se puede analizar con datos globales. A medida que los países experimentan mejoras en sus servicios de salud, la esperanza de vida tiende a aumentar y la mortalidad infantil disminuye. Sin embargo, estas mejoras suelen estar acompañadas por una caída en la tasa de fertilidad, lo que genera una desaceleración en el crecimiento de la población. Este fenómeno,  es clave para entender los retos y oportunidades del desarrollo global. A pesar de que las personas viven más tiempo, el número de nacimientos no suele ser suficiente para sostener el crecimiento poblacional, lo que plantea desafíos económicos y sociales que pueden estudiarse con técnicas de ciencia de datos.

Ejecuta la celda a continuación para generar dos histogramas superpuestos, uno para 1962 y otro para 2010, que muestran las distribuciones de las tasas de fertilidad total para estos dos años entre los 201 países en la tabla `fertility`.

```python
Table().with_columns(
    '1962', fertility.where('time', 1962).column(2),
    '2010', fertility.where('time', 2010).column(2)
).hist(bins=np.arange(0, 10.5, 0.5), unit='child per woman')
_ = plots.xlabel('Children per woman')
_ = plots.ylabel('Percent per children per woman')
_ = plots.xticks(np.arange(10))
```

```python
q=Table().with_columns(
    '1962', fertility.where('time', 1962).column(2),
    '2010', fertility.where('time', 2010).column(2)
).bin(bins=np.arange(0, 10.5, 0.5))
q.show()
```

```python
q.where("bin", are.between(0,3))
```

```python
a=3
b=10
q.where("bin", are.between(a,b))
```

```python
q.where("bin", are.between(a,b)).to_df().sum(axis=0)
```

```python
q.to_df().sum(axis=0)
```

```python
82/202
```


- Eje X : El eje X está etiquetado como "Children per woman", es decir, el número de hijos por mujer. Indica el número promedio de hijos que una mujer tendría en su vida en cada país, que esta relacionado con la tasa de fertilidad, que se refiere generalmente a la cantidad promedio de hijos por mujer en un país.


- Eje Y : El eje Y está etiquetado como "Percent per children per woman", Indica el porcentaje de países que tienen ese número específico de hijos por mujer, para los años 1962 y 2010.


**Pregunta 9.** Observa los dos histogramas que comparan las tasas de fertilidad de los países en 1962 y 2010. A continuación, se proporcionan varias afirmaciones que podrían inferirse de estos gráficos.
Tu tarea es asignar a `fertility_statements` un arreglo que contenga valores booleanos, asociados a cada  afirmacion,  que se pueden inferir correctamente a partir de los histogramas.

Ejemplo: Si las afirmaciones correctas son la 0, la 3 y la 5,  y las falsas son la 1, 2, 4,  entonces el arreglo se puede definir así:

`fertility_statements = make_array(False, False, True, True, True, True)`

Estas son las afirmaciones:

0. Aproximadamente el mismo número de países tuvo una tasa de fertilidad entre 3.5 y 4.5, tanto en 1962 como en 2010.

1. En 1962, menos del 20% de los países tenían una tasa de fertilidad inferior a 3.

2.  Al menos la mitad de los países tuvieron una tasa de fertilidad entre 5 y 8 en 1962.

3. En 2010, alrededor del 40% de los países tenían una tasa de fertilidad entre 1.5 y 2.

4.  Al menos la mitad de los países tuvieron una tasa de fertilidad inferior a 3 en 2010.

5. Más países tenían una tasa de fertilidad superior a 3 en 1962 que en 2010.

```python
fertility_statements = make_array(....)
```

```python
grader.check("q9")
```

**Pregunta 10.**  La **poblacion mundial** es la suma de las poblaciones de todos los países, cálcule esa suma en una tabla llamada `world_population` y después realice un gráfico de líneas de la poblacion mundial desde 1800 hasta 2020 (inclusive de ambos extremos). Debes usar la tabla `population` definida anteriormente en el proyecto, el nombre de las columnas de `world_population` deberá ser: 
`time`, `population_total sum`



```python

population = Table.read_table('population.csv')
world_population = ...
...
...
```

```python
grader.check("q10")
```

**Pregunta 11.** Crea una función `stats_for_year` que tome un `year` (año) como argumento y devuelva una tabla de estadísticas:

- La tabla que devuelva debe tener cuatro columnas: `geo` (geografía), `population_total` (población total), `children_per_woman_total_fertility` (número de hijos por mujer - tasa de fecundidad total), y `child_mortality_under_5_per_1000_born` (mortalidad infantil bajo los 5 años por cada 1000 nacidos).

- Cada fila debe contener un código de país único en formato Alpha-3 y tres estadísticas: población, tasa de fecundidad y mortalidad infantil para ese `year` (año) de las tablas `population` (población), `fertility` (fecundidad) y `child_mortality` (mortalidad infantil). Solo incluye filas para las cuales las tres estadísticas estén disponibles para el país y año. Además, restringe el resultado a los códigos de país que aparecen en `big_50`, un arreglo de los 50 países más poblados en 2020. Esta restricción acelerará los cálculos más adelante en el proyecto.

- Después de escribir `stats_for_year`, intenta llamar a `stats_for_year` en cualquier año entre 1960 y 2020. Trata de entender el resultado de `stats_for_year`.


*Pista 1*: Las pruebas para esta pregunta son bastante completas, así que si pasas las pruebas, tu función probablemente sea correcta. Sin embargo, sin llamar a tu función tú mismo y observar el resultado, será muy difícil entender cualquier problema que tengas, así que haz tu mejor esfuerzo para escribir la función correctamente y verifica que funcione antes de confiar en las pruebas del `grader` para confirmar tu trabajo.

*Pista 2*: ¿Qué tienen en común las tres tablas (presta atención a los nombres de las columnas)?

*Pista 3*: Crea celdas adicionales antes de escribir directamente la función.

```python


#Primero creamos una tabla de población que solo incluye los 
#50 países con las mayores poblaciones en 2020. Nos enfocamos 
#solo en estos 50 países para que los gráficos más adelante se ejecuten más rápido.


big_50 = population.where('time', are.equal_to(2020)).sort("population_total", descending=True).take(np.arange(50)).column('geo')
population_of_big_50 = population.where('time', are.above(1959)).where('geo', are.contained_in(big_50))

def stats_for_year(year):
    p = population_of_big_50.where('time', are.equal_to(year)).drop('time')
    f = fertility.where('time', are.equal_to(year)).drop('time')
    c = child_mortality.where('time', are.equal_to(year)).drop('time')
    ...
    ...

```

```python
grader.check("q11")
```

**Pregunta 12.** Crea una tabla llamada `pop_por_decada` con dos columnas: `decada` y `poblacion`, en ese orden. 

- La tabla debe tener una fila para cada año que inicia una década, en orden ascendente desde 1960 hasta 2020. Por ejemplo, 1960 marca el comienzo de la década de los 60.

- La columna `poblacion` debe contener la población total de todos los países incluidos en el resultado de `stats_for_year(year)` para el primer año de cada década. Deberías notar que estos países representan la mayor parte de la población mundial.

**Consejo:** Un enfoque útil sería definir una función llamada `poblacion_para_ano` que calcule la población total, y luego aplicar esa función a la columna `decada`. **Piensa en cómo podrías usar la función `stats_for_year`** de la pregunta anterior si decides implementar `poblacion_para_ano`. Esta primera prueba es simplemente una verificación de coherencia para tu función auxiliar si decides utilizarla. **No perderás puntos por no implementar la función `poblacion_para_ano`.**

**Nota:** La celda en la que generarás la tabla `pop_por_decada` está justo debajo de la celda donde puedes optar por definir la función auxiliar `poblacion_para_ano`. Debes definir tu tabla `pop_por_decada` en la celda que comienza con la tabla `decadas` ya definida.


```python
def pop_for_year(year):
    '''Return the total population for the specified year.'''
    ...
    
# Ahora que has definido tu función auxiliar (si has decidido hacerlo), define la tabla `pop_by_decade`.
pop_por_decada = ...

pop_por_decada.set_format(1, NumberFormatter)
```

```python
pop_por_decada
```

```python
grader.check("q12")
```

Ahora que has definido tu función auxiliar (si has decidido hacerlo), define la tabla `pop_by_decade`.

La tabla `countries` describe diversas características de los países. La columna `country` contiene los mismos códigos que la columna `geo` en cada una de las otras tablas de datos (`population`, `fertility` y `child_mortality`). La columna `world_6region` clasifica a cada país en una región del mundo. Ejecuta la celda de abajo para inspeccionar los datos.

```python
countries = Table.read_table('countries.csv').where('country', are.contained_in(population.group('geo').column('geo')))
countries.select('country', 'name', 'world_6region')
```

**Pregunta 13.** Crea una tabla llamada `region_counts`. Debe contener dos columnas llamadas `region` y `count`. La columna `region` debe contener las regiones del mundo, y la columna `count` debe contener el número de países en cada región que aparece en el resultado de `stats_for_year(2020)`.Por ejemplo, una fila tendría `south_asia` como su valor de `region` y un número entero como su valor de `count`: la cantidad de países grandes de Asia del Sur para los cuales tenemos cifras de población, fertilidad y mortalidad infantil del año 2020.


*Sugerencia*: Es posible que tengas que renombrar una columna para llamarla `region`.

```python
year_2020_data = stats_for_year(2020)
region_counts = ...
region_counts
```

```python
grader.check("q13")
```

El siguiente diagrama de dispersión compara la tasa de fecundidad total y la tasa de mortalidad infantil para cada país en 1960. El área de cada punto representa la población del país, y el color representa su región del mundo. Ejecuta la celda. ¿Crees que puedes identificar alguno de los puntos?

```python
from functools import lru_cache as cache

# This cache annotation makes sure that if the same year
# is passed as an argument twice, the work of computing
# the result is only carried out once.
@cache(None)
def stats_relabeled(year):
    """Relabeled and cached version of stats_for_year."""
    return stats_for_year(year).relabel(2, 'Children per woman').relabel(3, 'Child deaths per 1000 born')

def fertility_vs_child_mortality(year):
    """Draw a color scatter diagram comparing child mortality and fertility."""
    with_region = stats_relabeled(year).join('geo', countries.select('country', 'world_6region'), 'country')
    with_region.scatter(2, 3, sizes=1, group=4, s=500)
    plots.xlim(0,10)
    plots.ylim(-50, 500)
    plots.title(year)
    plots.show()

fertility_vs_child_mortality(1960)
```

```python
with_region = stats_relabeled(1960).join('geo', countries.select('country', 'world_6region'), 'country')

with_region.sort("population_total", descending=True)
```

```python
with_region.sort("population_total", descending=True).sort("Child deaths per 1000 born", descending=True)
```

**Pregunta 14.** Asigna `scatter_statements` a un arreglo con los valores booleanos,  de cada una de las afirmaciones que se pueden inferir de este diagrama de dispersión para 1960, asigna True o False al arreglo, ejemplo :

make_array(True, True, True, False, False)` :

1. Todos los países en `europe_central_asia` tenían tasas de fertilidad uniformemente bajas.
2. La tasa de mortalidad infantil más baja de cualquier país fue en un país de `east_asia_pacific`.
3. La mayoría de los países tiene una tasa de fertilidad superior a 5.
4. Hay una asociación entre la mortalidad infantil y la fertilidad.
5. Los dos países más grandes por población también tuvieron las dos tasas de mortalidad infantil más altas.



```python
scatter_statements = make_array(...)
```

```python
grader.check("q14")
```

El resultado de la celda a continuación es interactivo. Arrastra el deslizador hacia la derecha para ver cómo han cambiado los países a lo largo del tiempo. Descubrirás que, en términos de crecimiento poblacional, la división entre los países del Norte global y el Sur global que existía en los años 1960 se ha reducido significativamente.Este cambio en las tasas de fertilidad es la razón por la cual se espera que la población global crezca más lentamente en el siglo XXI de lo que lo hizo en los siglos XIX y XX. Las tasas de fertilidad cambian por razones que incluyen patrones culturales, mejores perspectivas de supervivencia infantil hasta la adultez y la planificación familiar (como el uso de anticonceptivos y el mayor control de las mujeres sobre su reproducción).**Nota:** No te preocupes si aparece una advertencia roja al ejecutar la celda a continuación. ¡Aún podrás ejecutarla!

```python
_ = widgets.interact(fertility_vs_child_mortality, 
                     year=widgets.IntSlider(min=1960, max=2020, value=1960))
```

Ahora es un buen momento para tomar un descanso y ver los mismos datos presentados por [Hans Rosling en una charla TEDx de 2010](https://www.ted.com/talks/hans_rosling_the_good_news_of_the_decade_we_re_winning_the_war_against_child_mortality?subtitle=en&geo=es) con una animación y algunos comentarios importantes.

Cuando analizamos la población y la fertilidad como científicos de datos, necesitamos aprender sobre las experiencias de las personas en la vida real, no solo de forma abstracta como datos. También debemos reconocer que los estudios de población a veces han tenido trasfondos políticos. Estos trasfondos han incluido el control de la población, el control de la reproducción de las mujeres o temores de cambios entre grupos raciales. Para mejorar como científicos de datos, debemos revisar nuestras suposiciones para evitar reproducir, sin pensar, patrones del pasado.

## 2. Pobreza Global

En 1800, el 85% de los 1,000 millones de personas en el mundo vivían en [pobreza extrema](https://en.wikipedia.org/wiki/Extreme_poverty), definida por las Naciones Unidas como "una condición caracterizada por una grave privación de las necesidades humanas básicas, incluyendo alimentos, agua potable segura, instalaciones sanitarias, salud, refugio, educación e información".  En el momento en que se recopilaron los datos para este proyecto, una definición común de pobreza extrema era una persona viviendo con menos de \$1.25 (dolares) al día.  

En 2018, se estimó que la proporción de personas viviendo en pobreza extrema era de [aproximadamente el 9%](https://unstats.un.org/sdgs/report/2019/goal-01/). Aunque la tasa mundial de pobreza extrema ha disminuido consistentemente durante cientos de años, el número de personas que viven en pobreza extrema todavía supera los 600 millones. Las Naciones Unidas adoptaron un [objetivo ambicioso](http://www.un.org/sustainabledevelopment/poverty/): "Para 2030, erradicar la pobreza extrema para todas las personas en todas partes".  

En esta parte del proyecto, examinaremos algunos aspectos de la pobreza global que podrían afectar si el objetivo es alcanzable. Las causas de la pobreza son complejas. Incluyen historias globales, como el colonialismo, así como factores tales como: la atención sanitaria, la economía y la desigualdad social en cada país.

Primero, carga la población y la tasa de pobreza por país y año, así como las descripciones de los países. Mientras que la tabla de `población` tiene valores para cada año reciente para muchos países, la tabla de `pobreza` solo incluye ciertos años para cada país en los cuales se disponía de una medición de la tasa de pobreza extrema.

```python
population = Table.read_table('population.csv')
countries = Table.read_table('countries.csv').where('country', are.contained_in(population.group('geo').column('geo')))
poverty = Table.read_table('poverty.csv')
poverty.show(3)
```

**Pregunta 1.** Asigna el nombre de  `latest_poverty` a una tabla de tres columnas, con una fila para cada país, que aparece en la tabla `poverty`, con las siguientes caracteristicas:

- La primera columna debe contener el código de 3 letras del país, el nombre de esta columna deberá ser  `geo`

- La segunda columna debe contener el año más reciente para el cual está disponible una tasa de pobreza extrema para el país, el nombre de esta columna deberá ser `time`,

- La tercera columna debe contener la tasa de pobreza en ese año. El nombre de esta columna deberá ser `extreme_poverty_percent_people_below_125_a_day`
<!-- *Pista*: piensa en cómo funciona `group`: realiza una búsqueda secuencial de la tabla (de arriba hacia abajo) y recopila valores en el arreglo en el orden en que aparecen, y luego aplica una función a ese arreglo. La función `first` puede ser útil, pero no estás obligado a usarla. -->

```python

latest_poverty = ...
latest_poverty
```

```python
grader.check("q15")
```

**Pregunta 2.** Usando tanto `latest_poverty` como `population`, crea una tabla de cuatro columnas llamada `recent_poverty_total` con una fila para cada país en `latest_poverty`. Las cuatro columnas deben tener las siguientes etiquetas y contenidos:

1. `geo` contiene el código de país de 3 letras.
  
2. `poverty_percent` contiene el porcentaje de pobreza más reciente.

3. `population_total` contiene la población del país en 2010.

4. `poverty_total` contiene el número de personas en situación de pobreza **redondeado al número entero más cercano**, basado en la población de 2010 y el porcentaje de pobreza más reciente


```python
...
...
recent_poverty_total = ...
recent_poverty_total

```

```python
grader.check("q16")
```

**Pregunta 3.** Asigna el nombre `poverty_percent` al porcentaje conocido de la población mundial en 2010 que vivía en extrema pobreza. Supón que los números de `poverty_total` en la tabla `recent_poverty_total` describen a **todas** las personas en 2010 que vivían en extrema pobreza. Deberías obtener un número que esté por encima de la estimación global de 2018 del 9%, ya que muchas de las tasas de pobreza específicas de los países son anteriores a 2018.

*Pista*: La suma de la columna `population_total` en la tabla `recent_poverty_total` no es la población mundial, porque solo un subconjunto de los países del mundo están incluidos en la tabla `recent_poverty_total` (solo algunos países tienen tasas de pobreza conocidas). Utiliza la tabla `population` para calcular la población total mundial de 2010.

*Pista*: Estamos calculando un porcentaje (valor entre 0 y 100), no una proporción (valor entre 0 y 1).

```python
...
total_population = ...
poverty_total = ...
poverty_percent = ...

```

```python
grader.check("q17")
```

La tabla `countries` incluye no solo el nombre y la región de los países, sino también sus posiciones en el globo terráqueo.

```python
countries.select('country', 'name', 'world_4region', 'latitude', 'longitude')
```

**Pregunta 4.** Usando tanto `countries` como `recent_poverty_total`, crea una tabla llamada `poverty_map` con cinco columnas y una fila para cada país en . Las cinco columnas deben tener las siguientes etiquetas y contenidos, en este orden:

1. `latitude` contiene la latitud del país,
2. `longitude` contiene la longitud del país.
3. `name` contiene el nombre del país no el código de tres letras.
4. `region` contiene la región del país desde la columna `world_4region` de `countries`
5. `poverty_total` contiene el total de pobreza del país.

```python
...
...
poverty_map=...
```

```python
grader.check("q18")
```

Ejecuta la celda a continuación para dibujar un mapa del mundo en el que las áreas de los círculos representan el número de personas que viven en la pobreza extrema. Haz doble clic en el mapa para acercarte.> **Nota:** Si la celda a continuación no se está cargando, puedes ver el resultado [aquí](https://github.jonathanferrari.com/static/poverty_map.html)

```python
# It may take a few seconds to generate this map.
colors = {'africa': 'blue', 'europe': 'black', 'asia': 'red', 'americas': 'green'}
scaled = poverty_map.with_columns(
    'labels', poverty_map.column('name'),
    'colors', poverty_map.apply(colors.get, 'region'),
    'areas', 1e-4 * poverty_map.column('poverty_total')
).drop('name', 'region', 'poverty_total')

Circle.map_table(scaled)
```

Aunque en 2010 las personas vivían en la pobreza extrema en todo el mundo (con más de 5 millones en los Estados Unidos), la mayor cantidad se encontraba en Asia y África.

**Pregunta 5.** Asigna `largest` a una tabla de dos columnas con el `nombre` (no el código de 3 letras) y `pobreza_total` de los 10 países con el mayor número de personas viviendo en extrema pobreza.

*Pista*: ¿Cómo podemos usar `take` y `np.arange` en conjunto?

```python
largest=...
largest.set_format('poverty_total', NumberFormatter)
```

```python
grader.check("q19")
```

**Pregunta 6.** Es crucial analizar el número absoluto de personas que viven en pobreza, además del porcentaje. Este número absoluto es esencial para determinar cuántos recursos se necesitan para apoyar a las personas en situación de pobreza.  Diversas entidades, como gobiernos, organizaciones internacionales y ONGs, desempeñan un papel crucial en determinar y asignar estos recursos según las necesidades identificadas. En las siguientes preguntas, profundizaremos en este aspecto. 


Ahora vamos a crear la función `poverty_timeline`. Esta función debe recibir **el nombre de un país** como argumento (no el código Alpha-3). Su propósito es trazar un gráfico de líneas que muestre el número de personas viviendo en pobreza en ese país, con el tiempo en el eje horizontal. El gráfico debe incluir un punto por cada registro(filas) en la tabla `poverty` correspondiente al país. Para calcular la población en situación de pobreza a partir del porcentaje, multiplícalo por la población del país **en ese año**. 

**Nota:** Utiliza `plots.show()` antes de del return de la función

*Pista 1:* Esta pregunta es extensa. Siéntete libre de experimentar creando celdas adicionales. Puedes hacerlo desde la barra de herramientas presionando el botón `+`.

*Pista 2:* Considera el uso de `join` en tu código. Si lo deseas, utiliza la celda de markdown a continuación para planificar tu respuesta, aunque no es obligatorio completarla.

```python

def poverty_timeline(country_name):
    ...
    ...
    ...
    ...
    poverty_population.plot("time", "people_in_poverty")
    plots.title(country_name)
    plots.ylim(bottom=0)
    plots.show() # 
    
poverty_timeline('Poland')
poverty_timeline('India')
poverty_timeline('Nigeria')
poverty_timeline('China')
poverty_timeline('Colombia')
```

```python
grader.check("q20")
```

Aunque el número de personas que viven en extrema pobreza ha aumentado en algunos países, las disminuciones en otros países, especialmente las disminuciones masivas en China e India, han configurado la tendencia general de que la pobreza extrema está disminuyendo en todo el mundo, tanto en porcentaje como en número absoluto. Para aprender más, mira [Hans Rosling en una película de 2015](https://www.gapminder.org/videos/dont-panic-end-poverty/) sobre el objetivo de la ONU de erradicar la pobreza extrema del mundo. A continuación, hemos añadido un menú desplegable interactivo para que puedas visualizar gráficos de la `línea de tiempo de pobreza` para otros países. Ten en cuenta que cada selección del menú desplegable puede tardar unos segundos en ejecutarse.

```python
# Just run this cell
all_countries = poverty_map.column('name')
_ = widgets.interact(poverty_timeline, country_name=list(all_countries))
```

**Anahi quiere decirte, ¡has terminado!** Felicitaciones por descubrir muchos hechos importantes sobre la pobreza global y demostrar tu dominio en la manipulación de tablas y visualización de datos. Es momento de enviar tu trabajo.**<span style="color: #BC412B"></span>** <u></u>


<img src="anahi.jpeg" alt="Foto en primer plano de un golden retriever mirando a la cámara" width="300"/>


## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

These are some submission instructions.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1 (points: 1)
```
>>> import pandas as pd
>>> assert b_pop.to_df().time.mean() == 1995
>>> assert np.isclose(b_pop.to_df().population_total.mean(), 114373388.50980392, 0.001)
>>> assert 'time' in b_pop.labels, "La columna 'time' no está presente en b_pop."
>>> assert 'population_total' in b_pop.labels, "La columna 'population_total' no está presente en b_pop."
```


Tests for q10 (points: 1)
```
>>> assert 'time' in world_population.labels, "La columna 'time' no está presente en world_population."
>>> assert 'population_total sum' in world_population.labels, "La columna 'population_total sum' no está presente en world_population."
>>> assert np.isclose(world_population.column('time').mean(), 1910, 0.001)
>>> assert np.isclose(world_population.column('population_total sum').mean(), 2626917081.678733, 0.001)
```


Tests for q11 (points: 1)
```
>>> assert np.isclose(np.mean(stats_for_year(2018).column('population_total')), 132609584.28, 0.01)
>>> assert np.isclose(np.mean(stats_for_year(2010).column('population_total')), 121293482.64, 0.01)
>>> assert np.isclose(np.mean(stats_for_year(2000).column('population_total')), 107541243.38, 0.01)
>>> assert np.isclose(np.mean(stats_for_year(2010).column('children_per_woman_total_fertility')), 2.9691999999999994, 0.01)
>>> assert np.isclose(np.mean(stats_for_year(2018).column('children_per_woman_total_fertility')), 2.7053999999999996, 0.01)
>>> assert np.isclose(np.mean(stats_for_year(2000).column('children_per_woman_total_fertility')), 3.33559999999, 0.01)
>>> a = []
>>> a.append(stats_for_year(2000).to_df().describe()['population_total'].iloc[0])
>>> a.append(stats_for_year(2000).to_df().describe()['population_total'].iloc[1])
>>> a.append(stats_for_year(2000).to_df().describe()['population_total'].iloc[2])
>>> a.append(stats_for_year(2000).to_df().describe()['population_total'].iloc[3])
>>> assert a[0] == 50
>>> assert a[1] == 107541243.38
>>> assert a[2] == 227538018.45596266
>>> assert a[3] == 16395477.0
```


Tests for q12 (points: 1)
```
>>> assert pop_por_decada.to_df().shape == (7, 2), 'La tabla resultante no tiene las dimensiones correctas'
>>> assert np.isclose(pop_por_decada.to_df().loc[0, 'poblacion'], 2635123897, 0.01)
>>> assert np.isclose(pop_por_decada.to_df().loc[6, 'poblacion'], 6765161289, 0.01)
```


Tests for q13 (points: 1)
```
>>> df = region_counts.to_df()
>>> assert np.isclose(df['count'].sum(), 50, 1e-06)
```


Tests for q14 (points: 1)
```
>>> u = make_array(False, False, True, True, False)
>>> assert scatter_statements[0] == u[0]
>>> assert scatter_statements[1] == u[1]
>>> assert scatter_statements[2] == u[2]
>>> assert scatter_statements[3] == u[3]
>>> assert scatter_statements[4] == u[4]
```


Tests for q15 (points: 1)
```
>>> assert 'geo' in latest_poverty.labels, "La columna 'geo' no está presente en latest_poverty."
>>> assert 'time' in latest_poverty.labels, "La columna 'time' no está presente en latest_poverty."
>>> assert 'poverty_percent' in latest_poverty.labels, "La columna 'poverty_percent' no está presente en latest_poverty."
>>> assert len(latest_poverty.to_df().geo.unique()) == 145
>>> assert np.isclose(latest_poverty.to_df()['time'].mean(), 2008.53, 0.001)
>>> assert np.isclose(latest_poverty.to_df()['poverty_percent'].mean(), 17.13, 0.001)
```


Tests for q16 (points: 1)
```
>>> assert 'geo' in recent_poverty_total.labels, "La columna 'geo' no está presente en latest_poverty."
>>> assert 'poverty_percent' in recent_poverty_total.labels, "La columna 'poverty_percent' no está presente en latest_poverty."
>>> assert 'population_total' in recent_poverty_total.labels, "La columna 'population_total' no está presente en latest_poverty."
>>> assert 'poverty_total' in recent_poverty_total.labels, "La columna 'population_total' no está presente en poverty_total."
>>> assert np.isclose(recent_poverty_total.to_df().describe().loc['mean'].values[0], 17.1326897, 0.001)
>>> assert np.isclose(recent_poverty_total.to_df().describe().loc['mean'].values[1], 45521178.3, 0.001)
>>> assert np.isclose(recent_poverty_total.to_df().describe().loc['mean'].values[2], 682793067.0 / 100, 0.001)
```


Tests for q17 (points: 1)
```
>>> assert np.isclose(poverty_percent, 14.24, 0.001)
```


Tests for q18 (points: 1)
```
>>> assert 'latitude' in poverty_map.labels, 'La columna  no está presente en poverty_map.'
>>> assert 'longitude' in poverty_map.labels, 'La columna  no está presente en poverty_map.'
>>> assert 'name' in poverty_map.labels, 'La columna  no está presente en poverty_map.'
>>> assert 'region' in poverty_map.labels, 'La columna  no está presente en poverty_map.'
>>> assert 'poverty_total' in recent_poverty_total.labels, "La columna 'population_total' no está presente en poverty_total."
>>> assert np.isclose(poverty_map.to_df()['latitude'].mean(), 19.39, 0.001)
>>> assert np.isclose(poverty_map.to_df()['longitude'].mean(), 16.44, 0.001)
>>> assert np.isclose(poverty_map.to_df()['poverty_total'].mean(), 6827930.717241379, 0.001)
```


Tests for q19 (points: 1)
```
>>> assert 'name' in largest.labels, 'La columna  no está presente en largest.'
>>> assert 'poverty_total' in largest.labels, 'La columna  no está presente en largest.'
>>> assert np.isclose(largest.to_df()['poverty_total'].mean(), 72823360.9, 0.001)
```


Tests for q2 (points: 1)
```
>>> import pandas as pd
>>> assert np.isclose(b_five_growth.to_df().time.mean(), 1992.5, 0.001)
>>> assert np.isclose(b_five_growth.to_df().population_total.mean(), 109356945.1, 0.001)
>>> assert np.isclose(b_five_growth.to_df().annual_growth[0], 0.01753862351142521, 0.001)
>>> assert 'time' in b_five_growth.labels, "La columna 'time' no está presente en b_pop."
>>> assert 'population_total' in b_five_growth.labels, "La columna 'population_total' no está presente en b_pop."
```


Tests for q20 (points: 1)
```
>>> assert callable(poverty_timeline), 'La función poverty_timeline no existe.'
```


Tests for q3 (points: 1)
```
>>> assert np.isclose(bangladesh_life_expectancy.to_df().time.mean(), 2035.0, 1e-06)
>>> assert np.isclose(bangladesh_life_expectancy.to_df().life_expectancy_years.mean(), 74.0656822, 1e-06)
>>> assert 'geo' in bangladesh_life_expectancy.labels, "La columna 'geo' no esta presente"
>>> assert 'time' in bangladesh_life_expectancy.labels, 'La columna time no esta presente'
```


Tests for q5 (points: 1)
```
>>> assert 'Anio' in fertility_1970_.labels, "La columna 'Anio' no esta presente"
>>> assert 'Ninos_por_mujer' in fertility_1970_.labels, 'La columna Ninos_por_mujer no esta presente'
>>> assert bangladesh_code == 'bgd'
>>> assert np.isclose(fertility_1970_.to_df().Anio.mean(), 2035.0, 0.1)
>>> assert np.isclose(fertility_1970_.to_df().Ninos_por_mujer.mean(), 2.6768708, 0.001)
```


Tests for q7 (points: 1)
```
>>> assert 'Ninos_por_mujer' in post_1969_fertility_and_child_mortality.labels, "La columna 'Ninos_por_mujer' no esta presente"
>>> assert 'Muertes_infantiles_por_1000_nacidos' in post_1969_fertility_and_child_mortality.labels, "La columna 'Muertes_infantiles_por_1000_nacidos' no esta presente"
>>> assert 'Anios' in post_1969_fertility_and_child_mortality, "La columna 'Anio' no esta presente"
>>> assert np.isclose(time_colors.mean(), 1995.0, 1e-06)
>>> assert np.isclose(y_deaths.mean(), 121.15352941176468, 0.01)
>>> assert np.isclose(x_births.mean(), 4.166274509803921, 1e-06)
```


Tests for q9 (points: 1)
```
>>> u = make_array(False, False, True, True, True, True)
>>> assert fertility_statements.sum() == u.sum()
>>> assert fertility_statements[0] == u[0]
>>> assert fertility_statements[1] == u[1]
>>> assert fertility_statements[2] == u[2]
>>> assert fertility_statements[3] == u[3]
>>> assert fertility_statements[4] == u[4]
>>> assert fertility_statements[5] == u[5]
```
